=========================================
Ejercicio de Conteo de Palabras en Python
=========================================


.. tabbed:: quiz3

    .. tab:: Ejercicio 1

        .. activecode:: q3_1_eewcn
            :nocodelens:

            Para la bandera --count, implementa una función print_words(filename) que cuente cuántas veces aparece cada palabra en el texto e imprima: palabra1 contador1, palabra2 contador2, ... Imprime la lista anterior en orden ordenado por palabra (Python ordenará la puntuación para que aparezca antes que las letras, eso está bien). Almacena todas las palabras en minúsculas, así que 'The' y 'the' cuentan como la misma palabra.

            ~~~~
            def print_words(filename):
                # +++tu código aquí+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Agrega tus pruebas aquí
                pass

            myTests().main()

    .. tab:: Ejercicio 2

        .. activecode:: q3_2_eewcn
            :nocodelens:

            Para la bandera --topcount, implementa una función print_top(filename) que es similar a print_words() pero que imprime solo las 20 palabras más comunes ordenadas de manera que la palabra más común sea la primera, luego la siguiente palabra más común, y así sucesivamente.

            ~~~~
            def print_top(filename):
                # +++tu código aquí+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Agrega tus pruebas aquí
                pass

            myTests().main()
