======================
Python Ejercicio Mimic
======================


.. tabbed:: quiz2

    .. tab:: Ejercicio 1

        .. activecode:: q4_1_eemcn
            :nocodelens:

            Construye un diccionario "mimic" que mapee cada palabra que aparece en el archivo a una lista de todas las palabras que inmediatamente siguen a esa palabra en el archivo. La lista de palabras puede estar en cualquier orden y debe incluir duplicados. 

            ~~~~
            def mimic_dict(filename):
                # +++tu código aquí+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Agrega tus pruebas aquí
                pass

            myTests().main()

    .. tab:: Ejercicio 2

        .. activecode:: q4_2_eemcn
            :nocodelens:

            Dado un diccionario mimic y una palabra de inicio, imprime 200 palabras aleatorias.

            ~~~~
            def print_mimic(mimic_dict, word):
                # +++tu código aquí+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Agrega tus pruebas aquí
                pass

            myTests().main()
