=====================
Python Mimic Exercise
=====================

.. tabbed:: quiz2

    .. tab:: Exercise 1

        .. activecode:: q4_1_reto11_
            :nocodelens:

            Build a "mimic" dict that maps each word that appears in the file to a list of all the words that immediately follow that word in the file. The list of words can be in any order and should include duplicates.

            ~~~~
            def mimic_dict(filename):
                # +++your code here+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Add your tests here
                pass

            myTests().main()

    .. tab:: Exercise 2

        .. activecode:: q4_2_mcn
            :nocodelens:

            Given mimic dict and start word, prints 200 random words.

            ~~~~
            def print_mimic(mimic_dict, word):
                # +++your code here+++
                return
            

            ====
            from unittest.gui import TestCaseGui

            class myTests(TestCaseGui):
                # Add your tests here
                pass

            myTests().main()