<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Strings &#8212; PyZombis master@b54b58e</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=ef1af912ec3415fea854"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="For, Funciones y Azar" href="TWP20.html" />
    <link rel="prev" title="Listas" href="TWP17.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.7.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@b54b58e</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=TWP18'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Strings</a><ul>
<li><a class="reference internal" href="#comillas-de-varios-tipos">Comillas de varios tipos.</a></li>
<li><a class="reference internal" href="#rebanar">Rebanar</a></li>
<li><a class="reference internal" href="#id1">Rebanar</a></li>
<li><a class="reference internal" href="#incremento-en-el-corte">Incremento en el corte</a></li>
<li><a class="reference internal" href="#id2">Incremento en el corte</a></li>
<li><a class="reference internal" href="#string-no-se-puede-modificar">String no se puede modificar</a></li>
<li><a class="reference internal" href="#puedo-crear-nuevos-strings">Puedo crear nuevos strings</a></li>
<li><a class="reference internal" href="#concentracion">Concentración</a></li>
<li><a class="reference internal" href="#verificacion-parcial-de-strings">Verificación parcial de strings</a></li>
<li><a class="reference internal" href="#funciones-find-y-replace">Funciones <code class="docutils literal notranslate"><span class="pre">find</span></code> y <code class="docutils literal notranslate"><span class="pre">replace</span></code></a></li>
<li><a class="reference internal" href="#funciones-split-y-join">Funciones <code class="docutils literal notranslate"><span class="pre">split</span></code> y <code class="docutils literal notranslate"><span class="pre">join</span></code></a></li>
<li><a class="reference internal" href="#ejercicio">Ejercicio</a></li>
<li><a class="reference internal" href="#dojo-de-codificacion">Dojo de codificación</a></li>
<li><a class="reference internal" href="#id3">Dojo de codificación</a></li>
<li><a class="reference internal" href="#desarrollo-basado-en-pruebas">Desarrollo basado en pruebas</a></li>
<li><a class="reference internal" href="#pasos-de-bebe">Pasos de bebé</a></li>
<li><a class="reference internal" href="#programacion-en-pareja">Programación en pareja</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Listas" data-toggle="tooltip">
    <a href="TWP17.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - For, Funciones y Azar' data-toggle="tooltip" >              
    <a href="TWP20.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="strings">
<h1>Strings<a class="headerlink" href="#strings" title="Permalink to this headline">¶</a></h1>
<a class="reference internal image-reference" href="../_images/TWP10_001.jpeg"><img alt="" class="align-center" src="../_images/TWP10_001.jpeg" style="width: 9.258cm; height: 14.925cm;" /></a>
<div class="section" id="comillas-de-varios-tipos">
<h2>Comillas de varios tipos.<a class="headerlink" href="#comillas-de-varios-tipos" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>¿Puedo usar comillas simples, dobles o triples?</p></li>
</ul>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="1">
        <div class="pytutorVisualizer" id="cl_l18_1"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 1 (cl_l18_1)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_1_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_1"] = {"code": "x = 'aguacate'\nprint(x)\ny = \"McDonald's\"\nprint(y)\nformulario = \"\"\"\n<html>\n<head>\n    <title> Prueba </title>\n</head>\n<body>\n    <p>Probando</p>\n</body>\n<html>\"\"\"", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"x": "aguacate"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"x": "aguacate"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "aguacate\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"x": "aguacate", "y": "McDonald's"}, "ordered_globals": ["x", "y"], "stack_to_render": [], "heap": {}, "stdout": "aguacate\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"x": "aguacate", "y": "McDonald's"}, "ordered_globals": ["x", "y"], "stack_to_render": [], "heap": {}, "stdout": "aguacate\nMcDonald's\n"}, {"line": 5, "event": "return", "func_name": "<module>", "globals": {"x": "aguacate", "y": "McDonald's", "formulario": "\n<html>\n<head>\n    <title> Prueba </title>\n</head>\n<body>\n    <p>Probando</p>\n</body>\n<html>"}, "ordered_globals": ["x", "y", "formulario"], "stack_to_render": [], "heap": {}, "stdout": "aguacate\nMcDonald's\n"}]};
</script>
</div>
</div>
<div class="section" id="rebanar">
<h2>Rebanar<a class="headerlink" href="#rebanar" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Rebanada del primer índice al anterior del segundo</p></li>
</ul>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="2">
        <div class="pytutorVisualizer" id="cl_l18_2"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 2 (cl_l18_2)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_2_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_2"] = {"code": "x = \"0123456789\"\nprint(x[0:2])\nprint(x[1:2])\nprint(x[2:4])\nprint(x[0:5])\nprint(x[1:8])", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n1\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n1\n23\n"}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n1\n23\n01234\n"}, {"line": 6, "event": "return", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n1\n23\n01234\n1234567\n"}]};
</script>
</div>
</div>
<div class="section" id="id1">
<h2>Rebanar<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Podemos omitir índices, sustituyendo el extremo correspondiente y
también podemos tener índices negativos: -1 último, -2 penúltimo</p></li>
</ul>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="3">
        <div class="pytutorVisualizer" id="cl_l18_3"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 3 (cl_l18_3)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_3_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_3"] = {"code": "x = \"0123456789\"\nprint(x[:2])\nprint(x[4:])\nprint(x[4:-1])\nprint(x[-4:-1])\nprint(x[:])", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n456789\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n456789\n45678\n"}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n456789\n45678\n678\n"}, {"line": 6, "event": "return", "func_name": "<module>", "globals": {"x": "0123456789"}, "ordered_globals": ["x"], "stack_to_render": [], "heap": {}, "stdout": "01\n456789\n45678\n678\n0123456789\n"}]};
</script>
</div>
</div>
<div class="section" id="incremento-en-el-corte">
<h2>Incremento en el corte<a class="headerlink" href="#incremento-en-el-corte" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Puedo usar un incremento al cortar el string</p></li>
</ul>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="4">
        <div class="pytutorVisualizer" id="cl_l18_4"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 4 (cl_l18_4)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_4_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_4"] = {"code": "texto = \"papa cuando nace\"\nprint(texto[::2])\nprint(texto[::-1])", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace"}, "ordered_globals": ["texto"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace"}, "ordered_globals": ["texto"], "stack_to_render": [], "heap": {}, "stdout": "pp unonc\n"}, {"line": 3, "event": "return", "func_name": "<module>", "globals": {"texto": "papa cuando nace"}, "ordered_globals": ["texto"], "stack_to_render": [], "heap": {}, "stdout": "pp unonc\necan odnauc apap\n"}]};
</script>
</div>
</div>
<div class="section" id="id2">
<h2>Incremento en el corte<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Comprobar si una palabra es palíndrome</p></li>
</ul>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_1 data-question_label="5">
<div id=ac_l18_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_l18_1_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
palabra = input("Palabra: ")
if palabra == palabra[::-1]:
    print("%s es palíndrome" % palabra)
else:
    print("%s no es un palíndrome" % palabra)
</textarea>
</div>
</div>
</div>
<div class="section" id="string-no-se-puede-modificar">
<h2>String no se puede modificar<a class="headerlink" href="#string-no-se-puede-modificar" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_2 data-question_label="6">
<div id=ac_l18_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_l18_2_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
texto = " Hola mundo!"
texto[0] = "@"
</textarea>
</div>
</div>
</div>
<div class="section" id="puedo-crear-nuevos-strings">
<h2>Puedo crear nuevos strings<a class="headerlink" href="#puedo-crear-nuevos-strings" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Usando la concatenación resolvemos este problema</p></li>
</ul>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_3 data-question_label="7">
<div id=ac_l18_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_l18_3_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
texto = "Hola mundo"
texto = "@" + texto[1:]
print(texto)
</textarea>
</div>
</div>
</div>
<div class="section" id="concentracion">
<h2>Concentración<a class="headerlink" href="#concentracion" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Ejemplo de un programa que lee una palabra y reemplaza las vocales por <code class="docutils literal notranslate"><span class="pre">&quot;*&quot;</span></code>.
La función <code class="docutils literal notranslate"><span class="pre">lower</span></code> transforma las letras en minúsculas.</p></li>
</ul>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_4 data-question_label="8">
<div id=ac_l18_4_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_l18_4_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
palabra = input("Palabra: ")
k = 0
intercambio = ""
while k < len(palabra):
    if palabra[k].lower() in "aeiou":
        intercambio = intercambio + "*"
    else:
        intercambio = intercambio + palabra[k]
    k += 1
print("Nueva palabra %s" % intercambio)
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_5 data-question_label="9">
<div id=ac_l18_5_question class="ac_question col-md-12">
<p>Ahora haga un programa que lea una palabra, la guarde en la variable <code class="docutils literal notranslate"><span class="pre">palabra</span></code>,
y reemplace las consonantes con <code class="docutils literal notranslate"><span class="pre">&quot;*&quot;</span></code>. Guardar el resultado en la variable <code class="docutils literal notranslate"><span class="pre">intercambio</span></code>.
Puede apoyarse en el programa anterior.</p>

</div>
<textarea data-lang="python" id="ac_l18_5_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
# Utilice la función input para leer la palabra del usuario.


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            intercambio,
            "".join(["*" if c.lower() not in "aeiou" else c for c in palabra]),
            "Probando que intercambio esté asignado correctamente",
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="verificacion-parcial-de-strings">
<h2>Verificación parcial de strings<a class="headerlink" href="#verificacion-parcial-de-strings" title="Permalink to this headline">¶</a></h2>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="10">
        <div class="pytutorVisualizer" id="cl_l18_5"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 10 (cl_l18_5)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_5_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_5"] = {"code": "archivo = \"prog.py\"\nprint(archivo.startswith(\"p\"))\nprint(archivo.endswith(\"p\"))\ncontestar = \"Si\"\nprint(contestar.lower())\nprint(contestar.upper())\nprint(contestar.lower() in \"si no yes no\")", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py"}, "ordered_globals": ["archivo"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py"}, "ordered_globals": ["archivo"], "stack_to_render": [], "heap": {}, "stdout": "True\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py"}, "ordered_globals": ["archivo"], "stack_to_render": [], "heap": {}, "stdout": "True\nFalse\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py", "contestar": "Si"}, "ordered_globals": ["archivo", "contestar"], "stack_to_render": [], "heap": {}, "stdout": "True\nFalse\n"}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py", "contestar": "Si"}, "ordered_globals": ["archivo", "contestar"], "stack_to_render": [], "heap": {}, "stdout": "True\nFalse\nsi\n"}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"archivo": "prog.py", "contestar": "Si"}, "ordered_globals": ["archivo", "contestar"], "stack_to_render": [], "heap": {}, "stdout": "True\nFalse\nsi\nSI\n"}, {"line": 7, "event": "return", "func_name": "<module>", "globals": {"archivo": "prog.py", "contestar": "Si"}, "ordered_globals": ["archivo", "contestar"], "stack_to_render": [], "heap": {}, "stdout": "True\nFalse\nsi\nSI\nTrue\n"}]};
</script>
</div>
</div>
<div class="section" id="funciones-find-y-replace">
<h2>Funciones <code class="docutils literal notranslate"><span class="pre">find</span></code> y <code class="docutils literal notranslate"><span class="pre">replace</span></code><a class="headerlink" href="#funciones-find-y-replace" title="Permalink to this headline">¶</a></h2>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="11">
        <div class="pytutorVisualizer" id="cl_l18_6"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 11 (cl_l18_6)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_6_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_6"] = {"code": "s = \"un tigre, dos tigres, tres tigres\"\nprint(s.find(\"tigre\"))\nprint(s.find(\"tigre\", 4))\nprint(s.find(\"tigre\", 16))\nprint(s.replace(\"tigre\", \"gato\"))\ns = s.replace(\"tigre\", \"gato\")\nprint(s)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"s": "un tigre, dos tigres, tres tigres"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"s": "un tigre, dos tigres, tres tigres"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"s": "un tigre, dos tigres, tres tigres"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n14\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"s": "un tigre, dos tigres, tres tigres"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n14\n27\n"}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"s": "un tigre, dos tigres, tres tigres"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n14\n27\nun gato, dos gatos, tres gatos\n"}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"s": "un gato, dos gatos, tres gatos"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n14\n27\nun gato, dos gatos, tres gatos\n"}, {"line": 7, "event": "return", "func_name": "<module>", "globals": {"s": "un gato, dos gatos, tres gatos"}, "ordered_globals": ["s"], "stack_to_render": [], "heap": {}, "stdout": "3\n14\n27\nun gato, dos gatos, tres gatos\nun gato, dos gatos, tres gatos\n"}]};
</script>
</div>
</div>
<div class="section" id="funciones-split-y-join">
<h2>Funciones <code class="docutils literal notranslate"><span class="pre">split</span></code> y <code class="docutils literal notranslate"><span class="pre">join</span></code><a class="headerlink" href="#funciones-split-y-join" title="Permalink to this headline">¶</a></h2>

<div class="runestone" style="max-width: none;">
    <div class="alert alert-warning cd_section" data-component="codelens" data-question_label="12">
        <div class="pytutorVisualizer" id="cl_l18_7"
           data-params='{"embeddedMode": true, "lang": "python", "jumpToEnd": false}'>
        </div>
        <p class="runestone_caption"><span class="runestone_caption_text">Activity: CodeLens 12 (cl_l18_7)</span> </p>
    </div>

<script type="text/javascript">
var cl_l18_7_vis;

if (allTraceData === undefined) {
   var allTraceData = {};
}

allTraceData["cl_l18_7"] = {"code": "texto = \"papa cuando nace\"\nprint(texto.split())\nfecha = \"21/02/2011\"\nprint(fecha.split(\"/\"))\nip = \"198.188.10.144\"\nprint(ip.split(\".\"))\nlugares = [\"Palmeiras\", \"Santos\", \"Corintios\"]\nprint(\"/\".join(lugares))", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace"}, "ordered_globals": ["texto"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace"}, "ordered_globals": ["texto"], "stack_to_render": [], "heap": {}, "stdout": "['papa', 'cuando', 'nace']\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011"}, "ordered_globals": ["texto", "fecha"], "stack_to_render": [], "heap": {}, "stdout": "['papa', 'cuando', 'nace']\n"}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011"}, "ordered_globals": ["texto", "fecha"], "stack_to_render": [], "heap": {}, "stdout": "['papa', 'cuando', 'nace']\n['21', '02', '2011']\n"}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011", "ip": "198.188.10.144"}, "ordered_globals": ["texto", "fecha", "ip"], "stack_to_render": [], "heap": {}, "stdout": "['papa', 'cuando', 'nace']\n['21', '02', '2011']\n"}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011", "ip": "198.188.10.144"}, "ordered_globals": ["texto", "fecha", "ip"], "stack_to_render": [], "heap": {}, "stdout": "['papa', 'cuando', 'nace']\n['21', '02', '2011']\n['198', '188', '10', '144']\n"}, {"line": 8, "event": "step_line", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011", "ip": "198.188.10.144", "lugares": ["REF", 1]}, "ordered_globals": ["texto", "fecha", "ip", "lugares"], "stack_to_render": [], "heap": {"1": ["LIST", "Palmeiras", "Santos", "Corintios"]}, "stdout": "['papa', 'cuando', 'nace']\n['21', '02', '2011']\n['198', '188', '10', '144']\n"}, {"line": 8, "event": "return", "func_name": "<module>", "globals": {"texto": "papa cuando nace", "fecha": "21/02/2011", "ip": "198.188.10.144", "lugares": ["REF", 1]}, "ordered_globals": ["texto", "fecha", "ip", "lugares"], "stack_to_render": [], "heap": {"1": ["LIST", "Palmeiras", "Santos", "Corintios"]}, "stdout": "['papa', 'cuando', 'nace']\n['21', '02', '2011']\n['198', '188', '10', '144']\nPalmeiras/Santos/Corintios\n"}]};
</script>
</div>
</div>
<div class="section" id="ejercicio">
<h2>Ejercicio<a class="headerlink" href="#ejercicio" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l18_6 data-question_label="13">
<div id=ac_l18_6_question class="ac_question col-md-12">
<p>Haga un programa que solicite la fecha de nacimiento en formato “dd/mm/aaaa” y
convierta esta fecha a formato “&lt;día&gt; de &lt;mes&gt; de &lt;año&gt;” usando la lista <code class="docutils literal notranslate"><span class="pre">meses</span></code>
que ya está escrita. Guardar el resultado en la variable <code class="docutils literal notranslate"><span class="pre">fecha_de_nacimiento</span></code> e
imprimir esta variable. Recordar que <code class="docutils literal notranslate"><span class="pre">.split()</span></code> regresa una lista y se le puede pasar
como argumento el caracter con el cual separar un string. Guardar el día, mes y el año en formato de cadena
en las variables <code class="docutils literal notranslate"><span class="pre">dia</span></code>, <code class="docutils literal notranslate"><span class="pre">mes</span></code> y <code class="docutils literal notranslate"><span class="pre">anio</span></code>.</p>

</div>
<textarea data-lang="python" id="ac_l18_6_editor" 
      data-timelimit=25000     
    data-audio=''   data-stdin=''   
           data-wasm=/_static
      style="visibility: hidden;">
fecha = input("fecha (dd/mm/aaaa): ").split("/")
meses = [
    "enero",
    "febrero",
    "marzo",
    "abril",
    "mayo",
    "junio",
    "julio",
    "agosto",
    "septiembre",
    "octubre",
    "noviembre",
    "diciembre",
]

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            fecha_de_nacimiento,
            "{} de {} de {}".format(dia, meses[int(mes) - 1], anio),
            "Probando que fecha_de_nacimiento esté asignada correctamente",
        )
        self.assertEqual(dia, fecha[0], "Probando que dia esté asignada correctamente")
        self.assertEqual(mes, fecha[1], "Probando que mes esté asignado correctamente")
        self.assertEqual(anio, fecha[2], "Probando que anio esté asignado correctamente")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="dojo-de-codificacion">
<h2>Dojo de codificación<a class="headerlink" href="#dojo-de-codificacion" title="Permalink to this headline">¶</a></h2>
<a class="reference internal image-reference" href="../_images/TWP18_015.jpeg"><img alt="" class="align-center" src="../_images/TWP18_015.jpeg" style="width: 19.001cm; height: 14.251cm;" /></a>
</div>
<div class="section" id="id3">
<h2>Dojo de codificación<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h2>
<ul class="simple">
<li><p>Desarrollo basado en pruebas</p></li>
<li><p>Pasos de bebé</p></li>
<li><p>Programación de pares</p></li>
</ul>
</div>
<div class="section" id="desarrollo-basado-en-pruebas">
<h2>Desarrollo basado en pruebas<a class="headerlink" href="#desarrollo-basado-en-pruebas" title="Permalink to this headline">¶</a></h2>
<a class="reference internal image-reference" href="../_images/TWP18_016.png"><img alt="" class="align-center" src="../_images/TWP18_016.png" style="width: 17.726cm; height: 11.032cm;" /></a>
</div>
<div class="section" id="pasos-de-bebe">
<h2>Pasos de bebé<a class="headerlink" href="#pasos-de-bebe" title="Permalink to this headline">¶</a></h2>
<a class="reference internal image-reference" href="../_images/TWP18_017.jpeg"><img alt="" class="align-center" src="../_images/TWP18_017.jpeg" style="width: 17.704cm; height: 12.624cm;" /></a>
</div>
<div class="section" id="programacion-en-pareja">
<h2>Programación en pareja<a class="headerlink" href="#programacion-en-pareja" title="Permalink to this headline">¶</a></h2>
<a class="reference internal image-reference" href="../_images/TWP18_018.png"><img alt="" class="align-center" src="../_images/TWP18_018.png" style="width: 18.201cm; height: 13.711cm;" /></a>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Listas' data-toggle="tooltip">
    <a href="TWP17.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="TWP17.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - For, Funciones y Azar' data-toggle="tooltip" >
    <a href="TWP20.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="TWP20.html" >Next Section - For, Funciones y Azar</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.7.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>