<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 7 &#8212; PyZombis update-test-playwright@6da227a</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.6.2" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=4011C77C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 9" href="Quiz9.html" />
    <link rel="prev" title="Quiz - 6" href="Quiz6.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.6.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis update-test-playwright@6da227a</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz7'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 7</a><ul>
<li><a class="reference internal" href="#ejercicio-a">Ejercicio A</a></li>
<li><a class="reference internal" href="#ejercicio-b">Ejercicio B</a></li>
<li><a class="reference internal" href="#ejercicio-c">Ejercicio C</a></li>
<li><a class="reference internal" href="#ejercicio-d">Ejercicio D</a></li>
<li><a class="reference internal" href="#ejercicio-e">Ejercicio E</a></li>
<li><a class="reference internal" href="#ejercicio-f">Ejercicio F</a></li>
<li><a class="reference internal" href="#ejercicio-g">Ejercicio G</a></li>
<li><a class="reference internal" href="#ejercicio-h">Ejercicio H</a></li>
<li><a class="reference internal" href="#ejercicio-i">Ejercicio I</a></li>
<li><a class="reference internal" href="#ejercicio-j">Ejercicio J</a></li>
<li><a class="reference internal" href="#ejercicio-k">Ejercicio K</a></li>
<li><a class="reference internal" href="#ejercicio-l">Ejercicio L</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 6" data-toggle="tooltip">
    <a href="Quiz6.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 9' data-toggle="tooltip" >              
    <a href="Quiz9.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-7">
<h1>Quiz - 7<a class="headerlink" href="#quiz-7" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-a">
<h2>Ejercicio A<a class="headerlink" href="#ejercicio-a" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=multstring data-question_label="1">
<div id=multstring_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">multstring</span></code> que recibe como parametros una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y un entero positivo <code class="docutils literal notranslate"><span class="pre">n</span></code> y devuelve una nueva cadena que contiene <code class="docutils literal notranslate"><span class="pre">n</span></code> copias de la cadena original <br /> <br />
Ejemplo: <code class="docutils literal notranslate"><span class="pre">multstring('Hola',</span> <span class="pre">2)</span></code> -&gt; ‘HolaHola’</p>

</div>
<textarea data-lang="python" id="multstring_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def multstring(s,n):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(multstring('Hola', 2), 'HolaHola',"Esperado: HolaHola")
        self.assertEqual(multstring('Hola', 3), 'HolaHolaHola',"Esperado: HolaHolaHola")
        self.assertEqual(multstring('Hi', 1), 'Hi',"Esperado: Hi")
        self.assertEqual(multstring('Hi', 0), '',"Esperado: ''")
        self.assertEqual(multstring('Hi', 5), 'HiHiHiHiHi',"Esperado: HiHiHiHiHi")
        self.assertEqual(multstring('Oh Boy!', 2), 'Oh Boy!Oh Boy!',"Esperado: Oh Boy!Oh Boy!")
        self.assertEqual(multstring('x', 4), 'xxxx',"Esperado: xxxx")
        self.assertEqual(multstring('', 4), '',"Esperado: ''")
        self.assertEqual(multstring('code', 2), 'codecode',"Esperado: codecode")
        self.assertEqual(multstring('code', 3), 'codecodecode',"Esperado: codecodecode")

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-b">
<h2>Ejercicio B<a class="headerlink" href="#ejercicio-b" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=string_splosion data-question_label="2">
<div id=string_splosion_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">string_splosion</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelva una nueva cadena siguiendo el patron descrito en los ejemplos <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">string_splosion('Code')</span></code> -&gt; ‘CCoCodCode’ <br />
<code class="docutils literal notranslate"><span class="pre">string_splosion('abc')</span></code> -&gt; ‘aababc’ <br />
<code class="docutils literal notranslate"><span class="pre">string_splosion('ab')</span></code> -&gt; ‘aab’ <br /></p>

</div>
<textarea data-lang="python" id="string_splosion_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def string_splosion(s):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(string_splosion('Code'), 'CCoCodCode',"Esperado: CCoCodCode")
        self.assertEqual(string_splosion('abc'), 'aababc',"Esperado: aababc")
        self.assertEqual(string_splosion('ab'), 'aab',"Esperado: aab")
        self.assertEqual(string_splosion('x'), 'x',"Esperado: x")
        self.assertEqual(string_splosion('fade'), 'ffafadfade',"Esperado: ffafadfade")
        self.assertEqual(string_splosion('There'), 'TThTheTherThere',"Esperado: TThTheTherThere")
        self.assertEqual(string_splosion('Kitten'), 'KKiKitKittKitteKitten',"Esperado: KKiKitKittKitteKitten")
        self.assertEqual(string_splosion('Bye'), 'BByBye',"Esperado: BByBye")
        self.assertEqual(string_splosion('Good'), 'GGoGooGood',"Esperado: GGoGooGood")
        self.assertEqual(string_splosion('Bad'), 'BBaBad',"Esperado: BBaBad")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-c">
<h2>Ejercicio C<a class="headerlink" href="#ejercicio-c" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=array_count9 data-question_label="3">
<div id=array_count9_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">array_count9</span></code> que recibe como parametro una lista no vacía de números enteros <code class="docutils literal notranslate"><span class="pre">nums</span></code> y devuelva cuántas veces aparece el número 9 en la lista <br /> <br />
Ejemplo: <code class="docutils literal notranslate"><span class="pre">array_count9([1,</span> <span class="pre">99,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code></p>

</div>
<textarea data-lang="python" id="array_count9_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def array_count9(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(array_count9([1, 99, 9]), 1,"Esperado: 1")
        self.assertEqual(array_count9([1, 9, 9]), 2,"Esperado: 2")
        self.assertEqual(array_count9([1, 9, 9, 3, 9]), 3,"Esperado: 3")
        self.assertEqual(array_count9([1, 2, 3]), 0,"Esperado: 0")
        self.assertEqual(array_count9([1]), 0,"Esperado: 1")
        self.assertEqual(array_count9([4, 2, 4, 3, 1]), 0,"Esperado: 0")
        self.assertEqual(array_count9([9, 2, 99, 3, 1]), 1,"Esperado: 1")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-d">
<h2>Ejercicio D<a class="headerlink" href="#ejercicio-d" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=array_front9 data-question_label="4">
<div id=array_front9_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">array_front9</span></code> que recibe una lista de números enteros <code class="docutils literal notranslate"><span class="pre">nums</span></code> y verifica si al menos uno de los primeros cuatro números es 9 <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">array_front9([1,</span> <span class="pre">2,</span> <span class="pre">9,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">array_front9([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">array_front9([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4,</span> <span class="pre">5])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">array_front9([1,</span> <span class="pre">2,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="array_front9_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def array_front9(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(array_front9([1, 2, 9, 3, 4]), True,"Esperado: True")
        self.assertEqual(array_front9([1, 2, 3, 4, 9]), False,"Esperado: False")
        self.assertEqual(array_front9([1, 2, 3, 4, 5]), False,"Esperado: False")
        self.assertEqual(array_front9([9, 2, 3]), True,"Esperado: True")
        self.assertEqual(array_front9([1, 9, 9]), True,"Esperado: True")
        self.assertEqual(array_front9([1, 2, 3]), False,"Esperado: False")
        self.assertEqual(array_front9([1, 9]), True,"Esperado: True")
        self.assertEqual(array_front9([5, 5]), False,"Esperado: False")
        self.assertEqual(array_front9([2]), False,"Esperado: False")
        self.assertEqual(array_front9([9]), True,"Esperado: True")
        self.assertEqual(array_front9([]), False,"Esperado: False")
        self.assertEqual(array_front9([3, 9, 2, 3, 3]), True,"Esperado: True")




myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-e">
<h2>Ejercicio E<a class="headerlink" href="#ejercicio-e" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=hello_name data-question_label="5">
<div id=hello_name_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">hello_name</span></code> que recibe como parametro un cadena representa un nombre <code class="docutils literal notranslate"><span class="pre">name</span></code> y devuelva un saludo con este nombre <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">hello_name('Bob')</span></code> -&gt; ‘¡Hola Bob!’ <br />
<code class="docutils literal notranslate"><span class="pre">hello_name('Alice')</span></code> -&gt; ‘¡Hola Alice!’ <br />
<code class="docutils literal notranslate"><span class="pre">hello_name('X')</span></code> -&gt; ‘Hola X!’ <br /></p>

</div>
<textarea data-lang="python" id="hello_name_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def hello_name(name):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(hello_name('Bob'), '¡Hola Bob!',"Esperado: ¡Hola Bob!")
        self.assertEqual(hello_name('Alice'), '¡Hola Alice!',"Esperado: ¡Hola Alice!")
        self.assertEqual(hello_name('X'), '¡Hola X!',"Esperado: ¡Hola X!")
        self.assertEqual(hello_name('Hola'), '¡Hola Hola!',"Esperado: ¡Hola Hola!")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-f">
<h2>Ejercicio F<a class="headerlink" href="#ejercicio-f" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=make_tags data-question_label="6">
<div id=make_tags_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">make_tags</span></code> que recibe dos cadenas <code class="docutils literal notranslate"><span class="pre">tag</span></code> y <code class="docutils literal notranslate"><span class="pre">word</span></code> y devuelva una nueva cadena formateada siguiendo el patrón de los ejemplos <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">make_tags('i',</span> <span class="pre">'Yay')</span></code> -&gt; ‘&lt;i&gt;Yay&lt;/i&gt;’ <br />
<code class="docutils literal notranslate"><span class="pre">make_tags('i',</span> <span class="pre">'Hello')</span></code> -&gt; ‘&lt;i&gt;Hello&lt;/i&gt;’ <br />
<code class="docutils literal notranslate"><span class="pre">make_tags('cite',</span> <span class="pre">'Yay')</span></code> -&gt; ‘&lt;cite&gt;Yay&lt;/cite&gt;’ <br /></p>

</div>
<textarea data-lang="python" id="make_tags_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def make_tags(tag,word):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(make_tags('i', 'Yay'), '<i>Yay</i>',"Esperado: <i>Yay</i>")
        self.assertEqual(make_tags('i', 'Hello'), '<i>Hello</i>',"Esperado:<i>Hello</i>")
        self.assertEqual(make_tags('cite', 'Yay'), '<cite>Yay</cite>',"Esperado: <cite>Yay</cite>")
        self.assertEqual(make_tags('address', 'here'), '<address>here</address>',"Esperado: <address>here</address>")
        self.assertEqual(make_tags('body', 'Heart'), '<body>Heart</body>',"Esperado: <body>Heart</body>")
        self.assertEqual(make_tags('i', 'i'), '<i>i</i>',"Esperado: <i>i</i>")
        self.assertEqual(make_tags('i', ''), '<i></i>',"Esperado: <i></i>")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-g">
<h2>Ejercicio G<a class="headerlink" href="#ejercicio-g" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=extra_end data-question_label="7">
<div id=extra_end_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">extra_end</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelve una nueva cadena con las dos últimas letras repetidas tres veces <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">extra_end('Hello')</span></code> -&gt; ‘lololo’ <br />
<code class="docutils literal notranslate"><span class="pre">extra_end('abb')</span></code> -&gt; ‘bbbbbb’ <br />
<code class="docutils literal notranslate"><span class="pre">extra_end('Hi')</span></code> -&gt; ‘HiHiHi’  <br /></p>

</div>
<textarea data-lang="python" id="extra_end_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def extra_end(s):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(extra_end('Hello'), 'lololo',"Esperado: lololo")
        self.assertEqual(extra_end('ab'), 'ababab',"Esperado: ababab")
        self.assertEqual(extra_end('Hi'), 'HiHiHi',"Esperado: HiHiHi")
        self.assertEqual(extra_end('Candy'), 'dydydy',"Esperado: dydydy")
        self.assertEqual(extra_end('Code'), 'dedede',"Esperado: dedede")



myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-h">
<h2>Ejercicio H<a class="headerlink" href="#ejercicio-h" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=new_extra_end data-question_label="8">
<div id=new_extra_end_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">new_extra_end</span></code> (variante de la función del Ejercicio G) que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y un entero positivo <code class="docutils literal notranslate"><span class="pre">n</span></code> y devuelva una nueva cadena con las dos últimas letras repetidas <code class="docutils literal notranslate"><span class="pre">n</span></code> veces <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">new_extra_end('Hello',3)</span></code> -&gt; ‘lololo’ <br />
<code class="docutils literal notranslate"><span class="pre">new_extra_end('abb',1)</span></code> -&gt; ‘bb’ <br />
<code class="docutils literal notranslate"><span class="pre">new_extra_end('Hi',5)</span></code> -&gt; ‘HiHiHiHiHi’  <br /></p>

</div>
<textarea data-lang="python" id="new_extra_end_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def new_extra_end(s,n):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(new_extra_end('Hello', 3), 'lololo',"Esperado: lololo")
        self.assertEqual(new_extra_end('ab', 2), 'abab',"Esperado: abab")
        self.assertEqual(new_extra_end('Hi', 6), 'HiHiHiHiHiHi',"Esperado: HiHiHi")
        self.assertEqual(new_extra_end('Candy', 1), 'dy',"Esperado: dy")
        self.assertEqual(new_extra_end('Code', 0), '',"Esperado: ''")



myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-i">
<h2>Ejercicio I<a class="headerlink" href="#ejercicio-i" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=first_half data-question_label="9">
<div id=first_half_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">first_half</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelve la primera mitad de la cadena original <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">first_half('WooHoo')</span></code> -&gt; ‘Woo’ <br />
<code class="docutils literal notranslate"><span class="pre">first_half('HelloThere')</span></code> -&gt; ‘Hello’ <br />
<code class="docutils literal notranslate"><span class="pre">first_half('abcdef')</span></code> -&gt; ‘abc’ <br /></p>

</div>
<textarea data-lang="python" id="first_half_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def first_half(s):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(first_half('WooHoo'), 'Woo',"Esperado: Woo")
        self.assertEqual(first_half('HelloThere'), 'Hello',"Esperado: Hello")
        self.assertEqual(first_half('abcdef'), 'abc',"Esperado: abc")
        self.assertEqual(first_half(''), '',"Esperado: ")
        self.assertEqual(first_half('ab'), 'a',"Esperado: a")
        self.assertEqual(first_half('0123456789'), '01234',"Esperado: 01234")
        self.assertEqual(first_half('kitten'), 'kit',"Esperado: kit")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-j">
<h2>Ejercicio J<a class="headerlink" href="#ejercicio-j" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=remove_first_last data-question_label="10">
<div id=remove_first_last_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">remove_first_last</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelva una cadena sin el primer y último carácter <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">remove_first_last('Hello')</span></code> -&gt; ‘ell’ <br />
<code class="docutils literal notranslate"><span class="pre">remove_first_last('python')</span></code> -&gt; ‘ytho’ <br />
<code class="docutils literal notranslate"><span class="pre">remove_first_last('coding')</span></code> -&gt; ‘odin’ <br /></p>

</div>
<textarea data-lang="python" id="remove_first_last_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def remove_first_last(s):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remove_first_last('Hello'), 'ell',"Esperado: ell")
        self.assertEqual(remove_first_last('Python'), 'ytho',"Esperado: ytho")
        self.assertEqual(remove_first_last('coding'), 'odin',"Esperado: odin")
        self.assertEqual(remove_first_last('code'), 'od',"Esperado: od")
        self.assertEqual(remove_first_last('ab'), '',"Esperado: ")
        self.assertEqual(remove_first_last(' PyZombies '), 'PyZombies',"Esperado: PyZombies")
        self.assertEqual(remove_first_last('Chocolate!'), 'hocolate',"Esperado: hocolate")
        self.assertEqual(remove_first_last('kitten'), 'itte',"Esperado: itte")
        self.assertEqual(remove_first_last('woohoo'), 'ooho',"Esperado: ooho")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-k">
<h2>Ejercicio K<a class="headerlink" href="#ejercicio-k" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=left2 data-question_label="11">
<div id=left2_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">left2</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelva la cadena original rotada 2 posiciones a la izquierda <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">left2('Hello')</span></code> -&gt; ‘lloHe’ <br />
<code class="docutils literal notranslate"><span class="pre">left2('Hi')</span></code> -&gt; ‘Hi’ <br /></p>

</div>
<textarea data-lang="python" id="left2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def left2(s):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(left2('Hello'), 'lloHe',"Esperado: lloHe")
        self.assertEqual(left2('python'), 'thonpy',"Esperado: thonpy")
        self.assertEqual(left2('Hi'), 'Hi',"Esperado: Hi")
        self.assertEqual(left2('code'), 'deco',"Esperado: deco")
        self.assertEqual(left2('cat'), 'tca',"Esperado: tca")
        self.assertEqual(left2('12345'), '34512',"Esperado: 34512")
        self.assertEqual(left2('Chocolate'), 'ocolateCh',"Esperado: ocolateCh")
        self.assertEqual(left2('bricks'), 'icksbr',"Esperado: icksbr")
        self.assertEqual(left2('isPyzomb'), 'PyZombies',"Esperado: PyZombies")



myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-l">
<h2>Ejercicio L<a class="headerlink" href="#ejercicio-l" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=remove_iesimo data-question_label="12">
<div id=remove_iesimo_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">remove_iesimo</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> no vacía y un entero positivo <code class="docutils literal notranslate"><span class="pre">i</span></code> y devuelva la cadena original sin el i-esimo carácter <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">remove_iesimo('Hello',1)</span></code> -&gt; ‘ello’ <br />
<code class="docutils literal notranslate"><span class="pre">remove_iesimo('Hi',2)</span></code> -&gt; ‘H’ <br />
<code class="docutils literal notranslate"><span class="pre">remove_iesimo('PyZombiess',</span> <span class="pre">10)</span></code> -&gt; ‘PyZombies’ <br /></p>

</div>
<textarea data-lang="python" id="remove_iesimo_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def remove_iesimo(s,i):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(remove_iesimo('Hello', 1), 'ello',"Esperado: ello")
        self.assertEqual(remove_iesimo('python', 6), 'pytho',"Esperado: pytho")
        self.assertEqual(remove_iesimo('Hi', 0), 'Hi',"Esperado: Hi")
        self.assertEqual(remove_iesimo('code', 5), 'code',"Esperado: code")
        self.assertEqual(remove_iesimo('cat', 1), 'ct',"Esperado: ct")
        self.assertEqual(remove_iesimo('12345', 3), '1245',"Esperado: 1245")
        self.assertEqual(remove_iesimo('Chocolate', 15), 'Chocolate',"Esperado: Chocolate")
        self.assertEqual(remove_iesimo('Hola mundo', 5), 'Holamundo',"Esperado: Holamundo")
        self.assertEqual(remove_iesimo('Pyzombis ', 9), 'PyZombies',"Esperado: PyZombies")



myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 6' data-toggle="tooltip">
    <a href="Quiz6.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz6.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 9' data-toggle="tooltip" >
    <a href="Quiz9.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz9.html" >Next Section - Quiz - 9</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.6.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>