<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 11 &#8212; PyZombis localstorage@710f0aa</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=1654bbde73e12a3366eb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis localstorage@710f0aa
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz11_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-11">
<h1>Quiz - 11<a class="headerlink" href="#quiz-11" title="Permalink to this headline">¶</a></h1>
<div id='quiz11' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 13)</p>
<p>Duplicate ID – see quiz/Quiz11, line 17</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  activecode:: q11_1
    :nocodelens:

    Develop the function ``verbo`` that receives a string ``s`` as a parameter. If the length of the string is at least 3, it should return the original string concatenated with ``&quot;ing&quot;`` at the end. If the string ``s`` already ends with ``&quot;ing&quot;``, concatenate the string ``&quot;ly&quot;``. If the length of the string is less than 3, it returns the original string. |br| |br|
    Examples: |br|
    ``verbo(&quot;singing&quot;)`` -&gt; ``&quot;singingly&quot;`` |br|
    ``verbo(&quot;travel&quot;)`` -&gt; ``&quot;traveling&quot;`` |br|
    ``verbo(&quot;do&quot;)`` -&gt; ``&quot;do&quot;`` |br|

    ~~~~
    def verbo(s):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(verbo(&quot;hail&quot;), &quot;hailing&quot;, &quot;Expected: hailing&quot;)
            self.assertEqual(verbo(&quot;swiming&quot;), &quot;swimingly&quot;, &quot;Expected: swimingly&quot;)
            self.assertEqual(verbo(&quot;do&quot;), &quot;do&quot;, &quot;Expected: do&quot;)
            self.assertEqual(verbo(&quot;singing&quot;), &quot;singingly&quot;, &quot;Expected: singingly&quot;)
            self.assertEqual(verbo(&quot;travel&quot;), &quot;traveling&quot;, &quot;Expected: traveling&quot;)
            self.assertEqual(verbo(&quot;lly&quot;), &quot;llying&quot;, &quot;Expected: llying&quot;)
            self.assertEqual(verbo(&quot;ing&quot;), &quot;ingly&quot;, &quot;Expected: ingly&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 48)</p>
<p>Duplicate ID – see quiz/Quiz11, line 52</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  activecode:: q11_2
    :nocodelens:

    Develop the function ``no_es_malo`` that receives a string ``s`` as a parameter. The function must search for the first occurrence of the string ``&quot;no es&quot;`` and the last occurrence of the string ``&quot;malo&quot;`` or the string ``&quot;mala&quot;``, if either appears after the first one, replace ``&quot;no es&quot; ... &quot;malo&quot;`` or ``&quot;no es&quot; ... &quot;mala&quot;`` with the strings ``&quot;es bueno&quot;`` or ``&quot;es buena&quot;`` respectively, then return the result. |br| |br|
    Examples: |br|
    ``no_es_malo(&quot;El té no es malo&quot;)`` -&gt; ``&quot;El té es bueno&quot;`` |br|
    ``no_es_malo(&quot;La película no es mala&quot;)`` -&gt; ``&quot;La película es buena&quot;`` |br|
    ``no_es_malo(&quot;El precio de esta casa no es para nada malo&quot;)`` -&gt; ``&quot;El precio de esta casa es bueno&quot;`` |br|
    ``no_es_malo(&quot;El teléfono es malo&quot;)`` -&gt; ``&quot;El teléfono es malo&quot;`` |br|

    ~~~~
    def no_es_malo(s):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(
                no_es_malo(&quot;El televisor no es malo&quot;),
                &quot;El televisor es bueno&quot;,
                &quot;Expected: El televisor es bueno&quot;
            )
            self.assertEqual(
                no_es_malo(&quot;El asado de la cena no es malo!&quot;),
                &quot;El asado de la cena es bueno!&quot;,
                &quot;Expected: El asado de la cena es bueno!&quot;
            )
            self.assertEqual(
                no_es_malo(&quot;El té no está caliente&quot;),
                &quot;El té no está caliente&quot;,
                &quot;Expected: El té no está caliente&quot;
            )
            self.assertEqual(
                no_es_malo(&quot;La película no es mala&quot;),
                &quot;La película es buena&quot;,
                &quot;Expected: La película es buena&quot;
            )
            self.assertEqual(no_es_malo(&quot;no es para nada malo&quot;), &quot;es bueno&quot;, &quot;Expected: es bueno&quot;)
            self.assertEqual(no_es_malo(&quot;no es malo&quot;), &quot;es bueno&quot;, &quot;Expected: es bueno&quot;)
            self.assertEqual(no_es_malo(&quot;malo&quot;), &quot;malo&quot;, &quot;Expected: malo&quot;)
            self.assertEqual(no_es_malo(&quot;no&quot;), &quot;no&quot;, &quot;Expected: no&quot;)
            self.assertEqual(no_es_malo(&quot;NO&quot;), &quot;NO&quot;, &quot;Expected: NO&quot;)
            self.assertEqual(no_es_malo(&quot;MALO&quot;), &quot;MALO&quot;, &quot;Expected: MALO&quot;)
            self.assertEqual(no_es_malo(&quot;NO es MALO&quot;), &quot;NO es MALO&quot;, &quot;Expected: NO es MALO&quot;)
            self.assertEqual(no_es_malo(&quot;no es MALO&quot;), &quot;no es MALO&quot;, &quot;Expected: no es MALO&quot;)
            self.assertEqual(no_es_malo(&quot;NO es malo&quot;), &quot;NO es malo&quot;, &quot;Expected: NO es malo&quot;)
            self.assertEqual(no_es_malo(&quot;no es malo ni mala&quot;), &quot;es buena&quot;, &quot;Expected: es buena&quot;)
            self.assertEqual(no_es_malo(&quot;no es ni mala ni malo&quot;), &quot;es bueno&quot;, &quot;Expected: es bueno&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 108)</p>
<p>Duplicate ID – see quiz/Quiz11, line 112</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q11_3
    :nocodelens:

    Develop the function ``inicio_final`` that receives two strings ``a`` and ``b``. The strings have to be divided into two, if either of the strings has an odd number of characters, the first half will be the longest substring (for example ``dog`` will be divided into: ``do`` and ``g``). Given the two strings, return a new string formed as follows ``a_start + b_start + a_end + b_end``. |br| |br|
    Examples: |br|
    ``inicio_final(&quot;abcd&quot;, &quot;1234&quot;)`` -&gt; ``&quot;ab12cd34&quot;`` |br|
    ``inicio_final(&quot;abc&quot;, &quot;1234&quot;)`` -&gt; ``&quot;ab12c34&quot;`` |br|
    ``inicio_final(&quot;abc&quot;, &quot;123&quot;)`` -&gt; ``&quot;ab12c3&quot;`` |br|

    ~~~~
    def inicio_final(a, b):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(inicio_final(&quot;abcd&quot;, &quot;xy&quot;), &quot;abxcdy&quot;, &quot;Expected: abxcdy&quot;)
            self.assertEqual(inicio_final(&quot;abcde&quot;, &quot;xyz&quot;), &quot;abcxydez&quot;, &quot;Expected: abcxydez&quot;)
            self.assertEqual(inicio_final(&quot;a&quot;, &quot;b&quot;), &quot;ab&quot;, &quot;Expected: ab&quot;)
            self.assertEqual(inicio_final(&quot;ac&quot;, &quot;b&quot;), &quot;abc&quot;, &quot;Expected: abc&quot;)
            self.assertEqual(inicio_final(&quot;a&quot;, &quot;bc&quot;), &quot;abc&quot;, &quot;Expected: abc&quot;)
            self.assertEqual(inicio_final(&quot;&quot;, &quot;&quot;), &quot;&quot;, &quot;Expected: &#39;&#39;&quot;)
            self.assertEqual(inicio_final(&quot;a&quot;, &quot;&quot;), &quot;a&quot;, &quot;Expected: &#39;a&#39;&quot;)
            self.assertEqual(inicio_final(&quot;&quot;, &quot;b&quot;), &quot;b&quot;, &quot;Expected: &#39;b&#39;&quot;)
            self.assertEqual(
                inicio_final(&quot;Kitten&quot;, &quot;Donut&quot;),
                &quot;KitDontenut&quot;,
                &quot;Expected: KitDontenut&quot;
            )


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 149)</p>
<p>Duplicate ID – see quiz/Quiz11, line 153</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q11_4
    :nocodelens:

    Develop the function ``cuantos_ceros`` that given a positive integer ``n``, returns the number of zeros at the end of the integer. |br| |br|
    Examples: |br|
    ``cuantos_ceros(10010)`` -&gt; ``1`` |br|
    ``cuantos_ceros(908007000)`` -&gt; ``3`` |br|

    ~~~~
    def cuantos_ceros(n):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(cuantos_ceros(10100100010000), 4, &quot;Expected: 4&quot;)
            self.assertEqual(cuantos_ceros(90000000000000000010), 1, &quot;Expected: 1&quot;)
            self.assertEqual(cuantos_ceros(10), 1, &quot;Expected: 1&quot;)
            self.assertEqual(cuantos_ceros(1050051222), 0, &quot;Expected: 0&quot;)
            self.assertEqual(cuantos_ceros(1010101010), 1, &quot;Expected: 1&quot;)
            self.assertEqual(cuantos_ceros(5000), 3, &quot;Expected: 3&quot;)
            self.assertEqual(cuantos_ceros(10000000000), 10, &quot;Expected: 10&quot;)
            self.assertEqual(cuantos_ceros(555), 0, &quot;Expected: 0&quot;)
            self.assertEqual(cuantos_ceros(1), 0, &quot;Expected: 0&quot;)
            self.assertEqual(cuantos_ceros(0), 0, &quot;Expected: 0&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 186)</p>
<p>Duplicate ID – see quiz/Quiz11, line 190</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q11_5
    :nocodelens:

    Develop the function ``contar_2`` that receives a positive integer ``n`` greater than 0. The function must return the number of times the digit 2 appears in the interval``[0, n-1]``. |br| |br|
    Examples: |br|
    ``contar_2(20)`` -&gt; ``2`` |br|
    ``contar_2(5)`` -&gt; ``1`` |br|
    ``contar_2(1)`` -&gt; ``0`` |br|

    ~~~~
    def contar_2(n):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(contar_2(20), 2, &quot;Expected: 2&quot;)
            self.assertEqual(contar_2(1), 0, &quot;Expected: 0&quot;)
            self.assertEqual(contar_2(5), 1, &quot;Expected: 1&quot;)
            self.assertEqual(contar_2(999), 300, &quot;Expected: 300&quot;)
            self.assertEqual(contar_2(555), 216, &quot;Expected: 216&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz11_en.rst</span>, line 219)</p>
<p>Duplicate ID – see quiz/Quiz11, line 223</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q11_6
    :nocodelens:

    Develop the function ``inicio_potencia`` that receives a positive integer ``n`` greater than 0. The function must return the first power of 2 that starts with ``n``. |br| |br|
    Examples: |br|
    ``inicio_potencia(65)`` -&gt; ``16`` |br|
    *Explanation*: for ``n = 65`` the power of ``2^16`` results in ``65536`` which contains ``n`` at the beginning. |br| |br|
    ``inicio_potencia(4)`` -&gt; ``2`` |br|
    *Explanation*: for ``n = 4`` the power of ``2^2`` results in ``4`` which contains ``n`` at the beginning. |br| |br|
    ``inicio_potencia(3)`` -&gt; ``5`` |br|
    *Explanation*: for ``n = 3`` the power of ``2^5`` results in ``32`` which contains ``n`` at the beginning. |br|

    ~~~~
    def inicio_potencia(n):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(inicio_potencia(7), 46, &quot;Expected: 46&quot;)
            self.assertEqual(inicio_potencia(3), 5, &quot;Expected: 5&quot;)
            self.assertEqual(inicio_potencia(133), 316, &quot;Expected: 316&quot;)
            self.assertEqual(inicio_potencia(1024), 10, &quot;Expected: 10&quot;)
            self.assertEqual(inicio_potencia(123), 90, &quot;Expected: 90&quot;)
            self.assertEqual(inicio_potencia(1), 0, &quot;Expected: 0&quot;)
            self.assertEqual(inicio_potencia(10), 10, &quot;Expected: 10&quot;)
            self.assertEqual(inicio_potencia(50), 102, &quot;Expected: 102&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>