<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>3. Probando la API de TVMaze &#8212; PyZombis aayush/issue#341@884246d</title>
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../../" src="../../_static/documentation_options.js"></script>
    <script src="../../_static/jquery.js"></script>
    <script src="../../_static/underscore.js"></script>
    <script src="../../_static/doctools.js"></script>
    <script src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../../_static/runtime.bundle.js?v=e476ce06ce356667798d"></script>
    <script src="../../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="4. Probando la API de University Domains" href="TWP45_4.html" />
    <link rel="prev" title="2. Obtener datos de Reddit" href="TWP45_2.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../../index.html" aria-label="index-page">
          
          PyZombis aayush/issue#341@884246d
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=TWP45_3'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">2. </span>Obtener datos de Reddit" data-toggle="tooltip">
    <a href="TWP45_2.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">4. </span>Probando la API de University Domains' data-toggle="tooltip" >              
    <a href="TWP45_4.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 2. Obtener datos de Reddit' data-toggle="tooltip">
  <a href="TWP45_2.html" aria-label="Previous - 2. Obtener datos de Reddit">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4. Probando la API de University Domains' data-toggle="tooltip" >
    <a href="TWP45_4.html" aria-label="Next - 4. Probando la API de University Domains">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="probando-la-api-de-tvmaze">
<h1><span class="section-number">3. </span>Probando la API de TVMaze<a class="headerlink" href="#probando-la-api-de-tvmaze" title="Permalink to this headline">¶</a></h1>
<p>TVMaze es una herramienta que:</p>
<blockquote>
<div><p>Te permite buscar programas de televisión y obtener información sobre ellos.
Proporciona una API RESTful para acceder a los datos.</p>
</div></blockquote>
<p>En el siguiente ejemplo, utilizaremos la API de TVMaze para buscar programas de televisión. La URL base es API de <a class="reference external" href="http://api.tvmaze.com/search/shows">TVMaze</a>.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l45_3a_es data-question_label="3.1">
<div id=ac_l45_3a_es_question class="ac_question col-md-12">
<p>En este caso, utilizaremos la biblioteca <code class="docutils literal notranslate"><span class="pre">requests</span></code> para hacer la solicitud a la API. La URL base
es <code class="docutils literal notranslate"><span class="pre">&quot;https://api.tvmaze.com/search/shows&quot;</span></code>. A esta URL se le pasará un parámetro <code class="docutils literal notranslate"><span class="pre">q</span></code> con el
valor del programa “big bang theory”. Finalmente, la URL se verá así: <code class="docutils literal notranslate"><span class="pre">&quot;https://api.tvmaze.com/search/shows?q=big+bang+theory&quot;</span></code>.
Ten en cuenta que después de la URL base, se escribe un <code class="docutils literal notranslate"><span class="pre">?</span></code> para indicar que siguen los parámetros.</p>

</div>
<textarea data-lang="python" id="ac_l45_3a_es_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
import requests
import json

api_url = "http://api.tvmaze.com/search/shows"

# Los parámetros que se pasarán a la URL se escriben dentro de un diccionario
# no necesitamos agregar espacios en los parámetros, porque la biblioteca requests lo hará por nosotros
parameters = {"q": "big bang theory"}


# Solicitamos los datos de la API
response = requests.get(api_url, params=parameters)

# Ahora imprimimos la URL
print(response.url)
print()

# Transformamos los datos en formato JSON a datos de Python
data = json.loads(response.text)

print(data)
</textarea>
</div>
</div>
<p>En el ejemplo anterior, pudiste ver que la API devuelve texto, que si se pasa a través de <code class="docutils literal notranslate"><span class="pre">json.loads</span></code>,
se transforma en un diccionario de Python. Sin embargo, no es completamente legible. Esto se puede solucionar con
<code class="docutils literal notranslate"><span class="pre">json.dumps</span></code>.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l45_3b_es data-question_label="3.2">
<div id=ac_l45_3b_es_question class="ac_question col-md-12">
<p>Ahora, solicitaremos información sobre el programa “golden girls”. Esta vez imprimiremos los datos en un
formato legible para el usuario. Utilizaremos <code class="docutils literal notranslate"><span class="pre">urllib</span></code> para hacer la solicitud.</p>

</div>
<textarea data-lang="python3" id="ac_l45_3b_es_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     data-python3_interpreter='brython'  style="visibility: hidden;">
import urllib.request
import urllib.parse
import json

api_url = "http://api.tvmaze.com/search/shows?"
# La siguiente línea es para los parámetros de la URL
# no necesitamos agregar espacios en los parámetros, porque la biblioteca requests lo hará por nosotros
parameters = urllib.parse.urlencode({"q": "golden girls"})

request = urllib.request.urlopen(api_url + parameters)
data = json.loads(request.read())

# Imprimimos los datos en un formato legible para el usuario
print(json.dumps(data, indent=4))
print("La consulta devolvió " + str(len(data)) + " resultados")
</textarea>
</div>
</div>
<div class="line-block">
<div class="line"><br /></div>
</div>
<p>El siguiente ejercicio viene con calificación automática.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_l45_3c_es data-question_label="3.3">
<div id=ac_l45_3c_es_question class="ac_question col-md-12">
<p>Ahora pediremos a TVMaze el programa “suits”. Luego, al diccionario <code class="docutils literal notranslate"><span class="pre">parameters</span></code> se le deberá asignar el valor
<code class="docutils literal notranslate"><span class="pre">&quot;suits&quot;</span></code> a la clave <code class="docutils literal notranslate"><span class="pre">&quot;q&quot;</span></code>.</p>
<p>Primero, solicitarás desde la API lo que se describió anteriormente, y guardarás esto en la variable <code class="docutils literal notranslate"><span class="pre">request</span></code>.
En otra variable, <code class="docutils literal notranslate"><span class="pre">request_url</span></code>, guarda la URL de la solicitud. Luego, asigna los datos a la variable <code class="docutils literal notranslate"><span class="pre">data</span></code>.
A continuación, asignarás a la variable <code class="docutils literal notranslate"><span class="pre">results</span></code> el número de resultados que devolvió la solicitud
(como se hizo en el ejemplo anterior).</p>
<p>Ahora, calcularás el promedio de clasificación de todos los programas que se devolvieron.
Guardarás esto en la variable <code class="docutils literal notranslate"><span class="pre">average_rating</span></code>. <strong>Sugerencia</strong>: las clasificaciones se encuentran dentro de
<code class="docutils literal notranslate"><span class="pre">data[&quot;show&quot;][&quot;rating&quot;][&quot;average&quot;]</span></code>. Necesitarás usar un bucle for para calcular el promedio.</p>
<p>Por último, buscarás el número de veces que aparece la palabra <code class="docutils literal notranslate"><span class="pre">&quot;Drama&quot;</span></code> en los géneros relacionados con suits.
Guarda ese número en la variable <code class="docutils literal notranslate"><span class="pre">drama_count</span></code>. <strong>Sugerencia</strong>: los géneros se encuentran dentro de
<code class="docutils literal notranslate"><span class="pre">data[&quot;show&quot;][&quot;genres&quot;]</span></code>. Necesitarás usar un bucle for para calcular el número de dramas.</p>

</div>
<textarea data-lang="python" id="ac_l45_3c_es_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
import requests
import json

api_url = "http://api.tvmaze.com/search/shows"
parameters = {}
request =
request_url =
data =
results =

total_rating = 0
rated_shows = 0
for show in data:
# completar el bucle

drama_count = 0
for show in data:
# completar el bucle

print("URL de la solicitud:", request_url)
print("Resultados:", results)
print("Clasificación promedio:", average_rating)
print("Número de dramas:", drama_count)

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            request_url,
            "http://api.tvmaze.com/search/shows?q=suits",
            "Probando que la URL sea: http://api.tvmaze.com/search/shows?q=suits",
        )
        self.assertEqual(results, 10, "Probando que los resultados se asignen correctamente.")
        self.assertEqual(average_rating, 7.8, "Probando que la clasificación promedio sea: 7.8")
        self.assertEqual(drama_count, 6, "Probando que drama_count se asigne correctamente.")


myTests().main()
</textarea>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - 2. Obtener datos de Reddit' data-toggle="tooltip">
  <a href="TWP45_2.html" aria-label="Previous - 2. Obtener datos de Reddit">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - 4. Probando la API de University Domains' data-toggle="tooltip" >
    <a href="TWP45_4.html" aria-label="Next - 4. Probando la API de University Domains">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>