<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Quiz - 9 &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=-7F7467D6" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css?v=-1DDF53CD" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css?v=C6382F3" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css?v=-2F4CB85B" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css?v=1CE53237" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=-6EF20E1F" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css?v=424F1CC9" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css?v=-46455F2F" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css?v=3BA0E63E" />
    <link rel="stylesheet" type="text/css" href="../_static/reveal.css?v=4A81D381" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css?v=-4FCA5E27" />
    <link rel="stylesheet" type="text/css" href="../_static/js_lib/prettify.css?v=2A364B6B" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css?v=239FC54E" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css?v=-75E5BD1C" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css?v=-1AC7A4C5" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css?v=58C4E18A" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css?v=-3E7B1476" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js?v=-7CD635AB"></script>
    <script type="text/javascript" src="../_static/timedmc.js?v=-4FEF5BA9"></script>
    <script type="text/javascript" src="../_static/timed.js?v=6789FE30"></script>
    <script type="text/javascript" src="../_static/poll.js?v=6C59AE9F"></script>
    <script type="text/javascript" src="../_static/showEval.js?v=-3A31F37F"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js?v=-E3B6938"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js?v=-127BC6C3"></script>
    <script type="text/javascript" src="../_static/codemirror.js?v=-6B0AAD37"></script>
    <script type="text/javascript" src="../_static/xml.js?v=236EDFF7"></script>
    <script type="text/javascript" src="../_static/css.js?v=623153CE"></script>
    <script type="text/javascript" src="../_static/python.js?v=14A28F1F"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js?v=33B91944"></script>
    <script type="text/javascript" src="../_static/javascript.js?v=6F38CA04"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js?v=599B42B0"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js?v=-775EBCF1"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js?v=-51EA5B5F"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js?v=-75420181"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js?v=747279FF"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js?v=7B020E13"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js?v=660E85C3"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js?v=54B21948"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js?v=3FBAE230"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js?v=-347AA2A9"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js?v=3C6652E0"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js?v=-5FB79FE6"></script>
    <script type="text/javascript" src="../_static/clike.js?v=2DCE0756"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js?v=5150FA78"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js?v=-7BF73669"></script>
    <script type="text/javascript" src="../_static/Blob.js?v=AEC2447"></script>
    <script type="text/javascript" src="../_static/matrixeq.js?v=-29B305A6"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js?v=5BE6106"></script>
    <script type="text/javascript" src="../_static/dragndrop.js?v=2B6403E8"></script>
    <script type="text/javascript" src="../_static/timeddnd.js?v=378F237E"></script>
    <script type="text/javascript" src="../_static/shortanswer.js?v=5EB6170D"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js?v=-4C5EE829"></script>
    <script type="text/javascript" src="../_static/fitb.js?v=-71BF18ED"></script>
    <script type="text/javascript" src="../_static/timedfitb.js?v=2AA17071"></script>
    <script type="text/javascript" src="../_static/reveal.js?v=7C095993"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js?v=-358176D8"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js?v=73056367"></script>
    <script type="text/javascript" src="../_static/parsons.js?v=-7278C3F6"></script>
    <script type="text/javascript" src="../_static/timedparsons.js?v=2EEA9E5A"></script>
    <script type="text/javascript" src="../_static/animationbase.js?v=-5FBC3C59"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js?v=836E01C"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js?v=-49B30F58"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js?v=42F9EBEA"></script>
    <script type="text/javascript" src="../_static/pytutor.js?v=-7D4FDD47"></script>
    <script type="text/javascript" src="../_static/codelens.js?v=-354132B7"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js?v=53E888DE"></script>
    <script type="text/javascript" src="../_static/clickable.js?v=-6A9F19A2"></script>
    <script type="text/javascript" src="../_static/timedclickable.js?v=-6C6763DD"></script>
    <script type="text/javascript" src="../_static/datafile.js?v=3ED1A80C"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 10" href="Quiz10.html" />
    <link rel="prev" title="Quiz - 7" href="Quiz7.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '3.3.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis Python Course</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz9'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 9</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 7" data-toggle="tooltip">
    <a href="Quiz7.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 10' data-toggle="tooltip" >              
    <a href="Quiz10.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-9">
<h1>Quiz - 9<a class="headerlink" href="#quiz-9" title="Permalink to this headline">¶</a></h1>

<div data-childcomponent="first_last6" class="runestone explainer ac_section alert alert-warning">
A. first_last6 <br />
verifica se 6 é o primeiro ou último elemento da lista nums <br />
first_last6([1, 2, 6]) -&gt; True <br />
first_last6([6, 1, 2, 3]) -&gt; True <br />
first_last6([3, 2, 1]) -&gt; False <br />
<textarea data-component="activecode" id=first_last6 data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def first_last6(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(first_last6([1, 2, 6]), True,"Esperado: True")
        self.assertEqual(first_last6([6, 1, 2, 3]), True,"Esperado: True")
        self.assertEqual(first_last6([3, 2, 1]), False,"Esperado: False")
        self.assertEqual(first_last6([3, 6, 1]), False,"Esperado: False")
        self.assertEqual(first_last6([3, 6]), True,"Esperado: True")
        self.assertEqual(first_last6([3, 6]), True,"Esperado: True")
        self.assertEqual(first_last6([3]), False,"Esperado: False")


myTests().main()
</textarea>
</div>

<div data-childcomponent="same_first_last" class="runestone explainer ac_section alert alert-warning">
B. same_first_last <br />
retorna True se a lista nums possui pelo menos um elemento <br />
e o primeiro elemento é igual ao último <br />
same_first_last([1, 2, 3]) -&gt; False <br />
same_first_last([1, 2, 3, 1]) -&gt; True <br />
same_first_last([1, 2, 1]) -&gt; True <br />
<textarea data-component="activecode" id=same_first_last data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def same_first_last(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(same_first_last([1, 2, 3]), False,"Esperado: False")
        self.assertEqual(same_first_last([1, 2, 3, 1]), True,"Esperado: True")
        self.assertEqual(same_first_last([1, 2, 1]), True,"Esperado: True")
        self.assertEqual(same_first_last([7]), True,"Esperado: True")
        self.assertEqual(same_first_last([]), False,"Esperado: False")
        self.assertEqual(same_first_last([7, 7]), True,"Esperado: True")


myTests().main()
</textarea>
</div>

<div data-childcomponent="common_end" class="runestone explainer ac_section alert alert-warning">
C. common_end <br />
Dada duas listas a e b verifica se os dois primeiros são <br />
iguais ou os dois últimos são iguais <br />
suponha que as listas tenham pelo menos um elemento <br />
common_end([1, 2, 3], [7, 3]) -&gt; True <br />
common_end([1, 2, 3], [7, 3, 2]) -&gt; False <br />
common_end([1, 2, 3], [1, 3]) -&gt; True <br />
<textarea data-component="activecode" id=common_end data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def common_end(a, b):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(common_end([1, 2, 3], [7, 3]), True,"Esperado: True")
        self.assertEqual(common_end([1, 2, 3], [7, 3, 2]), False,"Esperado: False")
        self.assertEqual(common_end([1, 2, 3], [1, 3]), True,"Esperado: True")
        self.assertEqual(common_end([1, 2, 3], [1]), True,"Esperado: True")
        self.assertEqual(common_end([1, 2, 3], [2]), False,"Esperado: False")



myTests().main()
</textarea>
</div>

<div data-childcomponent="maior_ponta" class="runestone explainer ac_section alert alert-warning">
D. maior_ponta <br />
Dada uma lista não vazia, cria uma nova lista onde todos <br />
os elementos são o maior das duas pontas <br />
obs.: não é o maior de todos, mas entre as duas pontas <br />
maior_ponta([1, 2, 3]) -&gt; [3, 3, 3] <br />
maior_ponta([1, 3, 2]) -&gt; [2, 2, 2] <br />
<textarea data-component="activecode" id=maior_ponta data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def maior_ponta(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(maior_ponta([1, 2, 3]), [3, 3, 3],"Esperado: [3,3,3]")
        self.assertEqual(maior_ponta([11, 5, 9]), [11, 11, 11],"Esperado: [11,11,11]")
        self.assertEqual((maior_ponta([2, 11, 3]), [3, 3, 3],"Esperado: [3,3,3]")
        self.assertEqual(maior_ponta([11, 3, 3]), [11, 11, 11],"Esperado: [11,11,11]")
        self.assertEqual(maior_ponta([3, 11, 11]), [11, 11, 11],"Esperado: [11,11,11]")
        self.assertEqual(maior_ponta([2, 2, 2]), [2, 2, 2],"Esperado: [2,2,2]"")
        self.assertEqual(maior_ponta([2, 11, 2]), [2, 2, 2],"Esperado: [2,2,2]")
        self.assertEqual(maior_ponta([0, 0, 1]), [1, 1, 1],"Esperado: [1,1,1]")


myTests().main()
</textarea>
</div>

<div data-childcomponent="sum2" class="runestone explainer ac_section alert alert-warning">
E. sum2 <br />
Dada uma lista de inteiros de qualquer tamanho <br />
retorna a soma dos dois primeiros elementos <br />
se a lista tiver menos de dois elementos, soma o que for possível <br />
<textarea data-component="activecode" id=sum2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def sum2(nums):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum2([1, 2, 3]), 3,"Esperado: 3")
        self.assertEqual(sum2([1, 1]), 2,"Esperado: 2")
        self.assertEqual(sum2([1, 1, 1, 1]), 2,"Esperado: 2")
        self.assertEqual(sum2([1, 2]), 3,"Esperado: 3")
        self.assertEqual(sum2([1]), 1,"Esperado: 1")
        self.assertEqual(sum2([]), 0,"Esperado: 0")
        self.assertEqual(sum2([4, 5, 6]), 9,"Esperado: 9")
        self.assertEqual(sum2([4]), 4,"Esperado: 4")


myTests().main()
</textarea>
</div>

<div data-childcomponent="middle_way" class="runestone explainer ac_section alert alert-warning">
F. middle_way <br />
sejam duas listas de inteiros a e b <br />
retorna uma lista de tamanho 2 contendo os elementos do <br />
meio de a e b, suponha que as listas tem tamanho ímpar <br />
middle_way([1, 2, 3], [4, 5, 6]) -&gt; [2, 5] <br />
middle_way([7, 7, 7], [3, 8, 0]) -&gt; [7, 8] <br />
middle_way([5, 2, 9], [1, 4, 5]) -&gt; [2, 4] <br />
<textarea data-component="activecode" id=middle_way data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def middle_way(a, b):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(middle_way([1, 2, 3], [4, 5, 6]), [2, 5],"Esperado: [2,5]")
        self.assertEqual(middle_way([7, 7, 7], [3, 8, 0]), [7, 8],"Esperado: [7,8]")
        self.assertEqual(middle_way([5, 2, 9], [1, 4, 5]), [2, 4],"Esperado: [2,4]")
        self.assertEqual(middle_way([1, 9, 7], [4, 8, 8]), [9, 8]),"Esperado: [9,8]")
        self.assertEqual(middle_way([1, 2, 3], [3, 1, 4]), [2, 1],"Esperado: [2,1]")
        self.assertEqual(middle_way([1, 2, 3], [4, 1, 1]), [2, 1],"Esperado: [2,1]")



myTests().main()
</textarea>
</div>

<div data-childcomponent="date_fashion" class="runestone explainer ac_section alert alert-warning">
G. date_fashion <br />
você e sua namorada(o) vão a um restaurante <br />
eu e par são as notas das suas roupas de 0 a 10 <br />
quanto maior a nota mais chique vocês estão vestidos <br />
o resultado é se vocês conseguiram uma mesa no restaurante: <br />
0=não 1=talvez e 2=sim <br />
se a nota da roupa de um dos dois for menor ou igual a 2 <br />
vocês não terão direito à uma mesa (0) <br />
se as notas são maiores, então caso um dos dois esteja <br />
bem chique (nota &gt;= 8) então a resposta é sim (2) <br />
caso contrário a resposta é talvez (1) <br />
date_fashion(5, 10) -&gt; 2 <br />
date_fashion(5, 2) -&gt; 0 <br />
date_fashion(5, 5) -&gt; 1 <br />
<textarea data-component="activecode" id=date_fashion data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def date_fashion(eu, par):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(date_fashion(5, 10), 2,"Esperado: 2")
        self.assertEqual(date_fashion(5, 2), 0,"Esperado: 0")
        self.assertEqual(date_fashion(5, 5), 1,"Esperado: 1")
        self.assertEqual(date_fashion(3, 3), 1,"Esperado: 1")
        self.assertEqual(date_fashion(10, 2), 0,"Esperado: 0")
        self.assertEqual(date_fashion(2, 9), 0,"Esperado: 0")
        self.assertEqual(date_fashion(9, 9), 2,"Esperado: 2")
        self.assertEqual(date_fashion(10, 5), 2,"Esperado: 2")
        self.assertEqual(date_fashion(2, 2), 0,"Esperado: 0")
        self.assertEqual(date_fashion(3, 7), 1,"Esperado: 1")
        self.assertEqual(date_fashion(2, 7), 0,"Esperado: 0")
        self.assertEqual(date_fashion(6, 2), 0,"Esperado: 0")


myTests().main()
</textarea>
</div>

<div data-childcomponent="squirrel_play" class="runestone explainer ac_section alert alert-warning">
H. squirrel_play <br />
os esquilos na FATEC brincam quando a temperatura está entre 60 e 90 <br />
graus Fahreneit (são estrangeiros e o termômetro é diferente rs) <br />
caso seja verão, então a temperatura superior é 100 no lugar de 90 <br />
retorne True caso os esquilos brinquem <br />
squirrel_play(70, False) -&gt; True <br />
squirrel_play(95, False) -&gt; False <br />
squirrel_play(95, True) -&gt; True <br />
<textarea data-component="activecode" id=squirrel_play data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def squirrel_play(temp, is_summer):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(squirrel_play(70, False), True,"Esperado: True")
        self.assertEqual(squirrel_play(95, False), False,"Esperado: False")
        self.assertEqual(squirrel_play(95, True), True,"Esperado: True")
        self.assertEqual(squirrel_play(90, False), True,"Esperado: True")
        self.assertEqual(squirrel_play(90, True), True,"Esperado: True")
        self.assertEqual(squirrel_play(50, False), False,"Esperado: False")
        self.assertEqual(squirrel_play(50, True), False,"Esperado: True")
        self.assertEqual(squirrel_play(100, False), False,"Esperado: False")
        self.assertEqual(squirrel_play(100, True), True,"Esperado: True")
        self.assertEqual(squirrel_play(105, True), False,"Esperado: True")
        self.assertEqual(squirrel_play(59, False), False,"Esperado: False")
        self.assertEqual(squirrel_play(59, True), False,"Esperado: True")
        self.assertEqual(squirrel_play(60, False), True,"Esperado: False")



myTests().main()
</textarea>
</div>

<div data-childcomponent="pego_correndo" class="runestone explainer ac_section alert alert-warning">
I. pego_correndo <br />
você foi pego correndo <br />
o resultado será: <br />
sem multa = 0 <br />
multa média = 1 <br />
multa grave = 2 <br />
velocide &lt;= 60 sem multa <br />
velocidade entre 61 e 80 multa média <br />
velocidade maior que 81 multa grave (cidade do interior) <br />
caso seja seu aniversário a velocidade pode ser 5 km/h maior em todos os casos <br />
pego_correndo(60, False) -&gt; 0 <br />
pego_correndo(65, False) -&gt; 1 <br />
pego_correndo(65, True) -&gt; 0  <br />
<textarea data-component="activecode" id=pego_correndo data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def pego_correndo(speed, is_birthday):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(pego_correndo(60, False), 0,"Esperado: 0")
        self.assertEqual(pego_correndo(65, False), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(65, True), 0,"Esperado: 0")
        self.assertEqual(pego_correndo(80, False), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(85, False), 2,"Esperado: 2")
        self.assertEqual(pego_correndo(85, True), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(70, False), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(75, False), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(75, True), 1,"Esperado: 1")
        self.assertEqual(pego_correndo(40, False), 0,"Esperado: 0")
        self.assertEqual(pego_correndo(40, True), 0,"Esperado: 0")
        self.assertEqual(pego_correndo(90, False), 2,"Esperado: 2")



myTests().main()
</textarea>
</div>

<div data-childcomponent="alarm_clock" class="runestone explainer ac_section alert alert-warning">
J. alarm_clock <br />
day: 0=domingo, 1=segunda, 2=terça, …, 6=sábado <br />
vacation = True caso você esteja de férias <br />
o retorno é uma string que diz quando o despertador tocará <br />
dias da semana ‘07:00’ <br />
finais de semana ‘10:00’ <br />
a menos que você esteja de férias, neste caso: <br />
dias da semana ‘10:00’ <br />
finais de semana ‘off’ <br />
alarm_clock(1, False) -&gt; ‘7:00’ <br />
alarm_clock(5, False) -&gt; ‘7:00’ <br />
alarm_clock(0, False) -&gt; ‘10:00’ <br />
<textarea data-component="activecode" id=alarm_clock data-lang="python" 
      data-timelimit=25000    
    data-audio=''     
          >
def alarm_clock(day, vacation):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(alarm_clock(1, False), '7:00',"Esperado: 7:00")
        self.assertEqual(alarm_clock(5, False), '7:00',"Esperado: 7:00")
        self.assertEqual(alarm_clock(0, False), '10:00',"Esperado: 10:00")
        self.assertEqual(alarm_clock(6, False), '10:00',"Esperado: 10:00")
        self.assertEqual(alarm_clock(0, True), 'off',"Esperado: off")
        self.assertEqual(alarm_clock(6, True), 'off',"Esperado: off")
        self.assertEqual(alarm_clock(1, True), '10:00',"Esperado: 10:00")
        self.assertEqual(alarm_clock(3, True), '10:00',"Esperado: 10:00")
        self.assertEqual(alarm_clock(5, True), '10:00',"Esperado: 10:00")



myTests().main()
</textarea>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 7' data-toggle="tooltip">
    <a href="Quiz7.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz7.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 10' data-toggle="tooltip" >
    <a href="Quiz10.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz10.html" >Next Section - Quiz - 10</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2017 Atharva Kulkarni.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.3.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>