<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Hangman Challenge &#8212; PyZombis localstorage@710f0aa</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=1654bbde73e12a3366eb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis localstorage@710f0aa
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Reto01_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="hangman-challenge">
<h1>Hangman Challenge<a class="headerlink" href="#hangman-challenge" title="Permalink to this headline">¶</a></h1>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 1)</p>
<p>Title overline too short.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">=============</span>
<span class="n">Hangman</span> <span class="n">Challenge</span>
<span class="o">=============</span>
</pre></div>
</div>
</div>
<p>This challenge consists in creating the <strong>hangman</strong> game from Python code.</p>
<div id='hangman' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >
<p>First, we’ll make the different hangman drawings with text strings.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 13)</p>
<p>Duplicate ID – see challenges/Reto01, line 17</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_1
    :nocodelens:

    In the ``drawings`` list we will place the text strings that represent
    the hangman drawings.

    ~~~~
    drawings = [
        &quot;&quot;&quot;
      +------+
             |
             |
             |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
             |
             |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
             |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
      |      |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
     /|      |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
     /|\\     |
             |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
     /|\\     |
     /       |
             |
    +--------+ &quot;&quot;&quot;,
        &quot;&quot;&quot;
      +------+
      |      |
      o      |
     /|\\     |
     / \\     |
             |
    +--------+ &quot;&quot;&quot;,
    ]
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >
<p>We will solve the challenge progressively.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 92)</p>
<p>Duplicate ID – see challenges/Reto01, line 96</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">ac_r01_2</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Create</span> <span class="n">two</span> <span class="n">variables</span><span class="p">,</span> <span class="n">one</span> <span class="n">called</span> <span class="s2">&quot;correct&quot;</span> <span class="ow">and</span> <span class="n">the</span> <span class="n">other</span> <span class="s2">&quot;incorrect&quot;</span><span class="o">.</span> <span class="n">Both</span> <span class="n">should</span> <span class="n">be</span> <span class="n">assigned</span> <span class="n">to</span>
    <span class="n">an</span> <span class="n">empty</span> <span class="n">string</span><span class="o">.</span>

    <span class="o">~~~~</span>

    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">correct</span><span class="p">,</span> <span class="s2">&quot;&quot;</span><span class="p">,</span> <span class="s2">&quot;Testing that correct is assigned correctly&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">incorrect</span><span class="p">,</span> <span class="s2">&quot;&quot;</span><span class="p">,</span> <span class="s2">&quot;Testing that incorrect is assigned correctly&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >
<p>Do you remember the lesson on APIs? Let’s apply what we learned from that lesson.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 118)</p>
<p>Duplicate ID – see challenges/Reto01, line 122</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_3
    :nocodelens:

    To play hangman we need words to guess. In this case we will retrieve those words from the University Domains API.
    These words will be world cities.
    Use the ``requests`` library to access the University Domains API. The url is the following
    http://universities.hipolabs.com/search
    Develop the ``choose`` function that will **randomly** choose a city within the cities of the
    ``cities`` list. This city will be returned by the function. It should be noted that the name of each city
    is given in English.

    ~~~~
    import requests
    import json
    from random import choice

    api_url = &quot;http://universities.hipolabs.com/search&quot;

    solicitation = requests.get(api_url)
    data = json.loads(solicitation.text)

    cities = []
    for university in data:
        if university[&quot;country&quot;] not in cities:
            cities.append(university[&quot;country&quot;].lower().replace(&quot; &quot;, &quot;&quot;))

    def choose(cities):
        # Develop the function
        # It can be achieved with a single line of code


    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(str(type(choose)), &quot;&lt;class &#39;function&#39;&gt;&quot;, &quot;Testing that choose has been defined&quot;)
            self.assertEqual(choose([&quot;Mexico&quot;]), &quot;Mexico&quot;, &quot;Testing that the function works&quot;)
            self.assertEqual(choice([2]), 2, &quot;Testing that choice has been imported&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >
<p>Let’s define another function to print the hangman.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 168)</p>
<p>Duplicate ID – see challenges/Reto01, line 173</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_4
    :nocodelens:
    :include: ac_r01_1, ac_r01_2, ac_r01_3

    The ``print_hangman`` function prints the corresponding hangman drawing according to
    the number of correct and incorrect letters so far. It takes a single parameter
    ``p_aleatoria`` which represents the word to guess.

    ~~~~
    p_aleatoria = choose(cities)
    def print_hangman(p_aleatoria):
        print(drawings[len(incorrect)])
        for c in p_aleatoria:
            print(c if c in correct else &quot;_&quot;, end=&quot; &quot;)
        print()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >
<p>It’s time to create the main function of our program, which asks the user for a
letter to guess.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 190)</p>
<p>Duplicate ID – see challenges/Reto01, line 195</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_5
    :nocodelens:

    Develop the ``guess`` function. It receives a string as a parameter (``letters``) with all the letters already
    tested (correct + incorrect). It returns a lowercase letter that has not been tested before.
    If the user enters more than one letter or another special character, a message should be printed
    alerting the user of their error. For example: ``&quot;Invalid character&quot;``
    This function will have a loop that will make sure that all the data is coherent and will only pass
    to the main program the lowercase letter that has not been already tested. Save what the user enters in the
    variable ``x``. **Note**: The function will be run once and will ask for a letter to verify if the unit tests pass.

    ~~~~
    from string import digits, punctuation
    import time

    def guess(letters):
        while True:
            x =      # Start here

            # All your code must go inside this while loop

            # Wait before asking for input again
            # Do not remove this line of code
            time.sleep(2)


    ====
    from unittest.gui import TestCaseGui
    import string


    class myTests(TestCaseGui):
        def testOne(self):
            x = guess(&quot;aeiou&quot;)
            self.assertEqual(x in string.ascii_lowercase, True, &quot;Testing that x has been returned correctly&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >
<p>Let’s make it possible for the user to decide whether or not they want to play again.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 234)</p>
<p>Duplicate ID – see challenges/Reto01, line 240</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_6
    :nocodelens:

    Develop the ``play_again`` function that asks the user if they want to play again. The answer should be *s* or *n* (S/N).
    Regardless of whether the user enters the response in uppercase or lowercase, the program should transform it to lowercase.
    The function returns ``True`` or ``False`` depending on the response. **Note**: The function will be run twice to check if it passes the unit tests.
    In the console, you will see instructions on what to enter for each test.

    ~~~~
    def play_again():

        return


    ====
    from unittest.gui import TestCaseGui
    import time


    class myTests(TestCaseGui):
        def testOne(self):
            print(&quot;Enter &#39;S&#39;&quot;)
            time.sleep(2)
            r = play_again()
            self.assertEqual(r, True, &quot;Testing that the lower method has been used&quot;)


        def testTwo(self):
            print(&quot;Enter &#39;n&#39;&quot;)
            time.sleep(2)
            r = play_again()
            self.assertEqual(r, False, &quot;Testing that it returns False when entering &#39;n&#39;&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 273)</p>
<p>Duplicate ID – see challenges/Reto01, line 280</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: ac_r01_7
    :nocodelens:

    Now develop the ``win`` function. It takes two parameters: ``p_aleatoria``, which represents the word to be guessed, and ``letras_adivinadas``. It should return ``True`` if all the letters of ``p_aleatoria`` are in the ``letras_adivinadas`` string. Otherwise, it returns False.

    ~~~~
    def win(p_aleatoria, letras_adivinadas):


    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):
            self.assertEqual(win(&quot;argentina&quot;, &quot;anitnegra&quot;), True, &quot;Testing function&quot;)
            self.assertEqual(win(&quot;argentina&quot;, &quot;argent&quot;), False, &quot;Testing function&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 8" >
<p>You have already defined all the necessary functions and variables for the program. If you have made it this far, you should have completed everything and passed all the unit tests. If so, you can run the following program to test your game. You have 5 minutes to play.</p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/challenges/Reto01_en.rst</span>, line 298)</p>
<p>Duplicate ID – see challenges/Reto01, line 309</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">ac_r01_8</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>
    <span class="p">:</span><span class="n">include</span><span class="p">:</span> <span class="n">ac_r01_1</span><span class="p">,</span> <span class="n">ac_r01_2</span><span class="p">,</span> <span class="n">ac_r01_3</span><span class="p">,</span> <span class="n">ac_r01_4</span><span class="p">,</span> <span class="n">ac_r01_5</span><span class="p">,</span> <span class="n">ac_r01_6</span><span class="p">,</span> <span class="n">ac_r01_7</span>

    <span class="c1"># This helps to increase the program duration to 5 minutes</span>
    <span class="kn">import</span> <span class="nn">sys</span>
    <span class="n">sys</span><span class="o">.</span><span class="n">setExecutionLimit</span><span class="p">(</span><span class="mi">300000</span><span class="p">)</span>

    <span class="n">p_aleatoria</span> <span class="o">=</span> <span class="n">choose</span><span class="p">(</span><span class="n">cities</span><span class="p">)</span>

    <span class="k">while</span> <span class="kc">True</span><span class="p">:</span>
        <span class="n">print_hangman</span><span class="p">(</span><span class="n">p_aleatoria</span><span class="p">)</span>
        <span class="n">time</span><span class="o">.</span><span class="n">sleep</span><span class="p">(</span><span class="mi">2</span><span class="p">)</span>
        <span class="n">x</span> <span class="o">=</span> <span class="n">guess</span><span class="p">(</span><span class="n">corrects</span> <span class="o">+</span> <span class="n">incorrects</span><span class="p">)</span>
        <span class="k">if</span> <span class="n">x</span> <span class="ow">in</span> <span class="n">p_aleatoria</span><span class="p">:</span>
            <span class="n">corrects</span> <span class="o">=</span> <span class="n">corrects</span> <span class="o">+</span> <span class="n">x</span>
        <span class="k">else</span><span class="p">:</span>
            <span class="n">incorrects</span> <span class="o">=</span> <span class="n">incorrects</span> <span class="o">+</span> <span class="n">x</span>
        <span class="k">if</span> <span class="nb">len</span><span class="p">(</span><span class="n">incorrects</span><span class="p">)</span> <span class="o">==</span> <span class="nb">len</span><span class="p">(</span><span class="n">drawings</span><span class="p">):</span>
            <span class="nb">print</span><span class="p">(</span><span class="sa">f</span><span class="s2">&quot;You were hanged, the word was </span><span class="si">{</span><span class="n">p_aleatoria</span><span class="si">}</span><span class="s2">&quot;</span><span class="p">)</span>
            <span class="n">time</span><span class="o">.</span><span class="n">sleep</span><span class="p">(</span><span class="mi">1</span><span class="p">)</span>
            <span class="k">if</span> <span class="n">play_again</span><span class="p">():</span>
                <span class="n">corrects</span> <span class="o">=</span> <span class="n">incorrects</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
                <span class="n">p_aleatoria</span> <span class="o">=</span> <span class="n">choose</span><span class="p">(</span><span class="n">cities</span><span class="p">)</span>
            <span class="k">else</span><span class="p">:</span>
                <span class="k">break</span>
        <span class="k">elif</span> <span class="n">win</span><span class="p">(</span><span class="n">p_aleatoria</span><span class="p">,</span> <span class="n">corrects</span><span class="p">):</span>
            <span class="nb">print</span><span class="p">(</span><span class="sa">f</span><span class="s2">&quot;You guessed the word: </span><span class="si">{</span><span class="n">p_aleatoria</span><span class="si">}</span><span class="s2">&quot;</span><span class="p">)</span>
            <span class="n">time</span><span class="o">.</span><span class="n">sleep</span><span class="p">(</span><span class="mi">1</span><span class="p">)</span>
            <span class="k">if</span> <span class="n">play_again</span><span class="p">():</span>
                <span class="n">corrects</span> <span class="o">=</span> <span class="n">incorrects</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
                <span class="n">p_aleatoria</span> <span class="o">=</span> <span class="n">choose</span><span class="p">(</span><span class="n">cities</span><span class="p">)</span>
            <span class="k">else</span><span class="p">:</span>
                <span class="k">break</span>

    <span class="o">====</span>
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>