<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Reto Ahorcado &#8212; PyZombis master@8d046a9</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.bundle.js?v=ef1af912ec3415fea854"></script>
    <script defer="" src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script defer="" src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Reto Pergamino" href="Reto02.html" />
    <link rel="prev" title="Quiz - Extras 2" href="../quiz/QuizExtras2.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@8d046a9</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Reto01'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Reto Ahorcado</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - Extras 2" data-toggle="tooltip">
    <a href="../quiz/QuizExtras2.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Reto Pergamino' data-toggle="tooltip" >              
    <a href="Reto02.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="reto-ahorcado">
<h1>Reto Ahorcado<a class="headerlink" href="#reto-ahorcado" title="Permalink to this headline">¶</a></h1>
<p>Este reto consiste en crear el juego de <strong>ahorcado</strong> a partir de código Python.</p>
<div id='ahorcado' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="ac_1" >
<p>Primero haremos los distintos dibujos del ahorcado con cadenas de texto.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_1 data-question_label="1">
<div id=ac_r01_1_question class="ac_question col-md-12">
<p>En la lista <code class="docutils literal notranslate"><span class="pre">dibujos</span></code> colocaremos las cadenas de texto que representan
los dibujos del ahorcado.</p>

</div>
<textarea data-lang="python" id="ac_r01_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
dibujos = [
    """
  +------+
         |
         |
         |
         |
         |
+--------+ """,
    """
  +------+
  |      |
         |
         |
         |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
         |
         |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
  |      |
         |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
 /|      |
         |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
 /|\     |
         |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
 /|\     |
 /       |
         |
+--------+ """,
    """
  +------+
  |      |
  o      |
 /|\     |
 / \     |
         |
+--------+ """,
]
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_2" >
<p>Va a resolver el reto de manera progresiva.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_2 data-question_label="2">
<div id=ac_r01_2_question class="ac_question col-md-12">
<p>Cree dos variables, una llamada <code class="docutils literal notranslate"><span class="pre">correctas</span></code> y la otra <code class="docutils literal notranslate"><span class="pre">incorrectas</span></code>. Ambas asígnelas a
una cadena vacía.</p>

</div>
<textarea data-lang="python" id="ac_r01_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(correctas, "", "Probando que correctas esté asignada correctamente")
        self.assertEqual(incorrectas, "", "Probando que incorrectas esté asignada correctamente")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_3" >
<p>¿Recuerda la lección de las APIs? Vamos a aplicar lo aprendido de esa lección.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_3 data-question_label="3">
<div id=ac_r01_3_question class="ac_question col-md-12">
<p>Para poder jugar ahorcado necesitamos palabras para adivinar. En este caso vamos a
recuperar esas palabras de la API de University Domains. Estas palabras serán ciudades
del mundo.
Con la librería <code class="docutils literal notranslate"><span class="pre">requests</span></code> acceda a la API de University Domains. La url es la siguiente
<a class="reference external" href="http://universities.hipolabs.com/search">http://universities.hipolabs.com/search</a>
Desarrolle la función <code class="docutils literal notranslate"><span class="pre">escoger</span></code> que escoja de manera <strong>aleatoria</strong> una ciudad dentro de las ciudades de la
lista <code class="docutils literal notranslate"><span class="pre">ciudades</span></code>. Esta ciudad será devuelta por la función. Cabe aclarar que el nombre de cada ciudad
está en inglés.</p>

</div>
<textarea data-lang="python" id="ac_r01_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
import requests
import json
from random import choice

api_url = "http://universities.hipolabs.com/search"

solicitud = requests.get(api_url)
datos = json.loads(solicitud.text)

ciudades = []
for universidad in datos:
    if universidad["country"] not in ciudades:
        ciudades.append(universidad["country"].lower().replace(" ", ""))

def escoger(ciudades):
    # Desarrolle la función
    # Se puede lograr con una sola línea de código


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(str(type(escoger)), "<class 'function'>", "Probando que escoger haya sido definida")
        self.assertEqual(escoger(["Mexico"]), "Mexico", "Probando que la función sirva")
        self.assertEqual(choice([2]), 2, "Probando que choice haya sido importado")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_4" >
<p>Ahora definamos otra función para imprimir el ahorcado.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_4 data-question_label="4">
<div id=ac_r01_4_question class="ac_question col-md-12">
<p>La función <code class="docutils literal notranslate"><span class="pre">imprimir_ahorcado</span></code> imprime el dibujo del ahorcado correspondiente al
número de letras incorrectas y correctas hasta el momento. Toma un único parámetro
<code class="docutils literal notranslate"><span class="pre">p_aleatoria</span></code> que representa la palabra a adivinar.</p>

</div>
<textarea data-lang="python" id="ac_r01_4_editor" 
     data-include="ac_r01_1 ac_r01_2 ac_r01_3" data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
p_aleatoria = escoger(ciudades)
def imprimir_ahorcado(p_aleatoria):
    print(dibujos[len(incorrectas)])
    for c in p_aleatoria:
        print(c if c in correctas else "_", end=" ")
    print()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_5" >
<p>Es momento de crear la función principal de nuestro programa, la cuál pide al usuario una
letra para adivinar.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_5 data-question_label="5">
<div id=ac_r01_5_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">adivinar</span></code>. Recibe como parámetro una cadena (<code class="docutils literal notranslate"><span class="pre">letras</span></code>) con todas las letras ya
probadas (correctas + incorrectas). Devuelve una letra minúscula que no se ha probado antes.
Si el usuario ingresa más de una letra u otro carácter especial, se debe imprimir un mensaje
alertando al usuario de su error. Por ejemplo: <code class="docutils literal notranslate"><span class="pre">&quot;Caracter</span> <span class="pre">inválido&quot;</span></code>
Esta función tendrá un bucle que hará que todos los datos sean coherentes y solo pasará
a tu programa principal la letra minúscula que no se probó antes. Guarde lo ingresado por el usuario en la
variable <code class="docutils literal notranslate"><span class="pre">x</span></code>. <strong>Nota</strong>: La función se va a ejecutar una vez y pedirá que se ingrese una letra
para verificar si pasa las pruebas unitarias.</p>

</div>
<textarea data-lang="python" id="ac_r01_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
from string import digits, punctuation
import time

def adivinar(letras):
    while True:
        x =      # Empiece aquí

        # Todo su código debe ir dentro de este ciclo while

        # Esperar antes de preguntar por un input de nuevo
        # No remover esta línea de código
        time.sleep(2)


====
from unittest.gui import TestCaseGui
import string


class myTests(TestCaseGui):
    def testOne(self):
        x = adivinar("aeiou")
        self.assertEqual(x in string.ascii_lowercase, True, "Probando que x haya sido devuelto correctamente")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_6" >
<p>Hagamos posible que el usuario decida si quiere volver a jugar o no.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_6 data-question_label="6">
<div id=ac_r01_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">jugar_nuevamente</span></code> que pregunte al usuario si quiere volver a
jugar. La respuesta debe ser una <em>s</em> o <em>n</em> (S/N). Sin importar si el usuario ingresa la respuesta
en mayúscula o minúscula, el programa debe transformarla a minúscula. La función devuelve <code class="docutils literal notranslate"><span class="pre">True</span></code> o <code class="docutils literal notranslate"><span class="pre">False</span></code>
dependiendo de la respuesta. <strong>Nota</strong>: La función se va a ejecutar dos veces para verificar si pasa las pruebas unitarias.
En la consola verá instrucciones sobre qué ingresar para cada prueba.</p>

</div>
<textarea data-lang="python" id="ac_r01_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def jugar_nuevamente():

    return


====
from unittest.gui import TestCaseGui
import time


class myTests(TestCaseGui):
    def testOne(self):
        print("Ingrese 'S'")
        time.sleep(2)
        r = jugar_nuevamente()
        self.assertEqual(r, True, "Probando que se haya utilizado el método lower")


    def testTwo(self):
        print("Ingrese 'n'")
        time.sleep(2)
        r = jugar_nuevamente()
        self.assertEqual(r, False, "Probando que al ingresar 'n' regrese False")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_7" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_7 data-question_label="7">
<div id=ac_r01_7_question class="ac_question col-md-12">
<p>Ahora desarrolle la función <code class="docutils literal notranslate"><span class="pre">ganar</span></code>. Ésta toma dos parámetros: <code class="docutils literal notranslate"><span class="pre">p_aleatoria</span></code> que representa
la palabra a adivinar, y <code class="docutils literal notranslate"><span class="pre">letras_adivinadas</span></code>. Debe devolver <code class="docutils literal notranslate"><span class="pre">True</span></code> si todas las letras de
<code class="docutils literal notranslate"><span class="pre">p_aleatoria</span></code> están en la cadena <code class="docutils literal notranslate"><span class="pre">letras_adivinadas</span></code>. De lo contrario, devuelve False.</p>

</div>
<textarea data-lang="python" id="ac_r01_7_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def ganar(p_aleatoria, letras_adivinadas):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(ganar("argentina", "anitnegra"), True, "Probando la función")
        self.assertEqual(ganar("argentina", "argent"), False, "Probando la función")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="ac_8" >
<p>Ya ha definido todas las funciones y variables necesarias para el programa. Si ha llegado hasta aquí,
debió haber cumplido con todo y haber pasado todas las pruebas unitarias. De ser así, puede correr el
siguiente programa para probar su juego. El tiempo que tiene para jugar son 5 minutos.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=ac_r01_8 data-question_label="8">
<div id=ac_r01_8_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="ac_r01_8_editor" 
     data-include="ac_r01_1 ac_r01_2 ac_r01_3 ac_r01_4 ac_r01_5 ac_r01_6 ac_r01_7" data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
# Esto nos sirve para aumentar la duración del programa a 5 minutos.
import sys
sys.setExecutionLimit(300000)

p_aleatoria = escoger(ciudades)

while True:
    imprimir_ahorcado(p_aleatoria)
    time.sleep(2)
    x = adivinar(correctas + incorrectas)
    if x in p_aleatoria:
        correctas = correctas + x
    else:
        incorrectas = incorrectas + x
    if len(incorrectas) == len(dibujos):
        print(f"Quedó ahorcado, la palabra era {p_aleatoria}")
        time.sleep(1)
        if jugar_nuevamente():
            correctas = incorrectas = ""
            p_aleatoria = escoger(ciudades)
        else:
            break
    elif ganar(p_aleatoria, correctas):
        print(f"Acertó la palabra: {p_aleatoria}")
        time.sleep(1)
        if jugar_nuevamente():
            correctas = incorrectas = ""
            p_aleatoria = escoger(ciudades)
        else:
            break

====
</textarea>
</div>
</div>
</div>
    </div>

</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - Extras 2' data-toggle="tooltip">
    <a href="../quiz/QuizExtras2.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../quiz/QuizExtras2.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Reto Pergamino' data-toggle="tooltip" >
    <a href="Reto02.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Reto02.html" >Next Section - Reto Pergamino</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>