<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 9 &#8212; PyZombis combined-poc@76690b4</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=e476ce06ce356667798d"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 10" href="Quiz10_en.html" />
    <link rel="prev" title="Quiz - 8" href="Quiz8_en.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis combined-poc@76690b4
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz9_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 8" data-toggle="tooltip">
    <a href="Quiz8_en.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 10' data-toggle="tooltip" >              
    <a href="Quiz10_en.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 8' data-toggle="tooltip">
  <a href="Quiz8_en.html" aria-label="Previous - Quiz - 8">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 10' data-toggle="tooltip" >
    <a href="Quiz10_en.html" aria-label="Next - Quiz - 10">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-9">
<h1>Quiz - 9<a class="headerlink" href="#quiz-9" title="Permalink to this headline">¶</a></h1>
<div id='quiz9' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_1_en data-question_label="1">
<div id=q9_1_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">near_ten</span></code> that receives an integer <code class="docutils literal notranslate"><span class="pre">n</span></code> and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> when <code class="docutils literal notranslate"><span class="pre">n</span></code> is a multiple of 10, or when <code class="docutils literal notranslate"><span class="pre">n</span></code> is one or two units away from it. Otherwise, return <code class="docutils literal notranslate"><span class="pre">False</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">near_ten(12)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">near_ten(17)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">near_ten(19)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_1_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def near_ten(n):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(near_ten(12), True, "Expected: True")
        self.assertEqual(near_ten(17), False, "Expected: False")
        self.assertEqual(near_ten(19), True, "Expected: True")
        self.assertEqual(near_ten(31), True, "Expected: True")
        self.assertEqual(near_ten(6), False, "Expected: False")
        self.assertEqual(near_ten(10), True, "Expected: True")
        self.assertEqual(near_ten(11), True, "Expected: True")
        self.assertEqual(near_ten(21), True, "Expected: True")
        self.assertEqual(near_ten(22), True, "Expected: True")
        self.assertEqual(near_ten(23), False, "Expected: False")
        self.assertEqual(near_ten(54), False, "Expected: False")
        self.assertEqual(near_ten(155), False, "Expected: False")
        self.assertEqual(near_ten(158), True, "Expected: True")
        self.assertEqual(near_ten(3), False, "Expected: False")
        self.assertEqual(near_ten(1), True, "Expected: True")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_2_en data-question_label="2">
<div id=q9_2_en_question class="ac_question col-md-12">
<p>Develop a function called “suma_loca” that receives three integers “a”, “b”, and “c”. The function must return the sum of the three numbers with the condition that if any number appears twice, it will not be counted in the sum. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(1,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">6</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(3,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(3,</span> <span class="pre">3,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_2_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def suma_loca(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_loca(1, 2, 3), 6, "Expected: 6")
        self.assertEqual(suma_loca(3, 2, 3), 2, "Expected: 2")
        self.assertEqual(suma_loca(3, 3, 3), 0, "Expected: 0")
        self.assertEqual(suma_loca(9, 2, 2), 9, "Expected: 9")
        self.assertEqual(suma_loca(2, 2, 9), 9, "Expected: 9")
        self.assertEqual(suma_loca(2, 9, 2), 9, "Expected: 9")
        self.assertEqual(suma_loca(2, 9, 3), 14, "Expected: 14")
        self.assertEqual(suma_loca(4, 2, 3), 9, "Expected: 9")
        self.assertEqual(suma_loca(1, 3, 1), 3, "Expected: 3")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_3_en data-question_label="3">
<div id=q9_3_en_question class="ac_question col-md-12">
<p>Develop the function called “suma_con_suerte” that receives three integers “a”, “b”, and “c”. The function must return the sum of the three numbers with the condition that if one of the numbers is <strong>13</strong>, it won’t count in the sum, as well as all the numbers to its right. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">6</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">2,</span> <span class="pre">13)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">13,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_3_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def suma_con_suerte(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_con_suerte(1, 2, 3), 6, "Expected: 6")
        self.assertEqual(suma_con_suerte(1, 2, 13), 3, "Expected: 3")
        self.assertEqual(suma_con_suerte(1, 13, 3), 1, "Expected: 1")
        self.assertEqual(suma_con_suerte(1, 13, 13), 1, "Expected: 1")
        self.assertEqual(suma_con_suerte(6, 5, 2), 13, "Expected: 13")
        self.assertEqual(suma_con_suerte(13, 2, 3), 0, "Expected: 0")
        self.assertEqual(suma_con_suerte(13, 2, 13), 0, "Expected: 0")
        self.assertEqual(suma_con_suerte(13, 13, 2), 0, "Expected: 0")
        self.assertEqual(suma_con_suerte(9, 4, 13), 13, "Expected: 13")
        self.assertEqual(suma_con_suerte(8, 13, 2), 8, "Expected: 8")
        self.assertEqual(suma_con_suerte(7, 2, 1), 10, "Expected: 10")
        self.assertEqual(suma_con_suerte(3, 3, 13), 6, "Expected: 6")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_4_en data-question_label="4">
<div id=q9_4_en_question class="ac_question col-md-12">
<p>Develop the function called “caracteres_dobles” that receives a string “s” and should return the original string duplicated. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('The')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;TThhee&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('AAbb')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;AAAAbbbb&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('Hi-There')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HHii--TThheerree&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_4_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def caracteres_dobles(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(caracteres_dobles("The"), "TThhee", "Expected: TThhee")
        self.assertEqual(caracteres_dobles("AAbb"), "AAAAbbbb", "Expected: AAAAbbbb")
        self.assertEqual(caracteres_dobles("Hi-There"), "HHii--TThheerree", "Expected: HHii--TThheerree")
        self.assertEqual(caracteres_dobles("Word!"), "WWoorrdd!!", "Expected: WWoorrdd!!")
        self.assertEqual(caracteres_dobles("!!"), "!!!!", "Expected: !!!!")
        self.assertEqual(caracteres_dobles(""), "", "Expected: ")
        self.assertEqual(caracteres_dobles("a"), "aa", "Expected: aa")
        self.assertEqual(caracteres_dobles("."), "..", "Expected: ..")
        self.assertEqual(caracteres_dobles("aa"), "aaaa", "Expected: aaaa")

myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_5_en data-question_label="5">
<div id=q9_5_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">count_greeting</span></code> which receives a string <code class="docutils literal notranslate"><span class="pre">s</span></code> as a parameter and returns the number of times the string <code class="docutils literal notranslate"><span class="pre">&quot;hello&quot;</span></code> appears in <code class="docutils literal notranslate"><span class="pre">s</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">count_greeting</span> <span class="pre">(&quot;abc</span> <span class="pre">hello</span> <span class="pre">ho&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_greeting</span> <span class="pre">(&quot;ABChello</span> <span class="pre">hello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_greeting</span> <span class="pre">(&quot;hellohello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_5_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def count_greeting(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(count_greeting("abc hello ho"), 1, "Expected: abc hello ho")
        self.assertEqual(count_greeting("ABChello hello"), 2, "Expected: ABCh i hello")
        self.assertEqual(count_greeting("hellohello"), 2, "Expected: hellohello")
        self.assertEqual(count_greeting("helloHoLAhoLa"), 1, "Expected: helloHOLAhoLa")
        self.assertEqual(count_greeting(""), 0, "Expected: ")
        self.assertEqual(count_greeting("h"), 0, "Expected: h")
        self.assertEqual(count_greeting("hello"), 1, "Expected: hello")
        self.assertEqual(count_greeting("Hello is no HOLA on aHolA"),0, "Expected: Hello is no HOLA on aHolA")
        self.assertEqual(count_greeting("helloh not HOholaholA"), 2, "Expected: helloh not HOholaholA")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_6_en data-question_label="6">
<div id=q9_6_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">cat_dog</span></code> which receives a string <code class="docutils literal notranslate"><span class="pre">s</span></code> as a parameter and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the string <code class="docutils literal notranslate"><span class="pre">&quot;cat&quot;</span></code> and the string <code class="docutils literal notranslate"><span class="pre">&quot;dog&quot;</span></code> appear the same number of times in the string <code class="docutils literal notranslate"><span class="pre">s</span></code> or returns <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. | br | | br |
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">cat_dog</span> <span class="pre">(&quot;catdog&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cat_dog</span> <span class="pre">(&quot;catcat&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cat_dog</span> <span class="pre">(&quot;1cat1dog&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_6_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def cat_dog(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(cat_dog("catdog"), True, "Expected: True")
        self.assertEqual(cat_dog("catcat"), False, "Expected: False")
        self.assertEqual(cat_dog("1cat1cadoperro"), True, "Expected: True")
        self.assertEqual(cat_dog("catxxperroxxxperro"), False, "Expected: False")
        self.assertEqual(cat_dog("catxperroxperroxcadogcat"), True, "Expected: True")
        self.assertEqual(cat_dog("catxperroxperroxca"), False, "Expected: False")
        self.assertEqual(cat_dog("perroperrocat"), False, "Expected: False")
        self.assertEqual(cat_dog("perroogcat"), True, "Expected: True")
        self.assertEqual(cat_dog("perro"), False, "Expected: False")
        self.assertEqual(cat_dog("cat"), False, "Expected: False")
        self.assertEqual(cat_dog("ca"), True, "Expected: True")
        self.assertEqual(cat_dog("c"), True, "Expected: True")
        self.assertEqual(cat_dog(""), True, "Expected: True")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_7_en data-question_label="7">
<div id=q9_7_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">count_code</span></code> which receives a string <code class="docutils literal notranslate"><span class="pre">s</span></code> as a parameter and returns the number of times the string <code class="docutils literal notranslate"><span class="pre">&quot;code&quot;</span></code> appears in <code class="docutils literal notranslate"><span class="pre">s</span></code>, but with the condition that the letter <code class="docutils literal notranslate"><span class="pre">'d'</span></code> can be exchanged for any other letter. Thus, the strings <code class="docutils literal notranslate"><span class="pre">&quot;coze&quot;</span></code> and <code class="docutils literal notranslate"><span class="pre">&quot;coze&quot;</span></code> should be counted equally. <br /> <br />
Examples:<br />
<code class="docutils literal notranslate"><span class="pre">count_code(&quot;aaacodebbb&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_code(&quot;codexxcode&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_code(&quot;cozexxcope&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_7_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def count_code(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(count_code("aaacodebbb"), 1, "Expected: 1")
        self.assertEqual(count_code("codexxcode"), 2, "Expected: 2")
        self.assertEqual(count_code("cozexxcope"), 2, "Expected: 2")
        self.assertEqual(count_code("cozfxxcope"), 1, "Expected: 1")
        self.assertEqual(count_code("xxcozeyycop"), 1, "Expected: 1")
        self.assertEqual(count_code("cozcop"), 0, "Expected: 0")
        self.assertEqual(count_code("abcxyz"), 0, "Expected: 0")
        self.assertEqual(count_code("code"), 1, "Expected: 1")
        self.assertEqual(count_code("ode"), 0, "Expected: 0")
        self.assertEqual(count_code("c"), 0, "Expected: 0")
        self.assertEqual(count_code(""), 0, "Expected: 0")
        self.assertEqual(count_code("AAcodeBBcoleCCccoreDD"), 3, "Expected: 3")
        self.assertEqual(count_code("AAcodeBBcoleCCccorfDD"), 2, "Expected: 2")
        self.assertEqual(count_code("coAcodeBcoleccoreDD"), 3, "Expected: 3")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 8" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_8_en data-question_label="8">
<div id=q9_8_en_question class="ac_question col-md-12">
<p>Knowing that the function <code class="docutils literal notranslate"><span class="pre">lower()</span></code> converts the characters of a string to lowercase. Develop a function called <code class="docutils literal notranslate"><span class="pre">end_equals</span></code> that receives two strings <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code> as parameters. The function will return <code class="docutils literal notranslate"><span class="pre">True</span></code> if the string <code class="docutils literal notranslate"><span class="pre">b</span></code> is found at the end of string <code class="docutils literal notranslate"><span class="pre">a</span></code> or if the string <code class="docutils literal notranslate"><span class="pre">a</span></code> is found at the end of string <code class="docutils literal notranslate"><span class="pre">b</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">end_equals(&quot;Hiabc&quot;,</span> <span class="pre">&quot;abc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">end_equals(&quot;AbC&quot;,</span> <span class="pre">&quot;HiaBc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">end_equals(&quot;abc&quot;,</span> <span class="pre">&quot;abXabc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_8_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def end_equals(a, b):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(end_equals("Hiabc", "abc"), True, "Expected: True")
        self.assertEqual(end_equals("AbC", "HiaBc"), True, "Expected: True")
        self.assertEqual(end_equals("abc", "abXabc"), True, "Expected: True")
        self.assertEqual(end_equals("Hiabc", "abcd"), False, "Expected: False")
        self.assertEqual(end_equals("Hiabc", "bc"), True, "Expected: True")
        self.assertEqual(end_equals("Hiabcx", "bc"), False, "Expected: False")
        self.assertEqual(end_equals("abc", "abc"), True, "Expected: True")
        self.assertEqual(end_equals("xyz", "12xyz"), True, "Expected: True")
        self.assertEqual(end_equals("yz", "12xz"), False, "Expected: False")
        self.assertEqual(end_equals("Z", "12xz"), True, "Expected: True")
        self.assertEqual(end_equals("12", "12"), True, "Expected: True")
        self.assertEqual(end_equals("abcXYZ", "abcDEF"), False, "Expected: False")
        self.assertEqual(end_equals("ab", "ab12"), False, "Expected: False")
        self.assertEqual(end_equals("ab", "12ab"), True, "Expected: True")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 9" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_9_en data-question_label="9">
<div id=q9_9_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">count_pairs</span></code> that receives a list <code class="docutils literal notranslate"><span class="pre">numbers</span></code> of integers and returns the number of even numbers in the list. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">count_pairs([2,</span> <span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_pairs([2,</span> <span class="pre">2,</span> <span class="pre">0])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">count_pairs([1,</span> <span class="pre">3,</span> <span class="pre">5])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_9_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def count_pairs(numbers):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(count_pairs([2, 1, 2, 3, 4]), 3, "Expected: 3")
        self.assertEqual(count_pairs([2, 2, 0]), 3, "Expected: 3")
        self.assertEqual(count_pairs([1, 3, 5]), 0, "Expected: 0")
        self.assertEqual(count_pairs([]), 0, "Expected: 0")
        self.assertEqual(count_pairs([11, 9, 0, 1]), 1, "Expected: 1")
        self.assertEqual(count_pairs([2, 11, 9, 0]), 2, "Expected: 2")
        self.assertEqual(count_pairs([2]), 1, "Expected: 1")
        self.assertEqual(count_pairs([2, 5, 12]), 2, "Expected: 2")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 10" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_10_en data-question_label="10">
<div id=q9_10_en_question class="ac_question col-md-12">
<p>Develop the function `` suma_con_mas_suerte ‘’ that receives a list `` numbers ‘’ of integers and returns the sum of the numbers in the list with the condition that if one of the numbers is ** 13 ** , this is not counted in the sum as well as all numbers that are to its right. | br | | br |
Examples: | br |
`` suma_con_mas_suerte ([1, 2, 2, 1]) ‘’ -&gt; `` 6 ‘’ | br |
`` suma_con_mas_suerte ([1, 1]) ‘’ -&gt; `` 2 ‘’ | br |
`` suma_con_mas_suerte ([1, 2, 13, 1, 13]) ‘’ -&gt; `` 3 ‘’ | br |
`` suma_con_mas_suerte ([13, 1, 2, 3, 4]) ‘’ -&gt; `` 0 ‘’ | br |</p>

</div>
<textarea data-lang="python" id="q9_10_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def suma_con_mas_suerte (numbers):


====
from unittest.gui import TestCaseGui


class myTests (TestCaseGui):
    def testOne (self):
        self.assertEqual (suma_con_mas_suerte ([1, 2, 2, 1]), 6, "Expected: 6")
        self.assertEqual (suma_con_mas_suerte ([1, 1]), 2, "Expected: 2")
        self.assertEqual (suma_con_mas_suerte ([1, 2, 13, 1, 13]), 6, "Expected: 6")
        self.assertEqual (suma_con_mas_suerte ([1, 2, 13, 2, 1, 13]), 3, "Expected: 3")
        self.assertEqual (suma_con_mas_suerte ([13, 1, 2, 13, 2, 1, 13]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([0, 1, 2, 13, 2, 1, 13]), 3, "Expected: 3")
        self.assertEqual (suma_con_mas_suerte ([]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([13]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([0, 5, 1, 2, 13, 100]), 8, "Expected: 8")
        self.assertEqual (suma_con_mas_suerte ([13, 13]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([13, 0, 13]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([13, 1, 13]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([5, 7, 2]), 14, "Expected: 14")
        self.assertEqual (suma_con_mas_suerte ([5, 13, 2]), 5, "Expected: 5")
        self.assertEqual (suma_con_mas_suerte ([0]), 0, "Expected: 0")
        self.assertEqual (suma_con_mas_suerte ([13, 0]), 0, "Expected: 0")


myTests (). main ()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 11" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_11_en data-question_label="11">
<div id=q9_11_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">tiene_2</span></code> that receives a list <code class="docutils literal notranslate"><span class="pre">numeros</span></code> of integers as parameter and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if there are two consecutive equal numbers in the list, and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([1,</span> <span class="pre">2,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([1,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([2,</span> <span class="pre">1,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_11_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def tiene_2(numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(tiene_2([1, 2, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([1, 2, 1, 2]), False, "Expected: False")
        self.assertEqual(tiene_2([2, 1, 2]), False, "Expected: False")
        self.assertEqual(tiene_2([2, 2, 1, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([1, 3, 2]), False, "Expected: False")
        self.assertEqual(tiene_2([1, 3, 2, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([2, 3, 2, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([4, 2, 4, 2, 2, 5]), True, "Expected: True")
        self.assertEqual(tiene_2([1, 2]), False, "Expected: False")
        self.assertEqual(tiene_2([2, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([2]), False, "Expected: False")
        self.assertEqual(tiene_2([]), False, "Expected: False")
        self.assertEqual(tiene_2([3, 3, 2, 2]), True, "Expected: True")
        self.assertEqual(tiene_2([5, 2, 5, 2]), False, "Expected: False")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 12" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_12_en data-question_label="12">
<div id=q9_12_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">suma_en_lista</span></code> that receives two parameters, an integer <code class="docutils literal notranslate"><span class="pre">n</span></code> and a list <code class="docutils literal notranslate"><span class="pre">numeros</span></code> of numbers. The function should return <code class="docutils literal notranslate"><span class="pre">True</span></code> if there are two different elements in the list whose sum is <code class="docutils literal notranslate"><span class="pre">n</span></code>, and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(5,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(9,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(0,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(8,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(4,</span> <span class="pre">[2,</span> <span class="pre">2,</span> <span class="pre">2,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(4,</span> <span class="pre">[2,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">3])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q9_12_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def suma_en_lista(n, numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_en_lista(5, [1, 2, 3, 4]), True, "Expected: True")
        self.assertEqual(suma_en_lista(9, [1, 2, 3, 4]), False, "Expected: False")
        self.assertEqual(suma_en_lista(0, [1, 2, 3, 4]), False, "Expected: False")
        self.assertEqual(suma_en_lista(8, [1, 2, 3, 4]), False, "Expected: False")
        self.assertEqual(suma_en_lista(4, [2, 2, 2, 2]), False, "Expected: False")
        self.assertEqual(suma_en_lista(4, [2, 2, 1, 3]), True, "Expected: True")
        self.assertEqual(suma_en_lista(42, [40, 2, 3, 39]), True, "Expected: True")
        self.assertEqual(suma_en_lista(10, [5, 5, 4, 6]), True, "Expected: True")
        self.assertEqual(suma_en_lista(8, [5, 5, 4, 4]), False, "Expected: False")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 13" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q9_13_en data-question_label="13">
<div id=q9_13_en_question class="ac_question col-md-12">
<p>In a construction it is desired to build a row of bricks. There are small bricks and large bricks. Small bricks are 1 unit in length while large bricks are 5 units in length. Develop the function <code class="docutils literal notranslate"><span class="pre">build_row</span></code> which receives three parameters, <code class="docutils literal notranslate"><span class="pre">num_small_bricks</span></code>, <code class="docutils literal notranslate"><span class="pre">num_large_bricks</span></code>, and <code class="docutils literal notranslate"><span class="pre">row_length</span></code>. These parameters are integer numbers that represent the quantity of small bricks, quantity of large bricks, and the length of the row you want to build, respectively. The function must return <code class="docutils literal notranslate"><span class="pre">True</span></code> if it is possible to build the row, or <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">build_row(3,</span> <span class="pre">1,</span> <span class="pre">8)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<em>Explanation</em>: To build a length of 8 meters, 1 large brick and 3 small bricks would be used. <br />
<em>Therefore</em>: 1 + 1 + 1 + 5 = 8 <br /></p>
<p><code class="docutils literal notranslate"><span class="pre">build_row(3,</span> <span class="pre">1,</span> <span class="pre">9)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<em>Explanation</em>: To build a length of 9 meters it wouldn’t be possible because there are only 3 small bricks (1 meter in length) and 1 large brick (5 meters in length). <br />
<em>Therefore</em>: 1 + 1 + 1 + 5 &lt; 9 in this case 1 large brick and 4 small bricks would be needed. <br /></p>
<p><code class="docutils literal notranslate"><span class="pre">build_row(3,</span> <span class="pre">2,</span> <span class="pre">10)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<em>Explanation</em>: To fill a length of 10 meters, only the 2 available large bricks would be used. <br />
<em>Therefore</em>: 5 + 5 = 10 <br /></p>

</div>
<textarea data-lang="python" id="q9_13_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def build_row(num_small_bricks, num_large_bricks, row_length):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(build_row(3, 1, 8), True, "Expected: True")
        self.assertEqual(build_row(3, 1, 9), False, "Expected: False")
        self.assertEqual(build_row(3, 2, 10), True, "Expected: True")
        self.assertEqual(build_row(3, 2, 8), True, "Expected: True")
        self.assertEqual(build_row(3, 2, 9), False, "Expected: False")
        self.assertEqual(build_row(6, 1, 11), True, "Expected: True")
        self.assertEqual(build_row(6, 0, 11), False, "Expected: False")
        self.assertEqual(build_row(3, 1, 7), True, "Expected: True")
        self.assertEqual(build_row(1, 1, 7), False, "Expected: False")
        self.assertEqual(build_row(2, 1, 7), True, "Expected: True")
        self.assertEqual(build_row(7, 1, 11), True, "Expected: True")
        self.assertEqual(build_row(7, 1, 8), True, "Expected: True")
        self.assertEqual(build_row(7, 1, 13), False, "Expected: False")
        self.assertEqual(build_row(43, 1, 46), True, "Expected: True")
        self.assertEqual(build_row(40, 1, 46), False, "Expected: False")
        self.assertEqual(build_row(22, 2, 33), False, "Expected: False")
        self.assertEqual(build_row(0, 2, 10), True, "Expected: True")
        self.assertEqual(build_row(1000000, 1000, 1000100), True, "Expected: True")
        self.assertEqual(build_row(2, 1000000, 100003), False, "Expected: False")
        self.assertEqual(build_row(12, 2, 21), True, "Expected: True")


myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 8' data-toggle="tooltip">
  <a href="Quiz8_en.html" aria-label="Previous - Quiz - 8">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 10' data-toggle="tooltip" >
    <a href="Quiz10_en.html" aria-label="Next - Quiz - 10">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>