<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta charset="utf-8" />
    <title>Quiz - 1 &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css?v=239FC54E" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css?v=8A1A42E4" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css?v=C6382F3" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css?v=3BA0E63E" />
    <link rel="stylesheet" type="text/css" href="../_static/reveal.css?v=4A81D381" />
    <link rel="stylesheet" type="text/css" href="../_static/spreadsheet.css?v=B5F018A" />
    <link rel="stylesheet" type="text/css" href="../_static/jexcel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/japp.css" />
    <link rel="stylesheet" type="text/css" href="../_static/shortanswer.css?v=7E18C432" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css?v=F471F95B" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css?v=9CF75563" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css?v=58C4E18A" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css?v=E220AC33" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css?v=E5385B3B" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=DB047029" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css?v=424F1CC9" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css?v=2F3E13EE" />
    <link rel="stylesheet" type="text/css" href="../_static/js_lib/prettify.css?v=2A364B6B" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css?v=1FF7F927" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css?v=C184EB8A" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/d3.v2.min.js?v=836E01C"></script>
    <script src="../_static/jquery.ba-bbq.min.js?v=B64CF0A8"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js?v=42F9EBEA"></script>
    <script src="../_static/pytutor.js?v=82B022B9"></script>
    <script src="../_static/codelens.js?v=CABECD49"></script>
    <script src="../_static/showEval.js?v=5B2DD7A6"></script>
    <script src="../_static/fitb.js?v=60630F3D"></script>
    <script src="../_static/timedfitb.js?v=2AA17071"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js?v=599B42B0"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js?v=88A1430F"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js?v=AE15A4A1"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js?v=8ABDFE7F"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js?v=747279FF"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js?v=7B020E13"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js?v=660E85C3"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js?v=54B21948"></script>
    <script src="../_static/fitb-i18n.en.js?v=389ED00C"></script>
    <script src="../_static/reveal.js?v=35AF34F1"></script>
    <script src="../_static/spreadsheet.js?v=ABDABEC0"></script>
    <script src="../_static/jexcel.js"></script>
    <script src="../_static/japp.js"></script>
    <script src="../_static/shortanswer.js?v=89F557ED"></script>
    <script src="../_static/timed_shortanswer.js?v=B3A117D7"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/bookfuncs.js?v=24766D61"></script>
    <script src="../_static/codemirror.js?v=25B54D0A"></script>
    <script src="../_static/xml.js?v=236EDFF7"></script>
    <script src="../_static/css.js?v=623153CE"></script>
    <script src="../_static/python.js?v=76B9CFA2"></script>
    <script src="../_static/htmlmixed.js?v=4EBE05BF"></script>
    <script src="../_static/javascript.js?v=6F38CA04"></script>
    <script src="../_static/sql.js?v=49D0BD74"></script>
    <script src="../_static/activecode-i18n.en.js?v=65CA0724"></script>
    <script src="../_static/skulpt.min.js?v=A901CC2"></script>
    <script src="../_static/skulpt-stdlib.js?v=24904821"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/activecode.js?v=AF9F4792"></script>
    <script src="../_static/clike.js?v=4247C891"></script>
    <script src="../_static/timed_activecode.js?v=20BFB7B1"></script>
    <script src="../_static/sql-wasm.js?v=44FE0709"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/mchoice.js?v=7B685250"></script>
    <script src="../_static/timedmc.js?v=B010A457"></script>
    <script src="../_static/timed.js?v=7A76C128"></script>
    <script src="../_static/mchoice-i18n.en.js?v=354E6CD"></script>
    <script src="../_static/clickable.js?v=846DC076"></script>
    <script src="../_static/timedclickable.js?v=93989C23"></script>
    <script src="../_static/poll.js?v=45F04365"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="../_static/runestonevideo.js?v=53E888DE"></script>
    <script src="../_static/tabbedstuff.js?v=5BE6106"></script>
    <script src="../_static/js_lib/prettify.js?v=CA7E8928"></script>
    <script src="../_static/js_lib/hammer.min.js?v=73056367"></script>
    <script src="../_static/parsons.js?v=CE79120E"></script>
    <script src="../_static/timedparsons.js?v=2EEA9E5A"></script>
    <script src="../_static/parsons-i18n.en.js?v=A9C11155"></script>
    <script src="../_static/dragndrop.js?v=86D99604"></script>
    <script src="../_static/timeddnd.js?v=378F237E"></script>
    <script src="../_static/dragndrop-i18n.en.js?v=40C28531"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/lp.js?v=5E742BDC"></script>
    <script src="../_static/gas.js?v=129CE036"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/datafile.js?v=156D8707"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 3" href="Quiz3.html" />
    <link rel="prev" title="Quiz - 6" href="../lectures/Quiz6.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '4.2.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis Python Course</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz1'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 1</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 6" data-toggle="tooltip">
    <a href="../lectures/Quiz6.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 3' data-toggle="tooltip" >              
    <a href="Quiz3.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-1">
<h1>Quiz - 1<a class="headerlink" href="#quiz-1" title="Permalink to this headline">¶</a></h1>

<div data-childcomponent="q1_1" class="runestone explainer ac_section alert alert-warning">
<p>Haga un programa que pida dos números enteros e imprima la suma de esos dos números. <br /></p>

<textarea data-component="activecode" id=q1_1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def suma(n,m):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(suma(24,42), 66,"Esperado: 66")
        self.assertEqual(suma(17,13), 30,"Esperado: 30")
        self.assertEqual(suma(-11,6), -5,"Esperado: -5")
        self.assertEqual(suma(0,9), 9,"Esperado: 9")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_2" class="runestone explainer ac_section alert alert-warning">
<p>Escriba un programa que lea un valor en metros y lo muestre convertido a milimetros. <br /></p>

<textarea data-component="activecode" id=q1_2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def metros_a_milimetros(n):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(metros_a_milimetros(1), 1000,"Esperado: 1000")
        self.assertEqual(metros_a_milimetros(0.2), 200,"Esperado: 200")
        self.assertEqual(metros_a_milimetros(30), 30000,"Esperado: 30000")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_3" class="runestone explainer ac_section alert alert-warning">
<p>Escriba un programa que lea el número de días, horas, minutos y segundos del usuario. Calcular el total en segundos. <br /></p>

<textarea data-component="activecode" id=q1_3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def tiempo_en_segundos(dias, horas, minutos, segundos):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(tiempo_en_segundos(2,5,2,5), 190925,"Esperado: 190925")
        self.assertEqual(tiempo_en_segundos(10,89,5,0), 1184700,"Esperado: 1184700")
        self.assertEqual(tiempo_en_segundos(8,0,2,0), 691320,"Esperado: 691320")
        self.assertEqual(tiempo_en_segundos(0,5,55,6), 21306,"Esperado: 21306")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_4" class="runestone explainer ac_section alert alert-warning">
<p>Haz un programa que calcule un aumento de salario. Debe solicitar el monto del salario y el porcentaje del aumento. Muestra el monto del aumento y el nuevo salario. <br /></p>

<textarea data-component="activecode" id=q1_4 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def aumento(salario, porcentaje):
   #devolver los valores en una tupla como: return (aumento, nuevo_salario)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(aumento(30500,10),(3050,33550),"Esperado: 3050 y 33550")
        self.assertEqual(aumento(10400,25), (2600,13000),"Esperado: 2600 y 13000")
        self.assertEqual(aumento(50100,8), (4008,54108),"Esperado: 4008 y 54108")
        self.assertEqual(aumento(25000,3), (750,25750),"Esperado: 750 y 25750")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_5" class="runestone explainer ac_section alert alert-warning">
<p>Solicita el precio de un comerciante y el porcentaje de descuento. Muestre el monto del descuento y el precio a pagar. <br /></p>

<textarea data-component="activecode" id=q1_5 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def precio_con_descuento(precio, porcentaje):
   #devolver los valores en una tupla como: return (descuento, precio_final)


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(precio_con_descuento(100100,10),(10010,90090),"Esperado: (10010,90090)")
        self.assertEqual(precio_con_descuento(20523,4),(820.92,19702.08),"Esperado: (820.92,19702.08)")
        self.assertEqual(precio_con_descuento(55566,50),(27783,27783),"Esperado: (27783,27783)")
        self.assertEqual(precio_con_descuento(75660,24),(18158.4,57501.6),"Esperado: (18158.4,57501.6)")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_6" class="runestone explainer ac_section alert alert-warning">
<p>Calcula el tiempo de un viaje en auto. Pregunte por la distancia a recorrer y la velocidad media esperada para el viaje. <br /></p>

<textarea data-component="activecode" id=q1_6 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def tiempo(distancia,velocidad):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(tiempo(5,5),1,"Esperado: 1")
        self.assertEqual(tiempo(100,3), 100/3 ,"Esperado: "+ str(100/3)")
        self.assertEqual(tiempo(10500,30),350,"Esperado: 350")
        self.assertEqual(tiempo(8600,50), 172,"Esperado: 172")
        self.assertEqual(tiempo(130,200), 0.65,"Esperado: 0.65")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_7" class="runestone explainer ac_section alert alert-warning">
<p>Convierta una temperatura escrita en Celsius a Fahrenheit. F = (9 * C) / 5 + 32 <br /></p>

<textarea data-component="activecode" id=q1_7 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def celsius_a_fahrenheit(c):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(celsius_a_fahrenheit(20),((9 * 20) / 5 )+ 32,"Esperado: 68")
        self.assertEqual(celsius_a_fahrenheit(68),((9 * 68) / 5 )+ 32 ,"Esperado: 154.4")
        self.assertEqual(celsius_a_fahrenheit(0), ((9 * 0) / 5 )+ 32,"Esperado: 32")
        self.assertEqual(celsius_a_fahrenheit(-10), ((9 * -10) / 5 )+ 32,"Esperado: 14")
        self.assertEqual(celsius_a_fahrenheit(-24), ((9 * -24) / 5 )+ 32,"Esperado: -11.2")


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_8" class="runestone explainer ac_section alert alert-warning">
<p>Ahora haz lo contrario, de Fahrenheit a Celsius. <br /></p>

<textarea data-component="activecode" id=q1_8 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def fahrenheit_a_celsius(f):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(fahrenheit_a_celsius(21),((21-32)*5)/9,"Esperado: "+str(((21-32)*5)/9))
        self.assertEqual(fahrenheit_a_celsius(108),((108-32)*5)/9 ,"Esperado: "+str(((108-32)*5)/9))
        self.assertEqual(fahrenheit_a_celsius(0), ((0-32)*5)/9,"Esperado: "+str(((0-32)*5)/9))
        self.assertEqual(fahrenheit_a_celsius(-10), ((-10-32)*5)/9,"Esperado "+str(((-10-32)*5)/9))
        self.assertEqual(fahrenheit_a_celsius(14), ((14-32)*5)/9,"Esperado: "+str(((14-32)*5)/9))


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_9" class="runestone explainer ac_section alert alert-warning">
<p>Escribe un programa que pregunte por la cantidad de kilómetros recorridos por un automóvil alquilado por usuario, así como el número de días que ha estado alquilado el coche. Calcule el precio a pagar, sabiendo que el coche cuesta R $ 60,00 por día y R $ 0,15 por km recorrido. <br /></p>

<textarea data-component="activecode" id=q1_9 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def precio(km,dias):



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(precio(123,3),(0.15*123)+(60*3),"Esperado: "+str((0.15*123)+(60*3)))
        self.assertEqual(precio(800,4),(0.15*800)+(60*4) ,"Esperado: "+str((0.15*800)+(60*4)))
        self.assertEqual(precio(60,1), (0.15*60)+(60*1),"Esperado: "+str((0.15*60)+(60*1)))
        self.assertEqual(precio(90,2), (0.15*90)+(60*2),"Esperado: "+str((0.15*90)+(60*2)))
        self.assertEqual(precio(1016,7), (0.15*1016)+(60*7),"Esperado: "+str((0.15*1016)+(60*7)))


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_10" class="runestone explainer ac_section alert alert-warning">
<p>Escribe un programa para calcular la reducción en la vida útil de un fumador. Preguntar cantidad de cigarrillos fumados por día y cuántos años ha fumado. Considere que un fumador pierde 10 minutos de vida por cada cigarrillo, calcula cuántos días de vida perderá un fumador. Mostrar los días totales. <br /></p>

<textarea data-component="activecode" id=q1_10 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def fumador(cigarrillos,annos):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(fumador(10,1),((10*1*365)*10)/1440,"Esperado: "+str(((10*1*365)*10)/1440))
        self.assertEqual(fumador(3,5),((3*5*365)*10)/1440 ,"Esperado: "+str(((3*5*365)*10)/1440))
        self.assertEqual(fumador(1,8), ((1*8*365)*10)/1440,"Esperado: "+str(((1*8*365)*10)/1440))
        self.assertEqual(fumador(2,3), ((2*3*365)*10)/1440,"Esperado: "+str(((2*3*365)*10)/1440))


myTests().main()
</textarea>
</div>

<div data-childcomponent="q1_11" class="runestone explainer ac_section alert alert-warning">
<p>Sabiendo que str() convierte valores numéricos en cadenas, calcule cuántos dígitos hay en 2 elevados a un millón. <br /></p>

<textarea data-component="activecode" id=q1_11 data-lang="python" 
      data-timelimit=25000    
    data-audio=''      
           data-wasm=/_static>
def digitos():



====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):

        self.assertEqual(digitos(),301030,"Esperado: 301030")


myTests().main()
</textarea>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 6' data-toggle="tooltip">
    <a href="../lectures/Quiz6.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../lectures/Quiz6.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 3' data-toggle="tooltip" >
    <a href="Quiz3.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz3.html" >Next Section - Quiz - 3</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2017 Atharva Kulkarni.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 4.2.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>