<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 8 &#8212; PyZombis main@281eb94</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.css?v=6ab2699256a7cd2f2dee" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=0ddbfb97cbf8a612e45f" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=c0391f13c957789bc8fb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=700db7378411b25a09f9"></script>
    <script src="../_static/runestone.bundle.js?v=dac904eec8b4456305b9"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis main@281eb94
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz8_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-8">
<h1>Quiz - 8<a class="headerlink" href="#quiz-8" title="Permalink to this headline">¶</a></h1>
<div id='quiz8' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 14)</p>
<p>Duplicate ID – see quiz/Quiz8, line 18</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_1
   :nocodelens:

   Define a function ``first_or_last_6`` that checks whether 6 is the first
   or the last element in the list of numbers ``numbers`` passed as a parameter.
   If it is, return ``True``, otherwise return ``False``. |br| |br|
   Examples: |br|
   ``first_or_last_6([1, 2, 6])`` -&gt; ``True`` |br|
   ``first_or_last_6([6, 1, 2, 3])`` -&gt; ``True`` |br|
   ``first_or_last_6([3, 2, 1])`` -&gt; ``False`` |br|

   ~~~~
   def first_or_last_6(numbers):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
       def testOne(self):
           self.assertEqual(first_or_last_6([1, 2, 6]), True, &quot;Expected: True&quot;)
           self.assertEqual(first_or_last_6([3, 6]), True, &quot;Expected: True&quot;)
           self.assertEqual(first_or_last_6([3, 6]), True, &quot;Expected: True&quot;)
           self.assertEqual(first_or_last_6([3]), False, &quot;Expected: False&quot;)
           self.assertEqual(first_or_last_6([6, 1, 2, 3]), True, &quot;Expected: True&quot;)
           self.assertEqual(first_or_last_6([3, 2, 1]), False, &quot;Expected: False&quot;)
           self.assertEqual(first_or_last_6([3, 6, 1]), False, &quot;Expected: False&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 49)</p>
<p>Duplicate ID – see quiz/Quiz8, line 53</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_2
   :nocodelens:

   Define a function ``same_ends`` that takes a list ``numbers`` as a parameter,
   and returns ``True`` if the list of numbers has at least one element and the
   first element is the same as the last. Otherwise, return ``False``. |br| |br|
   Examples: |br|
   ``same_ends([1, 2, 3])`` -&gt; ``False`` |br|
   ``same_ends([1, 2, 3, 1])`` -&gt; ``True`` |br|
   ``same_ends([1, 2, 1])`` -&gt; ``True`` |br|

   ~~~~
   def same_ends(numbers):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(same_ends([1, 2, 3]), False, &quot;Expected: False&quot;)
           self.assertEqual(same_ends([1, 2, 1]), True, &quot;Expected: True&quot;)
           self.assertEqual(same_ends([7]), True, &quot;Expected: True&quot;)
           self.assertEqual(same_ends([]), False, &quot;Expected: False&quot;)
           self.assertEqual(same_ends([7, 7]), True, &quot;Expected: True&quot;)
           self.assertEqual(
                   same_ends([1, 2, 3, 1]),
                   True,
                   &quot;Expected: True&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 87)</p>
<p>Duplicate ID – see quiz/Quiz8, line 91</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_3
   :nocodelens:

   Define a function ``common_ends`` that takes two lists ``a`` and ``b``
   as parameters, and checks that the first two numbers of both lists are equal
   or that the last two numbers of both lists are equal. Assume both lists have
   at least one element. |br| |br|
   Examples: |br|
   ``common_ends([1, 2, 3], [7, 3])`` -&gt; ``True`` |br|
   ``common_ends([1, 2, 3], [7, 3, 2])`` -&gt; ``False`` |br|
   ``common_ends([1, 2, 3], [1, 3])`` -&gt; ``True`` |br|

   ~~~~
   def common_ends(a, b):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(common_ends([1, 2, 3], [1]), True, &quot;Expected: True&quot;)
           self.assertEqual(
                   common_ends([1, 2, 3], [2]),
                   False,
                   &quot;Expected: False&quot;
           )
           self.assertEqual(
                   common_ends([1, 2, 3], [7, 3]),
                   True,
                   &quot;Expected: True&quot;
           )
           self.assertEqual(
                   common_ends([1, 2, 3], [7, 3, 2]),
                   False,
                   &quot;Expected: False&quot;
           )
           self.assertEqual(
                   common_ends([1, 2, 3], [1, 3]),
                   True,
                   &quot;Expected: True&quot;
           )


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 138)</p>
<p>Duplicate ID – see quiz/Quiz8, line 142</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_4
   :nocodelens:

   Define a function called ``mayor_extremo`` that takes the list &quot;numeros&quot; as parameter, compares the ends of the list and returns a new list of the same size where all elements are the greatest magnitude extreme. |br| |br|
   Examples: |br|
   ``mayor_extremo([1, 2, 3])`` -&gt; ``[3, 3, 3]`` |br|
   ``mayor_extremo([1, 3, 2])`` -&gt; ``[2, 2, 2]`` |br|

   ~~~~
   def mayor_extremo(numeros):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(
                  mayor_extremo([1, 2, 3]),
                  [3, 3, 3],
                  &quot;Expected: [3,3,3]&quot;
          )
          self.assertEqual(
                  mayor_extremo([11, 5, 9]),
                  [11, 11, 11],
                  &quot;Expected: [11,11,11]&quot;
          )
          self.assertEqual(
                  mayor_extremo([2, 11, 3]),
                  [3, 3, 3],
                  &quot;Expected: [3,3,3]&quot;
          )
          self.assertEqual(
                  mayor_extremo([11, 3, 3]),
                  [11, 11, 11],
                  &quot;Expected: [11,11,11]&quot;
          )
          self.assertEqual(
                  mayor_extremo([3, 11, 11]),
                  [11, 11, 11],
                  &quot;Expected: [11,11,11]&quot;
          )
          self.assertEqual(
                  mayor_extremo([2, 2, 2]),
                  [2, 2, 2],
                  &quot;Expected: [2,2,2]&quot;
          )
          self.assertEqual(
                  mayor_extremo([2, 11, 2]),
                  [2, 2, 2],
                  &quot;Expected: [2,2,2]&quot;
          )
          self.assertEqual(
                  mayor_extremo([0, 0, 1]),
                  [1, 1, 1],
                  &quot;Expected: [1,1,1]&quot;
          )


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 204)</p>
<p>Duplicate ID – see quiz/Quiz8, line 210</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_5
   :nocodelens:

   Define a function called ``sumar_primeros_dos`` that takes the integer list ``numeros`` of any length as parameter and returns the sum of the first two elements. If the list has less than two elements, add 0s. |br| |br|
   Examples: |br|
   ``sumar_primeros_dos([1, 2, 3])`` -&gt; ``3`` |br|
   ``sumar_primeros_dos([1, 1])`` -&gt; ``2`` |br|
   ``sumar_primeros_dos([])`` -&gt; ``0`` |br|

   ~~~~
   def sumar_primeros_dos(numeros):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
          self.assertEqual(sumar_primeros_dos([1, 2, 3]), 3, &quot;Expected: 3&quot;)
          self.assertEqual(sumar_primeros_dos([1, 1]), 2, &quot;Expected: 2&quot;)
          self.assertEqual(sumar_primeros_dos([1, 1, 1, 1]), 2, &quot;Expected: 2&quot;)
          self.assertEqual(sumar_primeros_dos([1, 2]), 3, &quot;Expected: 3&quot;)
          self.assertEqual(sumar_primeros_dos([1]), 1, &quot;Expected: 1&quot;)
          self.assertEqual(sumar_primeros_dos([]), 0, &quot;Expected: 0&quot;)
          self.assertEqual(sumar_primeros_dos([4, 5, 6]), 9, &quot;Expected: 9&quot;)
          self.assertEqual(sumar_primeros_dos([4]), 4, &quot;Expected: 4&quot;)
          self.assertEqual(sumar_primeros_dos([4, -12, 5]), -8, &quot;Expected: -8&quot;)
          self.assertEqual(
                  sumar_primeros_dos([-200, 35, 10, 8]),
                  -165,
                  &quot;Expected: -165&quot;
          )


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 244)</p>
<p>Duplicate ID – see quiz/Quiz8, line 253</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_6
   :nocodelens:

   Define a function called ``al_medio`` that takes two integer lists ``a`` and ``b`` as parameters, and returns a list of size 2 that contains the middle elements of ``a`` and ``b``. Assume that the lists have an odd length. |br| |br|
   Examples: |br|
   ``al_medio([1, 2, 3], [4, 5, 6])`` -&gt; ``[2, 5]`` |br|
   ``al_medio([7, 7, 7], [3, 8, 0])`` -&gt; ``[7, 8]`` |br|
   ``al_medio([5, 2, 9], [1, 4, 5])`` -&gt; ``[2, 4]`` |br|

   ~~~~
   def al_medio(a, b):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(
                  al_medio([1, 2, 3], [4, 5, 6]),
                  [2, 5],
                  &quot;Expected: [2,5]&quot;
          )
          self.assertEqual(
                  al_medio([7, 7, 7], [3, 8, 0]),
                  [7, 8],
                  &quot;Expected: [7,8]&quot;
          )
          self.assertEqual(
                  al_medio([5, 2, 9], [1, 4, 5]),
                  [2, 4],
                  &quot;Expected: [2,4]&quot;
          )
          self.assertEqual(
                  al_medio([1, 9, 7], [4, 8, 8]),
          [9, 8],
                  &quot;Expected: [9,8]&quot;
          )
          self.assertEqual(
                  al_medio([1, 2, 3, 10, 2], [3, 1, 4]),
                  [3, 1],
                  &quot;Expected: [2,1]&quot;
          )
          self.assertEqual(
                  al_medio([1, 2, -2, 4, 1], [1]),
                  [-2, 1],
                  &quot;Expected: [2,1]&quot;
          )


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 300)</p>
<p>Duplicate ID – see quiz/Quiz8, line 313</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_7
   :nocodelens:

   Define a function ``fancy_date``. The function will take two parameters,
   ``my_score`` and ``partner_score``. Assume that you are going to a
   restaurant with your partner. The parameters represent the score of the clothes
   you are wearing from 0 to 10. The higher the score, the more elegant you are dressed.
   Your clothes score will determine if you get a table at the restaurant or not, according to the following rules: |br|

   - If the score of one of the two clothing is less than or equal to 2,
   they will not be entitled to a table (``0``). |br|
   - If the scores are higher, then if one of them is very
   elegant (score &gt;= 8) they will be entitled to a table (``2``). |br|
   - Otherwise, the answer is maybe (``1``). |br|

   So ``fancy_date`` returns a number between 0, 1, and 2, which mean No, Maybe, and Yes, respectively. |br| |br|
   Examples: |br|
   ``fancy_date(5, 10)`` -&gt; ``2`` |br|
   ``fancy_date(5, 2)`` -&gt; ``0`` |br|
   ``fancy_date(5, 5)`` -&gt; ``1`` |br|

   ~~~~
   def fancy_date(my_score, partner_score):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(fancy_date(5, 10), 2, &quot;Expected: 2&quot;)
          self.assertEqual(fancy_date(5, 2), 0, &quot;Expected: 0&quot;)
          self.assertEqual(fancy_date(5, 5), 1, &quot;Expected: 1&quot;)
          self.assertEqual(fancy_date(3, 3), 1, &quot;Expected: 1&quot;)
          self.assertEqual(fancy_date(10, 2), 0, &quot;Expected: 0&quot;)
          self.assertEqual(fancy_date(2, 9), 0, &quot;Expected: 0&quot;)
          self.assertEqual(fancy_date(9, 9), 2, &quot;Expected: 2&quot;)
          self.assertEqual(fancy_date(10, 5), 2, &quot;Expected: 2&quot;)
          self.assertEqual(fancy_date(2, 2), 0, &quot;Expected: 0&quot;)
          self.assertEqual(fancy_date(3, 7), 1, &quot;Expected: 1&quot;)
          self.assertEqual(fancy_date(2, 7), 0, &quot;Expected: 0&quot;)
          self.assertEqual(fancy_date(6, 2), 0, &quot;Expected: 0&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 8" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 350)</p>
<p>Duplicate ID – see quiz/Quiz8, line 364</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_8
   :nocodelens:

   The *squirrels* soccer team normally plays when the temperature is
   between 60 and 90 degrees Fahrenheit. But in summer, the temperature exceeds
   the maximum temperature they will play, which is 90 degrees Fahrenheit.
   Define a function ``squirrel_play`` that takes 2 parameters, ``temp`` representing
   the temperature, and ``is_summer``, which can be ``True`` if it&#39;s summer or ``False``
   if not. The function should return ``True`` if the *squirrels* play or ``False`` if they don&#39;t. |br| |br|
   Examples: |br|
   ``squirrel_play(70, False)`` -&gt; ``True`` |br|
   ``squirrel_play(95, False)`` -&gt; ``False``  |br|
   ``squirrel_play(95, True)`` -&gt; ``True`` |br|

   ~~~~
   def squirrel_play(temp, is_summer):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(squirrel_play(70, False), True, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(95, False), False, &quot;Expected: False&quot;)
           self.assertEqual(squirrel_play(95, True), True, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(90, False), True, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(90, True), True, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(50, False), False, &quot;Expected: False&quot;)
           self.assertEqual(squirrel_play(50, True), False, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(100, False), False, &quot;Expected: False&quot;)
           self.assertEqual(squirrel_play(100, True), True, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(105, True), False, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(59, False), False, &quot;Expected: False&quot;)
           self.assertEqual(squirrel_play(59, True), False, &quot;Expected: True&quot;)
           self.assertEqual(squirrel_play(60, False), True, &quot;Expected: False&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 9" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 395)</p>
<p>Duplicate ID – see quiz/Quiz8, line 409</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_9
   :nocodelens:

   Let&#39;s simulate a speedometer that applies a fine if the maximum speed is exceeded:

   - Speed &lt;= 60: no fine |br|
   - Speed between 61 and 80: medium fine |br|
   - Speed above 80: severe fine |br|

   If it&#39;s your birthday, the speed can be 5 km/h higher in all cases. |br|
   Define a function ``apply_fine`` that takes the parameters ``speed`` and
   ``birthday``. The first represents the speed you were going, and the second
   can be ``True`` if it&#39;s your birthday or ``False`` if it&#39;s not. This function should
   return 0, 1 or 2 according to the following: |br|

   - no fine = 0 |br|
   - medium fine = 1 |br|
   - severe fine = 2 |br| |br|

   Examples: |br|
   ``apply_fine(60, False)`` -&gt; ``0`` |br|
   ``apply_fine(65, False)`` -&gt; ``1`` |br|
   ``apply_fine(65, True)`` -&gt; ``0``  |br|
   ``apply_fine(80, False)`` -&gt; ``1`` |br|

   ~~~~
   def apply_fine(speed, birthday):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(apply_fine(20, False), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(20, True), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(60, False), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(65, False), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(65, True), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(80, False), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(85, False), 2, &quot;Expected: 2&quot;)
          self.assertEqual(apply_fine(85, True), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(70, False), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(75, False), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(75, True), 1, &quot;Expected: 1&quot;)
          self.assertEqual(apply_fine(40, False), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(40, True), 0, &quot;Expected: 0&quot;)
          self.assertEqual(apply_fine(90, False), 2, &quot;Expected: 2&quot;)
          self.assertEqual(apply_fine(90, True), 2, &quot;Expected: 2&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 10" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz8_en.rst</span>, line 453)</p>
<p>Duplicate ID – see quiz/Quiz8, line 467</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q8_10
  :nocodelens:

  Define a function ``set_alarm`` that takes two parameters: The first parameter
  ``day`` is a number from 0-6 that represents a day of the week according to the following: |br|
  day: 0 = Sunday, 1 = Monday, 2 = Tuesday, ..., 6 = Saturday. |br|
  The second parameter ``vacation`` will be set to ``True`` if you&#39;re on vacation or
  ``False`` if you&#39;re not. The function will return a string that indicates when
  the alarm clock will ring according to the following: |br|

  - Weekdays: &#39;07:00&#39; |br|
  - Weekends: &#39;10:00&#39; |br|

  Unless you&#39;re on vacation, in which case: |br|

  - Weekdays: &#39;10:00&#39; |br|
  - Weekends: &#39;off&#39; |br| |br|

  Examples: |br|
  ``set_alarm(1, False)`` -&gt; ``&quot;7:00&quot;`` |br|
  ``set_alarm(1, True)`` -&gt; ``&quot;10:00&quot;`` |br|
  ``set_alarm(6, True)`` -&gt; ``&quot;off&quot;`` |br|
  ``set_alarm(0, False)`` -&gt; ``&quot;10:00&quot;`` |br|

  ~~~~
  def set_alarm(day, vacation):


  ====
  from unittest.gui import TestCaseGui


  class myTests(TestCaseGui):
          def testOne(self):
          self.assertEqual(set_alarm(1, False), &quot;7:00&quot;, &quot;Expected: 7:00&quot;)
          self.assertEqual(set_alarm(5, False), &quot;7:00&quot;, &quot;Expected: 7:00&quot;)
          self.assertEqual(set_alarm(0, False), &quot;10:00&quot;, &quot;Expected: 10:00&quot;)
          self.assertEqual(set_alarm(6, False), &quot;10:00&quot;, &quot;Expected: 10:00&quot;)
          self.assertEqual(set_alarm(0, True), &quot;off&quot;, &quot;Expected: off&quot;)
          self.assertEqual(set_alarm(6, True), &quot;off&quot;, &quot;Expected: off&quot;)
          self.assertEqual(set_alarm(1, True), &quot;10:00&quot;, &quot;Expected: 10:00&quot;)
          self.assertEqual(set_alarm(3, True), &quot;10:00&quot;, &quot;Expected: 10:00&quot;)
          self.assertEqual(set_alarm(5, True), &quot;10:00&quot;, &quot;Expected: 10:00&quot;)


  myTests().main()
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>