<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 5 &#8212; PyZombis main@765ab40</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.css?v=6ab2699256a7cd2f2dee" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=0ddbfb97cbf8a612e45f" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=c0391f13c957789bc8fb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=700db7378411b25a09f9"></script>
    <script src="../_static/runestone.bundle.js?v=dac904eec8b4456305b9"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis main@765ab40
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz5_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-5">
<h1>Quiz - 5<a class="headerlink" href="#quiz-5" title="Permalink to this headline">¶</a></h1>
<div id='quiz5_1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1_1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 14)</p>
<p>Duplicate ID – see quiz/Quiz5, line 18</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span>  <span class="n">clickablearea</span><span class="p">::</span> <span class="n">q5_1</span>
    <span class="p">:</span><span class="n">question</span><span class="p">:</span> <span class="n">The</span> <span class="n">following</span> <span class="n">program</span> <span class="n">declares</span> <span class="n">a</span> <span class="n">variable</span> <span class="ow">in</span> <span class="n">the</span> <span class="n">wrong</span> <span class="n">place</span> <span class="ow">in</span> <span class="n">the</span> <span class="n">code</span><span class="o">.</span> <span class="n">Your</span> <span class="n">task</span> <span class="ow">is</span> <span class="n">to</span> <span class="n">select</span> <span class="n">the</span> <span class="n">incorrect</span> <span class="n">declaration</span> <span class="ow">and</span> <span class="n">select</span> <span class="n">the</span> <span class="n">correct</span> <span class="n">line</span> <span class="n">where</span> <span class="n">it</span> <span class="n">should</span> <span class="n">go</span><span class="o">.</span>
    <span class="p">:</span><span class="n">iscode</span><span class="p">:</span>
    <span class="p">:</span><span class="n">feedback</span><span class="p">:</span> <span class="n">Remember</span><span class="p">,</span> <span class="n">you</span> <span class="n">cannot</span> <span class="n">use</span> <span class="n">a</span> <span class="n">variable</span> <span class="n">without</span> <span class="n">declaring</span> <span class="n">it</span> <span class="n">previously</span><span class="o">.</span>

    <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="n">x</span> <span class="o">=</span> <span class="mi">2</span><span class="p">:</span><span class="n">endclick</span><span class="p">:</span>
    <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">correct</span><span class="p">:</span>     <span class="p">:</span><span class="n">endclick</span><span class="p">:</span>
    <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="k">if</span> <span class="n">y</span> <span class="o">&gt;</span> <span class="mi">8</span> <span class="n">then</span><span class="p">::</span><span class="n">endclick</span><span class="p">:</span>
        <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">correct</span><span class="p">:</span><span class="n">y</span> <span class="o">=</span> <span class="mi">5</span><span class="p">:</span><span class="n">endclick</span><span class="p">:</span>
        <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="n">y</span> <span class="o">=</span> <span class="n">y</span> <span class="o">*</span> <span class="mi">2</span><span class="p">:</span><span class="n">endclick</span><span class="p">:</span>

    <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="k">else</span><span class="p">::</span><span class="n">endclick</span><span class="p">:</span>
        <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="n">x</span> <span class="o">=</span> <span class="n">x</span> <span class="o">*</span> <span class="mi">2</span><span class="p">:</span><span class="n">endclick</span><span class="p">:</span>
        <span class="p">:</span><span class="n">click</span><span class="o">-</span><span class="n">incorrect</span><span class="p">:</span><span class="nb">print</span><span class="p">(</span><span class="n">x</span> <span class="o">+</span> <span class="n">y</span><span class="p">)::</span><span class="n">endclick</span><span class="p">:</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 1_2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 32)</p>
<p>Duplicate ID – see quiz/Quiz5, line 36</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span>  <span class="n">activecode</span><span class="p">::</span> <span class="n">q5_2</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Translate</span> <span class="n">the</span> <span class="n">previous</span> <span class="n">program</span> <span class="p">(</span><span class="k">with</span> <span class="n">the</span> <span class="n">declaration</span> <span class="ow">in</span> <span class="n">the</span> <span class="n">right</span> <span class="n">place</span><span class="p">)</span> <span class="n">to</span> <span class="n">Python</span> <span class="n">using</span> <span class="n">the</span> <span class="n">following</span> <span class="n">code</span> <span class="n">block</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="o">~~~~</span>

    <span class="o">====</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 1_3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 43)</p>
<p>Duplicate ID – see quiz/Quiz5, line 47</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  fillintheblank:: q5_3

    What does the previous program print?

    -   :9: Your answer is correct, good job!
        :.*: Your answer is incorrect, please try again.
</pre></div>
</div>
</div>
</div>
    </div>

<div id='quiz5_2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 2_1" >
<div class="highlight-none notranslate"><div class="highlight"><pre><span></span>for i = 1 to 9:
    if i != 3 then:
        for j = 1 to 6:
            print(&quot;Hello&quot;)
</pre></div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2_2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 65)</p>
<p>Duplicate ID – see quiz/Quiz5, line 69</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span>  <span class="n">activecode</span><span class="p">::</span> <span class="n">q5_4</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Translate</span> <span class="n">the</span> <span class="n">previous</span> <span class="n">program</span> <span class="n">to</span> <span class="n">Python</span> <span class="n">using</span> <span class="n">the</span> <span class="n">following</span> <span class="n">code</span> <span class="n">block</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="o">**</span><span class="n">Note</span><span class="o">**</span><span class="p">:</span> <span class="ow">in</span> <span class="n">our</span> <span class="n">pseudolanguage</span><span class="p">,</span> <span class="n">the</span> <span class="n">loop</span> <span class="n">includes</span> <span class="n">the</span> <span class="n">endpoints</span><span class="p">,</span> <span class="n">that</span> <span class="ow">is</span><span class="p">,</span> <span class="mi">1</span> <span class="n">to</span> <span class="mi">4</span> <span class="n">means</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mf">4.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="o">~~~~</span>

    <span class="o">====</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2_3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 77)</p>
<p>Duplicate ID – see quiz/Quiz5, line 81</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  fillintheblank:: q5_5

    How many times does the previous program print ``&quot;Hello&quot;``?

    -   :48: Your answer is correct, good job!
        :.*: Your answer is incorrect, please try again.
</pre></div>
</div>
</div>
</div>
    </div>

<div id='quiz5_3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 3_1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 89)</p>
<p>Duplicate ID – see quiz/Quiz5, line 93</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  activecode:: q5_6
    :nocodelens:

    Develop the function ``pares_divisibles_7`` which receives two positive integers, ``inicio`` and ``fin``, that represent the start and end of a range of numbers. The function must return the amount of numbers that are even and also divisible by 7. |br| |br|
    Examples: |br|
    ``pares_divisibles_7(1, 7)`` -&gt; ``0`` |br|
    ``pares_divisibles_7(25, 123)`` -&gt; ``7`` |br|
    ``pares_divisibles_7(13, 245)`` -&gt; ``17`` |br|

    ~~~~
    def pares_divisibles_7(inicio, fin):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):
            self.assertEqual(pares_divisibles_7(1, 7), 0, &quot;Expected: 0&quot;)
            self.assertEqual(pares_divisibles_7(25, 123), 7, &quot;Expected: 7&quot;)
            self.assertEqual(pares_divisibles_7(13, 245), 17, &quot;Expected: 17&quot;)
            self.assertEqual(pares_divisibles_7(1300, 2460), 83, &quot;Expected: 83&quot;)
            self.assertEqual(pares_divisibles_7(1, 2460), 175, &quot;Expected: 175&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3_2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 121)</p>
<p>Duplicate ID – see quiz/Quiz5, line 125</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  fillintheblank:: q5_7

    Use the function created in the previous exercise to answer How many numbers are even and also divisible by 7 between 1067 and 3627 (inclusive)?

    -   :183: Your answer is correct, good job!
        :.*: Your answer is incorrect, try again.
</pre></div>
</div>
</div>
</div>
    </div>

<div id='quiz5_4' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 4_1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 133)</p>
<p>Duplicate ID – see quiz/Quiz5, line 137</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  activecode:: q5_8
    :nocodelens:

    Daniela is a very superstitious person. For her, a number is **lucky** if it contains the digit ``2`` but not the ``7``. She is also very curious and wants to know how many **lucky** numbers are in a range of numbers (including the ends)?. That&#39;s why she has asked you to make the function called ``cuantos_tienen_suerte`` which receives two positive integers, ``inicio`` and ``fin``, that represent the start and end of a range of numbers. The function must return the amount of numbers that are **lucky**. |br| |br|
    Examples: |br|
    ``cuantos_tienen_suerte(1, 7)`` -&gt; ``1`` |br|
    ``cuantos_tienen_suerte(1, 20)`` -&gt; ``3`` |br|
    ``cuantos_tienen_suerte(25, 123)`` -&gt; ``16`` |br|
    ``cuantos_tienen_suerte(13, 245)`` -&gt; ``74`` |br|

    ~~~~
    def cuantos_tienen_suerte(inicio, fin):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):
            self.assertEqual(cuantos_tienen_suerte(1, 1), 0, &quot;Expected: 0&quot;)
            self.assertEqual(cuantos_tienen_suerte(1, 7), 1, &quot;Expected: 1&quot;)
            self.assertEqual(cuantos_tienen_suerte(1, 20), 3, &quot;Expected: 3&quot;)
            self.assertEqual(cuantos_tienen_suerte(25, 123), 16, &quot;Expected: 16&quot;)
            self.assertEqual(cuantos_tienen_suerte(1300, 2460), 481, &quot;Expected: 481&quot;)
            self.assertEqual(cuantos_tienen_suerte(13, 245), 74, &quot;Expected: 74&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4_2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 167)</p>
<p>Duplicate ID – see quiz/Quiz5, line 171</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  fillintheblank:: q5_9

    Use the function created in the previous exercise (**MANDATORY** to have finished it before and passed all tests) to answer How many **lucky** numbers are there between 18644 and 33087 (inclusive)?

    -   :7995: Your answer is correct, good job!
        :.*: Your answer is incorrect, try again.
</pre></div>
</div>
</div>
</div>
    </div>

<div id='quiz5_5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 5_1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 179)</p>
<p>Error in “activecode” directive:
maximum 2 argument(s) allowed, 3 supplied.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode :: q5_10
    : nocodelens:

    In the quiet rural village of *Ponteironuloville*, all phones have 6 digits. The telephone company establishes the following rules about numbers: | br |

    1. There cannot be two identical consecutive digits, because this is boring
    2. The sum of the digits must be even, because this is legal
    3. The last digit cannot be the same as the first, because that is bad luck.

    Then, given these perfectly reasonable, well-designed and mature rules, develop the function called ``es_numero_ponteironuloville`` that receives a string of positive integers and returns ``True`` if the number is valid according to the rules of *Ponteironuloville* and ``False`` otherwise.
    | br | | br |

    Examples: | br |
    ``es_numero_ponteironuloville(&quot;123457&quot;)`` -&gt; ``True`` | br |
    ``es_numero_ponteironuloville(&quot;234562&quot;)`` -&gt; ``False`` | br |
    ``es_numero_ponteironuloville(&quot;222222&quot;)`` -&gt; ``False`` | br |
    ``es_numero_ponteironuloville(&quot;123456&quot;)`` -&gt; ``False`` | br |
    ``es_numero_ponteironuloville(&quot;312214&quot;)`` -&gt; ``False`` | br |
    ``es_numero_ponteironuloville(&quot;312312&quot;)`` -&gt; ``True`` | br |

    ~~~~
    def es_numero_ponteironuloville(numero):



    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):
            self.assertEqual(es_numero_ponteironuloville(&quot;123457&quot;), True, &quot;Expected: True&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;234562&quot;), False, &quot;Expected: False&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;304706&quot;), True, &quot;Expected: True&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;222222&quot;), False, &quot;Expected: False&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;123456&quot;), False, &quot;Expected: False&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;312214&quot;), False, &quot;Expected: False&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;312312&quot;), True, &quot;Expected: True&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;131313&quot;), True, &quot;Expected: True&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;249184&quot;), True, &quot;Expected: True&quot;)
            self.assertEqual(es_numero_ponteironuloville(&quot;012445&quot;), False, &quot;Expected: False&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5_2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 227)</p>
<p>Duplicate ID – see quiz/Quiz5, line 231</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  activecode:: q5_11
    :nocodelens:
    :include: q5_10

    Knowing that the ``split()`` function divides a string into multiple strings, use the function created in the previous exercise (**MANDATORY** to have finished it before and passed all tests) to complete the ``count_valid`` function that returns the number of valid numbers from the list of numbers given in the following block of code. |br| |br|
    **Note**: the list of numbers is a string, wisely use the ``split()`` function to obtain a list of strings.
    ~~~~
    numbers_list = &quot;&quot;&quot;236043 237330 239636 240138 242123 246224 249183 252936
        254711 257200 257607 261424 263814 266794 268649 273050
        275001 277606 278997 283331 287104 287953 289137 291591
        292559 292946 295180 295566 297529 300400 304707 306931
        310638 313595 318449 319021 322082 323796 326266 326880
        327249 329914 334392 334575 336723 336734 338808 343269
        346040 350113 353631 357154 361633 361891 364889 365746
        365749 366426 369156 369444 369689 372896 374983 375223
        379163 380712 385640 386777 388599 389450 390178 392943
        394742 395921 398644 398832 401149 402219 405364 408088
        412901 417683 422267 424767 426613 430474 433910 435054
        440052 444630 447852 449116 453865 457631 461750 462985
        463328 466458 469601 473108 476773 477956 481991 482422
        486195 488359 489209 489388 491928 496569 496964 497901
        500877 502386 502715 507617 512526 512827 513796 518232
        521455 524277 528496 529345 531231 531766 535067 535183
        536593 537360 539055 540582 543708 547492 550779 551595
        556493 558807 559102 562050 564962 569677 570945 575447
        579937 580112 580680 582458 583012 585395 586244 587393
        590483 593112 593894 594293 597525 598184 600455 600953
        601523 605761 608618 609198 610141 610536 612636 615233
        618314 622752 626345 626632 628889 629457 629643 633673
        637656 641136 644176 644973 647617 652218 657143 659902
        662224 666265 668010 672480 672695 676868 677125 678315&quot;&quot;&quot;

    def count_valid():
        count = 0
        &quot;&quot;&quot;
            Write your code here, using the created variable
        &quot;&quot;&quot;
        return count
    ====
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5_3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz5_en.rst</span>, line 270)</p>
<p>Duplicate ID – see quiz/Quiz5, line 274</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>..  fillintheblank:: q5_12

    Use the function created in the previous exercise to answer how many numbers from the list are valid according to the rules of *Ponteironuloville*?

    -   :39: Your answer is correct, good job!
        :.*: Your answer is incorrect, try again.
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>