<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 6 &#8212; PyZombis main@765ab40</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.css?v=6ab2699256a7cd2f2dee" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=0ddbfb97cbf8a612e45f" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=c0391f13c957789bc8fb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=700db7378411b25a09f9"></script>
    <script src="../_static/runestone.bundle.js?v=dac904eec8b4456305b9"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis main@765ab40
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz6_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-6">
<h1>Quiz - 6<a class="headerlink" href="#quiz-6" title="Permalink to this headline">¶</a></h1>
<div id='quiz6' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 14)</p>
<p>Duplicate ID – see quiz/Quiz6, line 18</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_1
   :nocodelens:

   Develop the function ``sleep`` that takes two parameters,
   ``weekday`` and ``holiday``, both parameters take boolean values,
   meaning they can be either ``True`` or ``False``. The function acts as follows:
   you can sleep when it&#39;s a holiday or when it&#39;s not a weekday.
   ``sleep`` will return ``True`` or ``False`` whether or not you will sleep. |br| |br|
   Examples: |br|
   ``sleep(False, False)`` -&gt; ``True`` |br|
   ``sleep(True, False)`` -&gt; ``False`` |br|

   ~~~~
   def sleep(weekday, holiday):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(sleep(False, False), True, &quot;Excellent! Obtained: True Expected: True&quot;)
         self.assertEqual(sleep(True, False), False, &quot;Excellent! Obtained: False Expected: False&quot;)
         self.assertEqual(sleep(False, True), True, &quot;Excellent! Obtained: True Expected: True&quot;)
         self.assertEqual(sleep(True, True), True, &quot;Excellent! Obtained: True Expected: True&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 47)</p>
<p>Duplicate ID – see quiz/Quiz6, line 52</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_2
   :nocodelens:

   Develop the function ``students_in_trouble`` that takes two parameters,
   ``a_smiling`` and ``b_smiling``. a and b represent two students.
   ``a_smiling`` and ``b_smiling`` indicate if a and b are smiling. When both
   are smiling or both are not smiling we have trouble. ``students_in_trouble``
   should return ``True`` when there is trouble. Otherwise, it will return ``False``.
   |br| |br|
   Examples: |br|
   ``students_in_trouble(True, True)`` -&gt; ``True`` |br|
   ``students_in_trouble(False, True)`` -&gt; ``False`` |br|

   ~~~~
   def students_in_trouble(a_smiling, b_smiling):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(students_in_trouble(True, True), True, &quot;Expected: True&quot;)
         self.assertEqual(students_in_trouble(False, False), True, &quot;Expected: True&quot;)
         self.assertEqual(students_in_trouble(True, False), False, &quot;Expected: False&quot;)
         self.assertEqual(students_in_trouble(False, True), False, &quot;Expected: False&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 81)</p>
<p>Duplicate ID – see quiz/Quiz6, line 86</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_3
   :nocodelens:

   Develop the function ``double_sum`` that takes two parameters, ``a`` and ``b``.
   Both are integers. The function should return the sum of ``a`` and ``b``.
   However, if the numbers are equal, it returns twice the sum. |br| |br|
   Examples: |br|
   ``double_sum(1, 2)`` -&gt; ``3`` |br|
   ``double_sum(2, 2)`` -&gt; ``8`` |br|

   ~~~~
   def double_sum(a, b):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(double_sum(1, 2), 3, &quot;Excellent! gotten: 3 expected: 3&quot;)
         self.assertEqual(double_sum(3, 2), 5, &quot;Excellent! gotten: 5 expected: 5&quot;)
         self.assertEqual(double_sum(2, 2), 8, &quot;Excellent! gotten: 8 expected: 8&quot;)
         self.assertEqual(double_sum(-1, 0), -1, &quot;Excellent! gotten: -1 expected: -1&quot;)
         self.assertEqual(double_sum(0, 0), 0, &quot;Excellent! gotten: 0 expected: 0&quot;)
         self.assertEqual(double_sum(0, 1), 1, &quot;Excellent! gotten: 1 expected: 1&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 113)</p>
<p>Duplicate ID – see quiz/Quiz6, line 119</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_4
   :nocodelens:

   Develop the function ``absolute_difference_21`` that takes a parameter, ``n``,
   and returns the absolute difference between ``n`` and 21 only if ``n``
   is less than or equal to 21. If it is greater, then it returns twice the
   absolute difference between the number and 21. **Remember**: ``abs(x)`` returns
   the absolute value of x. |br| |br|
   Examples: |br|
   ``absolute_difference_21(19)`` -&gt; ``2`` |br|
   ``absolute_difference_21(25)`` -&gt; ``8`` |br|

   ~~~~
   def absolute_difference_21(n):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(absolute_difference_21(19), 2, &quot;Expected: 2&quot;)
         self.assertEqual(absolute_difference_21(10), 11, &quot;Expected: 11&quot;)
         self.assertEqual(absolute_difference_21(21), 0, &quot;Expected: 0&quot;)
         self.assertEqual(absolute_difference_21(22), 2, &quot;Expected: 2&quot;)
         self.assertEqual(absolute_difference_21(25), 8, &quot;Expected: 8&quot;)
         self.assertEqual(absolute_difference_21(30), 18, &quot;Expected: 18&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 148)</p>
<p>Duplicate ID – see quiz/Quiz6, line 154</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_5
   :nocodelens:

   Develop the function ``parrot_trouble`` that takes two parameters,
   ``talking``, which can be ``True`` or ``False``, and ``hour``, which takes a value
   between 0 and 23. We have a parrot, and there is trouble if the parrot is talking
   before 7 hours or after 20 hours. Return ``True`` if there are
   problems or ``False`` if there are not. |br| |br|
   Examples: |br|
   ``parrot_trouble(True, 6)`` -&gt; ``True`` |br|
   ``parrot_trouble(True, 20)`` -&gt; ``False`` |br|

   ~~~~
   def parrot_trouble(talking, hour):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(parrot_trouble(True, 6), True, &quot;Expected: True&quot;)
         self.assertEqual(parrot_trouble(True, 7), False, &quot;Expected: False&quot;)
         self.assertEqual(parrot_trouble(True, 21), True, &quot;Expected: True&quot;)
         self.assertEqual(parrot_trouble(True, 23), True, &quot;Expected: True&quot;)
         self.assertEqual(parrot_trouble(True, 20), False, &quot;Expected: False&quot;)
         self.assertEqual(parrot_trouble(False, 6), False, &quot;Expected: False&quot;)
         self.assertEqual(parrot_trouble(False, 21), False, &quot;Expected: False&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 184)</p>
<p>Duplicate ID – see quiz/Quiz6, line 190</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_6
   :nocodelens:

   Develop the function ``makes10`` that takes two parameters, ``a`` and ``b``.
   Returns ``True`` if one of the parameters is 10, or if the sum of both
   is 10. Otherwise, it returns ``False``. |br| |br|
   Examples: |br|
   ``makes10(9,10)`` -&gt; ``True`` |br|
   ``makes10(1,9)`` -&gt; ``True`` |br|
   ``makes10(8,3)`` -&gt; ``False`` |br|

   ~~~~
   def makes10(a, b):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(makes10(9, 10), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(9, 9), False, &quot;Expected: False&quot;)
         self.assertEqual(makes10(1, 9), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(10, 1), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(10, 10), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(8, 2), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(8, 3), False, &quot;Expected: False&quot;)
         self.assertEqual(makes10(10, 42), True, &quot;Expected: True&quot;)
         self.assertEqual(makes10(12, -2), True, &quot;Expected: True&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 221)</p>
<p>Duplicate ID – see quiz/Quiz6, line 227</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_7
   :nocodelens:

   Develop the function ``near_hundred`` that takes ``n``
   as a parameter, which is an integer. Returns ``True`` if
   the absolute difference between ``n`` and 100 or ``n`` and 200 is less than or
   equal to 10. |br| |br|
   Examples: |br|
   ``near_hundred(93)`` -&gt; ``True`` |br|
   ``near_hundred(90)`` -&gt; ``True`` |br|
   ``near_hundred(89)`` -&gt; ``False`` |br|
   ``near_hundred(210)`` -&gt; ``True`` |br|
   ``near_hundred(211)`` -&gt; ``False`` |br|

   ~~~~
   def near_hundred(n):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(near_hundred(93), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(90), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(89), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(110), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(111), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(121), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(0), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(5), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(191), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(189), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(190), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(200), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(210), True, &quot;Expected: True&quot;)
         self.assertEqual(near_hundred(211), False, &quot;Expected: False&quot;)
         self.assertEqual(near_hundred(290), False, &quot;Expected: False&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 8" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 266)</p>
<p>Duplicate ID – see quiz/Quiz6, line 273</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_8
   :nocodelens:

   Develop the function ``remover_iesimo`` that receives a non-empty string ``s`` and a positive integer ``i``, and returns the original string without the i-th character. |br| |br|
   Examples: |br|
   ``remover_iesimo(&quot;Hello&quot;, 1)`` -&gt; ``&quot;ello&quot;`` |br|
   ``remover_iesimo(&quot;Hi&quot;, 2)`` -&gt; ``&quot;H&quot;`` |br|
   ``remover_iesimo(&quot;PyZombiess&quot;, 10)`` -&gt; ``&quot;PyZombies&quot;`` |br|

   ~~~~
   def remover_iesimo(s, i):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(quitar_caracter(&quot;Hi&quot;, 0), &quot;i&quot;, &quot;Expected: i&quot;)
         self.assertEqual(quitar_caracter(&quot;Hi&quot;, 1), &quot;H&quot;, &quot;Expected: H&quot;)
         self.assertEqual(quitar_caracter(&quot;code&quot;, 0), &quot;ode&quot;, &quot;Expected: ode&quot;)
         self.assertEqual(quitar_caracter(&quot;code&quot;, 1), &quot;cde&quot;, &quot;Expected: cde&quot;)
         self.assertEqual(quitar_caracter(&quot;code&quot;, 2), &quot;coe&quot;, &quot;Expected: coe&quot;)
         self.assertEqual(quitar_caracter(&quot;code&quot;, 3), &quot;cod&quot;, &quot;Expected: cod&quot;)
         self.assertEqual(quitar_caracter(&quot;gatito&quot;, 1), &quot;gtito&quot;, &quot;Expected: gtito&quot;)
         self.assertEqual(quitar_caracter(&quot;gatito&quot;, 0), &quot;atito&quot;, &quot;Expected: atito&quot;)
         self.assertEqual(quitar_caracter(&quot;gatito&quot;, 4), &quot;gatio&quot;, &quot;Expected: gatio&quot;)
         self.assertEqual(quitar_caracter(&quot;chocolate&quot;, 8), &quot;chocolat&quot;, &quot;Expected: chocolat&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 9" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6_en.rst</span>, line 302)</p>
<p>Duplicate ID – see quiz/Quiz6, line 311</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q6_9
   :nocodelens:

   Develop the function ``intercambiar`` that takes ``s`` as a parameter, representing a string. If ``s`` has a length of 1 or less, the same string is returned. Otherwise, the first and last letter of ``s`` are exchanged. |br| |br|
   Examples: |br|
   ``intercambiar(&quot;codigo&quot;)`` -&gt; ``&quot;oodigc&quot;`` |br|
   ``intercambiar(&quot;a&quot;)`` -&gt; ``&quot;a&quot;`` |br|
   ``intercambiar(&quot;ab&quot;)`` -&gt; ``&quot;ba&quot;`` |br|

   ~~~~
   def intercambiar(s):


   ====
   from unittest.gui import TestCaseGui


   class myTests(TestCaseGui):
      def testOne(self):
         self.assertEqual(intercambiar(&quot;code&quot;), &quot;eodc&quot;, &quot;Expected: eodc&quot;)
         self.assertEqual(intercambiar(&quot;a&quot;), &quot;a&quot;, &quot;Expected: a&quot;)
         self.assertEqual(intercambiar(&quot;ab&quot;), &quot;ba&quot;, &quot;Expected: ba&quot;)
         self.assertEqual(intercambiar(&quot;abc&quot;), &quot;cba&quot;, &quot;Expected: cba&quot;)
         self.assertEqual(intercambiar(&quot; &quot;), &quot; &quot;, &quot;Expected: &#39; &#39;&quot;)
         self.assertEqual(intercambiar(&quot;nythoP&quot;), &quot;Python&quot;, &quot;Expected: Python&quot;)
         self.assertEqual(intercambiar(&quot;hello&quot;), &quot;oellh&quot;, &quot;Expected: oellh&quot;)
         self.assertEqual(intercambiar(&quot;Chocolate&quot;), &quot;ehocolatC&quot;, &quot;Expected: ehocolatC&quot;)


   myTests().main()
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>