<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 1 &#8212; PyZombis main@281eb94</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" type="text/css" href="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.css?v=6ab2699256a7cd2f2dee" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.css?v=0ddbfb97cbf8a612e45f" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=c0391f13c957789bc8fb"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=700db7378411b25a09f9"></script>
    <script src="../_static/runestone.bundle.js?v=dac904eec8b4456305b9"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis main@281eb94
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz1_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-1">
<h1>Quiz - 1<a class="headerlink" href="#quiz-1" title="Permalink to this headline">¶</a></h1>
<div id='quiz1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 14)</p>
<p>Duplicate ID – see quiz/Quiz1, line 18</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_1</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Make</span> <span class="n">a</span> <span class="n">program</span> <span class="n">that</span> <span class="n">asks</span> <span class="k">for</span> <span class="n">two</span> <span class="n">integer</span> <span class="n">numbers</span> <span class="ow">and</span> <span class="n">prints</span> <span class="n">the</span> <span class="nb">sum</span> <span class="n">of</span> <span class="n">those</span> <span class="n">two</span> <span class="n">numbers</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">suma</span><span class="p">(</span><span class="n">n</span><span class="p">,</span> <span class="n">m</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma</span><span class="p">(</span><span class="mi">24</span><span class="p">,</span> <span class="mi">42</span><span class="p">),</span> <span class="mi">66</span><span class="p">,</span> <span class="s2">&quot;Expected: 66&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma</span><span class="p">(</span><span class="mi">17</span><span class="p">,</span> <span class="mi">13</span><span class="p">),</span> <span class="mi">30</span><span class="p">,</span> <span class="s2">&quot;Expected: 30&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma</span><span class="p">(</span><span class="o">-</span><span class="mi">11</span><span class="p">,</span> <span class="mi">6</span><span class="p">),</span> <span class="o">-</span><span class="mi">5</span><span class="p">,</span> <span class="s2">&quot;Expected: -5&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="mi">9</span><span class="p">),</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: 9&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 41)</p>
<p>Duplicate ID – see quiz/Quiz1, line 45</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_2</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Write</span> <span class="n">a</span> <span class="n">program</span> <span class="n">that</span> <span class="n">reads</span> <span class="n">a</span> <span class="n">value</span> <span class="ow">in</span> <span class="n">meters</span> <span class="ow">and</span> <span class="n">shows</span> <span class="n">it</span> <span class="n">converted</span> <span class="n">to</span> <span class="n">millimeters</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">metros_a_milimetros</span><span class="p">(</span><span class="n">n</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">metros_a_milimetros</span><span class="p">(</span><span class="mi">1</span><span class="p">),</span> <span class="mi">1000</span><span class="p">,</span><span class="s2">&quot;Expected: 1000&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">metros_a_milimetros</span><span class="p">(</span><span class="mf">0.2</span><span class="p">),</span> <span class="mi">200</span><span class="p">,</span><span class="s2">&quot;Expected: 200&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">metros_a_milimetros</span><span class="p">(</span><span class="mi">30</span><span class="p">),</span> <span class="mi">30000</span><span class="p">,</span><span class="s2">&quot;Expected: 30000&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 67)</p>
<p>Duplicate ID – see quiz/Quiz1, line 71</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_3</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Write</span> <span class="n">a</span> <span class="n">program</span> <span class="n">that</span> <span class="n">reads</span> <span class="n">the</span> <span class="n">number</span> <span class="n">of</span> <span class="n">days</span><span class="p">,</span> <span class="n">hours</span><span class="p">,</span> <span class="n">minutes</span><span class="p">,</span> <span class="ow">and</span> <span class="n">seconds</span> <span class="kn">from</span>
    <span class="nn">the</span> <span class="n">user</span><span class="o">.</span> <span class="n">Calculate</span> <span class="n">the</span> <span class="n">total</span> <span class="n">number</span> <span class="n">of</span> <span class="n">seconds</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">tiempo_en_segundos</span><span class="p">(</span><span class="n">dias</span><span class="p">,</span> <span class="n">horas</span><span class="p">,</span> <span class="n">minutos</span><span class="p">,</span> <span class="n">segundos</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">tiempo_en_segundos</span><span class="p">(</span><span class="mi">2</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">5</span><span class="p">),</span> <span class="mi">190925</span><span class="p">,</span> <span class="s2">&quot;Expected: 190925&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">tiempo_en_segundos</span><span class="p">(</span><span class="mi">10</span><span class="p">,</span> <span class="mi">89</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">0</span><span class="p">),</span> <span class="mi">1184700</span><span class="p">,</span> <span class="s2">&quot;Expected: 1184700&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">tiempo_en_segundos</span><span class="p">(</span><span class="mi">8</span><span class="p">,</span> <span class="mi">0</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">0</span><span class="p">),</span> <span class="mi">691320</span><span class="p">,</span> <span class="s2">&quot;Expected: 691320&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">tiempo_en_segundos</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">55</span><span class="p">,</span> <span class="mi">6</span><span class="p">),</span> <span class="mi">21306</span><span class="p">,</span> <span class="s2">&quot;Expected: 21306&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 95)</p>
<p>Duplicate ID – see quiz/Quiz1, line 99</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_4</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Make</span> <span class="n">a</span> <span class="n">program</span> <span class="n">that</span> <span class="n">calculates</span> <span class="n">a</span> <span class="n">salary</span> <span class="n">increase</span><span class="o">.</span> <span class="n">It</span> <span class="n">should</span> <span class="n">request</span> <span class="n">the</span>
    <span class="n">salary</span> <span class="n">amount</span> <span class="ow">and</span> <span class="n">the</span> <span class="n">percentage</span> <span class="n">of</span> <span class="n">the</span> <span class="n">increase</span><span class="o">.</span> <span class="n">Display</span> <span class="n">the</span> <span class="n">amount</span> <span class="n">of</span>
    <span class="n">the</span> <span class="n">increase</span> <span class="ow">and</span> <span class="n">the</span> <span class="n">new</span> <span class="n">salary</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">aumento</span><span class="p">(</span><span class="n">salario</span><span class="p">,</span> <span class="n">porcentaje</span><span class="p">):</span>
            <span class="c1">#Return the values in a tuple like: return (increase, new_salary)</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">aumento</span><span class="p">(</span><span class="mi">30500</span><span class="p">,</span> <span class="mi">10</span><span class="p">),</span> <span class="p">(</span><span class="mi">3050</span><span class="p">,</span> <span class="mi">33550</span><span class="p">),</span> <span class="s2">&quot;Expected: (3050,33550)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">aumento</span><span class="p">(</span><span class="mi">10400</span><span class="p">,</span> <span class="mi">25</span><span class="p">),</span> <span class="p">(</span><span class="mi">2600</span><span class="p">,</span> <span class="mi">13000</span><span class="p">),</span> <span class="s2">&quot;Expected: (2600,13000)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">aumento</span><span class="p">(</span><span class="mi">50100</span><span class="p">,</span> <span class="mi">8</span><span class="p">),</span> <span class="p">(</span><span class="mi">4008</span><span class="p">,</span> <span class="mi">54108</span><span class="p">),</span> <span class="s2">&quot;Expected: (4008,54108)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">aumento</span><span class="p">(</span><span class="mi">25000</span><span class="p">,</span> <span class="mi">3</span><span class="p">),</span> <span class="p">(</span><span class="mi">750</span><span class="p">,</span> <span class="mi">25750</span><span class="p">),</span> <span class="s2">&quot;Expected: (750,25750)&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 125)</p>
<p>Duplicate ID – see quiz/Quiz1, line 129</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_5</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Request</span> <span class="n">the</span> <span class="n">price</span> <span class="kn">from</span> <span class="nn">a</span> <span class="n">merchant</span> <span class="ow">and</span> <span class="n">the</span> <span class="n">percentage</span> <span class="n">of</span> <span class="n">the</span> <span class="n">discount</span><span class="o">.</span>
    <span class="n">Display</span> <span class="n">the</span> <span class="n">discount</span> <span class="n">amount</span> <span class="ow">and</span> <span class="n">the</span> <span class="n">final</span> <span class="n">price</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">precio_con_descuento</span><span class="p">(</span><span class="n">precio</span><span class="p">,</span> <span class="n">porcentaje</span><span class="p">):</span>
            <span class="c1">#Return the values in a tuple like: return (discount, final_price)</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">precio_con_descuento</span><span class="p">(</span><span class="mi">100100</span><span class="p">,</span> <span class="mi">10</span><span class="p">),</span> <span class="p">(</span><span class="mi">10010</span><span class="p">,</span> <span class="mi">90090</span><span class="p">),</span> <span class="s2">&quot;Expected: (10010,90090)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">precio_con_descuento</span><span class="p">(</span><span class="mi">20523</span><span class="p">,</span> <span class="mi">4</span><span class="p">),</span> <span class="p">(</span><span class="mf">820.92</span><span class="p">,</span> <span class="mf">19702.08</span><span class="p">),</span> <span class="s2">&quot;Expected: (820.92,19702.08)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">precio_con_descuento</span><span class="p">(</span><span class="mi">55566</span><span class="p">,</span> <span class="mi">50</span><span class="p">),</span> <span class="p">(</span><span class="mi">27783</span><span class="p">,</span> <span class="mi">27783</span><span class="p">),</span> <span class="s2">&quot;Expected: (27783,27783)&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">precio_con_descuento</span><span class="p">(</span><span class="mi">75660</span><span class="p">,</span> <span class="mi">24</span><span class="p">),</span> <span class="p">(</span><span class="mf">18158.4</span><span class="p">,</span> <span class="mf">57501.6</span><span class="p">),</span> <span class="s2">&quot;Expected: (18158.4,57501.6)&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 154)</p>
<p>Duplicate ID – see quiz/Quiz1, line 158</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_6</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Calculate</span> <span class="n">the</span> <span class="n">duration</span> <span class="n">of</span> <span class="n">a</span> <span class="n">car</span> <span class="n">trip</span><span class="o">.</span> <span class="n">Ask</span> <span class="k">for</span> <span class="n">the</span> <span class="n">distance</span> <span class="n">to</span> <span class="n">be</span> <span class="n">covered</span>
    <span class="ow">and</span> <span class="n">the</span> <span class="n">expected</span> <span class="n">average</span> <span class="n">speed</span> <span class="k">for</span> <span class="n">the</span> <span class="n">trip</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">duration</span><span class="p">(</span><span class="n">distance</span><span class="p">,</span> <span class="n">speed</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">duration</span><span class="p">(</span><span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">),</span> <span class="mi">1</span><span class="p">,</span> <span class="s2">&quot;Expected: 1&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">duration</span><span class="p">(</span><span class="mi">100</span><span class="p">,</span> <span class="mi">3</span><span class="p">),</span> <span class="mi">100</span><span class="o">/</span><span class="mi">3</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(</span><span class="mi">100</span><span class="o">/</span><span class="mi">3</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">duration</span><span class="p">(</span><span class="mi">10500</span><span class="p">,</span> <span class="mi">30</span><span class="p">),</span> <span class="mi">350</span><span class="p">,</span> <span class="s2">&quot;Expected: 350&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">duration</span><span class="p">(</span><span class="mi">8600</span><span class="p">,</span> <span class="mi">50</span><span class="p">),</span> <span class="mi">172</span><span class="p">,</span> <span class="s2">&quot;Expected: 172&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">duration</span><span class="p">(</span><span class="mi">130</span><span class="p">,</span> <span class="mi">200</span><span class="p">),</span> <span class="mf">0.65</span><span class="p">,</span> <span class="s2">&quot;Expected: 0.65&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 183)</p>
<p>Duplicate ID – see quiz/Quiz1, line 187</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_7</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Convert</span> <span class="n">a</span> <span class="n">temperature</span> <span class="n">written</span> <span class="ow">in</span> <span class="n">Celsius</span> <span class="n">to</span> <span class="n">Fahrenheit</span><span class="o">.</span> <span class="n">F</span> <span class="o">=</span> <span class="p">(</span><span class="mi">9</span> <span class="o">*</span> <span class="n">C</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span> <span class="o">+</span> <span class="mi">32</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">celsius_to_fahrenheit</span><span class="p">(</span><span class="n">c</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">celsius_to_fahrenheit</span><span class="p">(</span><span class="mi">20</span><span class="p">),</span> <span class="p">((</span><span class="mi">9</span> <span class="o">*</span> <span class="mi">20</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="mi">32</span><span class="p">,</span> <span class="s2">&quot;Expected: 68&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">celsius_to_fahrenheit</span><span class="p">(</span><span class="mi">68</span><span class="p">),</span> <span class="p">((</span><span class="mi">9</span> <span class="o">*</span> <span class="mi">68</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="mi">32</span><span class="p">,</span> <span class="s2">&quot;Expected: 154.4&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">celsius_to_fahrenheit</span><span class="p">(</span><span class="mi">0</span><span class="p">),</span> <span class="p">((</span><span class="mi">9</span> <span class="o">*</span> <span class="mi">0</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="mi">32</span><span class="p">,</span> <span class="s2">&quot;Expected: 32&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">celsius_to_fahrenheit</span><span class="p">(</span><span class="o">-</span><span class="mi">10</span><span class="p">),</span> <span class="p">((</span><span class="mi">9</span> <span class="o">*</span> <span class="o">-</span><span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="mi">32</span><span class="p">,</span> <span class="s2">&quot;Expected: 14&quot;</span><span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">celsius_to_fahrenheit</span><span class="p">(</span><span class="o">-</span><span class="mi">24</span><span class="p">),</span> <span class="p">((</span><span class="mi">9</span> <span class="o">*</span> <span class="o">-</span><span class="mi">24</span><span class="p">)</span> <span class="o">/</span> <span class="mi">5</span><span class="p">)</span> <span class="o">+</span> <span class="mi">32</span><span class="p">,</span> <span class="s2">&quot;Expected: -11.2&quot;</span><span class="p">)</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 8" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 211)</p>
<p>Duplicate ID – see quiz/Quiz1, line 215</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_8</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Now</span> <span class="n">do</span> <span class="n">the</span> <span class="n">opposite</span><span class="p">,</span> <span class="kn">from</span> <span class="nn">Fahrenheit</span> <span class="n">to</span> <span class="n">Celsius</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">fahrenheit_to_celsius</span><span class="p">(</span><span class="n">f</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">fahrenheit_to_celsius</span><span class="p">(</span><span class="mi">21</span><span class="p">),</span> <span class="p">((</span><span class="mi">21</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">21</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">fahrenheit_to_celsius</span><span class="p">(</span><span class="mi">108</span><span class="p">),</span> <span class="p">((</span><span class="mi">108</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">108</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">fahrenheit_to_celsius</span><span class="p">(</span><span class="mi">0</span><span class="p">),</span> <span class="p">((</span><span class="mi">0</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">0</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">fahrenheit_to_celsius</span><span class="p">(</span><span class="o">-</span><span class="mi">10</span><span class="p">),</span> <span class="p">((</span><span class="o">-</span><span class="mi">10</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="o">-</span><span class="mi">10</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">fahrenheit_to_celsius</span><span class="p">(</span><span class="mi">14</span><span class="p">),</span> <span class="p">((</span><span class="mi">14</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">14</span> <span class="o">-</span> <span class="mi">32</span><span class="p">)</span> <span class="o">*</span> <span class="mi">5</span><span class="p">)</span> <span class="o">/</span> <span class="mi">9</span><span class="p">))</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 9" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 239)</p>
<p>Duplicate ID – see quiz/Quiz1, line 243</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q1_9
    :nocodelens:

    Write a program that asks for the number of kilometers traveled
    by a rented car, as well as the number of days the car has been rented for.
    Calculate the price to be paid, knowing that the car costs R $ 60.00
    per day and R $ 0.15 per kilometer traveled. |br|

    ~~~~
    def price(km, days):


    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(price(123, 3), (0.15 * 123) + (60 * 3), &quot;Expected: &quot; + str((0.15 * 123) + (60 * 3)))
            self.assertEqual(price(800, 4), (0.15 * 800) + (60 * 4), &quot;Expected: &quot; + str((0.15 * 800) + (60 * 4)))
            self.assertEqual(price(60, 1), (0.15 * 60) + (60 * 1), &quot;Expected: &quot; + str((0.15 * 60) + (60 * 1)))
            self.assertEqual(price(90, 2), (0.15 * 90) + (60 * 2), &quot;Expected: &quot; + str((0.15 * 90) + (60 * 2)))
            self.assertEqual(price(1016, 7), (0.15 * 1016) + (60 * 7), &quot;Expected: &quot; + str((0.15 * 1016) + (60 * 7)))


    myTests().main()
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 10" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 270)</p>
<p>Duplicate ID – see quiz/Quiz1, line 274</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">q1_10</span>
    <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

    <span class="n">Write</span> <span class="n">a</span> <span class="n">program</span> <span class="n">to</span> <span class="n">calculate</span> <span class="n">the</span> <span class="n">reduction</span> <span class="ow">in</span> <span class="n">the</span> <span class="n">lifespan</span> <span class="n">of</span> <span class="n">a</span> <span class="n">smoker</span><span class="o">.</span>
    <span class="n">Ask</span> <span class="k">for</span> <span class="n">the</span> <span class="n">number</span> <span class="n">of</span> <span class="n">cigarettes</span> <span class="n">smoked</span> <span class="n">per</span> <span class="n">day</span> <span class="ow">and</span> <span class="n">how</span> <span class="n">many</span> <span class="n">years</span> <span class="n">the</span> <span class="n">person</span> <span class="n">has</span> <span class="n">smoked</span><span class="o">.</span>
    <span class="n">Consider</span> <span class="n">that</span> <span class="n">a</span> <span class="n">smoker</span> <span class="n">loses</span> <span class="mi">10</span> <span class="n">minutes</span> <span class="n">of</span> <span class="n">life</span> <span class="k">for</span> <span class="n">each</span> <span class="n">cigarette</span><span class="p">,</span> <span class="n">calculate</span>
    <span class="n">how</span> <span class="n">many</span> <span class="n">days</span> <span class="n">a</span> <span class="n">smoker</span> <span class="n">will</span> <span class="n">lose</span><span class="o">.</span> <span class="n">Show</span> <span class="n">the</span> <span class="n">total</span> <span class="n">days</span><span class="o">.</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>

    <span class="o">~~~~</span>
    <span class="k">def</span> <span class="nf">smoker</span><span class="p">(</span><span class="n">cigarettes</span><span class="p">,</span> <span class="n">years</span><span class="p">):</span>


    <span class="o">====</span>
    <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>


    <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>
        <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>

            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span>
                    <span class="n">smoker</span><span class="p">(</span><span class="mi">10</span><span class="p">,</span> <span class="mi">1</span><span class="p">),</span> <span class="p">((</span><span class="mi">10</span> <span class="o">*</span> <span class="mi">1</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">10</span> <span class="o">*</span> <span class="mi">1</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">)</span>
            <span class="p">)</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">smoker</span><span class="p">(</span><span class="mi">3</span><span class="p">,</span> <span class="mi">5</span><span class="p">),</span> <span class="p">((</span><span class="mi">3</span> <span class="o">*</span> <span class="mi">5</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">3</span> <span class="o">*</span> <span class="mi">5</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">smoker</span><span class="p">(</span><span class="mi">1</span><span class="p">,</span> <span class="mi">8</span><span class="p">),</span> <span class="p">((</span><span class="mi">1</span> <span class="o">*</span> <span class="mi">8</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">1</span> <span class="o">*</span> <span class="mi">8</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">))</span>
            <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">smoker</span><span class="p">(</span><span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">),</span> <span class="p">((</span><span class="mi">2</span> <span class="o">*</span> <span class="mi">3</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">,</span> <span class="s2">&quot;Expected: &quot;</span> <span class="o">+</span> <span class="nb">str</span><span class="p">(((</span><span class="mi">2</span> <span class="o">*</span> <span class="mi">3</span> <span class="o">*</span> <span class="mi">365</span><span class="p">)</span> <span class="o">*</span> <span class="mi">10</span><span class="p">)</span> <span class="o">/</span> <span class="mi">1440</span><span class="p">))</span>


    <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 11" >
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz1_en.rst</span>, line 302)</p>
<p>Duplicate ID – see quiz/Quiz1, line 306</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>.. activecode:: q1_11
    :nocodelens:

    Knowing that ``str()`` converts numerical values to strings,
    calculate how many digits there are in 2 raised to a million. |br|

    ~~~~
    def digits():


    ====
    from unittest.gui import TestCaseGui


    class myTests(TestCaseGui):
        def testOne(self):

            self.assertEqual(digits(), 301030, &quot;Expected: 301030&quot;)


    myTests().main()
</pre></div>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>