<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 5 &#8212; PyZombis main@d3289e3</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=e476ce06ce356667798d"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 6" href="Quiz6_en.html" />
    <link rel="prev" title="Quiz - 4" href="Quiz4_en.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis main@d3289e3
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz5_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 4" data-toggle="tooltip">
    <a href="Quiz4_en.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 6' data-toggle="tooltip" >              
    <a href="Quiz6_en.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 4' data-toggle="tooltip">
  <a href="Quiz4_en.html" aria-label="Previous - Quiz - 4">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 6' data-toggle="tooltip" >
    <a href="Quiz6_en.html" aria-label="Next - Quiz - 6">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-5">
<h1>Quiz - 5<a class="headerlink" href="#quiz-5" title="Permalink to this headline">¶</a></h1>
<div id='quiz5_1' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1_1" >

<div class="runestone">
<div data-component="clickablearea" class="runestone alert alert-warning" id="q5_1_en" data-question_label="1"     style="visibility: hidden;">
<span data-question>Q-1: The following program declares a variable in the wrong place in the code. Your task is to select the incorrect declaration and select the correct line where it should go.</span><span data-feedback>Remember, you cannot use a variable without declaring it previously.</span><pre><span data-incorrect>x = 2</span>
<span data-correct>     </span>
<span data-incorrect>if y > 8 then:</span>
    <span data-correct>y = 5</span>
    <span data-incorrect>y = y * 2</span>

<span data-incorrect>else:</span>
    <span data-incorrect>x = x * 2</span>
    <span data-incorrect>print(x + y):</span></pre>

</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 1_2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_2_en data-question_label="2">
<div id=q5_2_en_question class="ac_question col-md-12">
<p>Translate the previous program (with the declaration in the right place) to Python using the following code block. <br /></p>

</div>
<textarea data-lang="python" id="q5_2_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">

====
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 1_3" >

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="3" id="q5_3_en"  style="visibility: hidden;">
            <p>Q-3: What does the previous program print?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [9, 9], "feedback": "<p>Your answer is correct, good job!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Your answer is incorrect, please try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
    </div>

<div id='quiz5_2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 2_1" >
<div class="highlight-none notranslate"><div class="highlight"><pre><span></span>for i = 1 to 9:
    if i != 3 then:
        for j = 1 to 6:
            print(&quot;Hello&quot;)
</pre></div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2_2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_4_en data-question_label="4">
<div id=q5_4_en_question class="ac_question col-md-12">
<p>Translate the previous program to Python using the following code block. <br />
<strong>Note</strong>: in our pseudolanguage, the loop includes the endpoints, that is, 1 to 4 means 1, 2, 3, 4. <br /></p>

</div>
<textarea data-lang="python" id="q5_4_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">

====
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2_3" >

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5" id="q5_5_en"  style="visibility: hidden;">
            <p>Q-5: How many times does the previous program print <code class="docutils literal notranslate"><span class="pre">&quot;Hello&quot;</span></code>?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [48, 48], "feedback": "<p>Your answer is correct, good job!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Your answer is incorrect, please try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
    </div>

<div id='quiz5_3' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 3_1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_6_en data-question_label="6">
<div id=q5_6_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">pares_divisibles_7</span></code> which receives two positive integers, <code class="docutils literal notranslate"><span class="pre">inicio</span></code> and <code class="docutils literal notranslate"><span class="pre">fin</span></code>, that represent the start and end of a range of numbers. The function must return the amount of numbers that are even and also divisible by 7. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(1,</span> <span class="pre">7)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(25,</span> <span class="pre">123)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">7</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(13,</span> <span class="pre">245)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">17</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q5_6_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def pares_divisibles_7(inicio, fin):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(pares_divisibles_7(1, 7), 0, "Expected: 0")
        self.assertEqual(pares_divisibles_7(25, 123), 7, "Expected: 7")
        self.assertEqual(pares_divisibles_7(13, 245), 17, "Expected: 17")
        self.assertEqual(pares_divisibles_7(1300, 2460), 83, "Expected: 83")
        self.assertEqual(pares_divisibles_7(1, 2460), 175, "Expected: 175")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3_2" >

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="7" id="q5_7_en"  style="visibility: hidden;">
            <p>Q-7: Use the function created in the previous exercise to answer How many numbers are even and also divisible by 7 between 1067 and 3627 (inclusive)?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [183, 183], "feedback": "<p>Your answer is correct, good job!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Your answer is incorrect, try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
    </div>

<div id='quiz5_4' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 4_1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_8_en data-question_label="8">
<div id=q5_8_en_question class="ac_question col-md-12">
<p>Daniela is a very superstitious person. For her, a number is <strong>lucky</strong> if it contains the digit <code class="docutils literal notranslate"><span class="pre">2</span></code> but not the <code class="docutils literal notranslate"><span class="pre">7</span></code>. She is also very curious and wants to know how many <strong>lucky</strong> numbers are in a range of numbers (including the ends)?. That’s why she has asked you to make the function called <code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte</span></code> which receives two positive integers, <code class="docutils literal notranslate"><span class="pre">inicio</span></code> and <code class="docutils literal notranslate"><span class="pre">fin</span></code>, that represent the start and end of a range of numbers. The function must return the amount of numbers that are <strong>lucky</strong>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(1,</span> <span class="pre">7)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(1,</span> <span class="pre">20)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(25,</span> <span class="pre">123)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">16</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(13,</span> <span class="pre">245)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">74</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q5_8_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def cuantos_tienen_suerte(inicio, fin):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(cuantos_tienen_suerte(1, 1), 0, "Expected: 0")
        self.assertEqual(cuantos_tienen_suerte(1, 7), 1, "Expected: 1")
        self.assertEqual(cuantos_tienen_suerte(1, 20), 3, "Expected: 3")
        self.assertEqual(cuantos_tienen_suerte(25, 123), 16, "Expected: 16")
        self.assertEqual(cuantos_tienen_suerte(1300, 2460), 481, "Expected: 481")
        self.assertEqual(cuantos_tienen_suerte(13, 245), 74, "Expected: 74")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4_2" >

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="9" id="q5_9_en"  style="visibility: hidden;">
            <p>Q-9: Use the function created in the previous exercise (<strong>MANDATORY</strong> to have finished it before and passed all tests) to answer How many <strong>lucky</strong> numbers are there between 18644 and 33087 (inclusive)?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [7995, 7995], "feedback": "<p>Your answer is correct, good job!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Your answer is incorrect, try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
    </div>

<div id='quiz5_5' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 5_1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_10_en data-question_label="10">
<div id=q5_10_en_question class="ac_question col-md-12">
<p>In the quiet rural village of <em>Ponteironuloville</em>, all phones have 6 digits. The telephone company establishes the following rules about numbers: | br |</p>
<ol class="arabic simple">
<li><p>There cannot be two identical consecutive digits, because this is boring</p></li>
<li><p>The sum of the digits must be even, because this is legal</p></li>
<li><p>The last digit cannot be the same as the first, because that is bad luck.</p></li>
</ol>
<p>Then, given these perfectly reasonable, well-designed and mature rules, develop the function called <code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville</span></code> that receives a string of positive integers and returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the number is valid according to the rules of <em>Ponteironuloville</em> and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.
| br | | br |</p>
<p>Examples: | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;123457&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;234562&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;222222&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;123456&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;312214&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> | br |
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;312312&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> | br |</p>

</div>
<textarea data-lang="python" id="q5_10_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def es_numero_ponteironuloville(numero):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(es_numero_ponteironuloville("123457"), True, "Expected: True")
        self.assertEqual(es_numero_ponteironuloville("234562"), False, "Expected: False")
        self.assertEqual(es_numero_ponteironuloville("304706"), True, "Expected: True")
        self.assertEqual(es_numero_ponteironuloville("222222"), False, "Expected: False")
        self.assertEqual(es_numero_ponteironuloville("123456"), False, "Expected: False")
        self.assertEqual(es_numero_ponteironuloville("312214"), False, "Expected: False")
        self.assertEqual(es_numero_ponteironuloville("312312"), True, "Expected: True")
        self.assertEqual(es_numero_ponteironuloville("131313"), True, "Expected: True")
        self.assertEqual(es_numero_ponteironuloville("249184"), True, "Expected: True")
        self.assertEqual(es_numero_ponteironuloville("012445"), False, "Expected: False")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5_2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_11_en data-question_label="11">
<div id=q5_11_en_question class="ac_question col-md-12">
<p>Knowing that the <code class="docutils literal notranslate"><span class="pre">split()</span></code> function divides a string into multiple strings, use the function created in the previous exercise (<strong>MANDATORY</strong> to have finished it before and passed all tests) to complete the <code class="docutils literal notranslate"><span class="pre">count_valid</span></code> function that returns the number of valid numbers from the list of numbers given in the following block of code. <br /> <br />
<strong>Note</strong>: the list of numbers is a string, wisely use the <code class="docutils literal notranslate"><span class="pre">split()</span></code> function to obtain a list of strings.</p>

</div>
<textarea data-lang="python" id="q5_11_en_editor" 
     data-include="q5_10_en" data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
numbers_list = """236043 237330 239636 240138 242123 246224 249183 252936
    254711 257200 257607 261424 263814 266794 268649 273050
    275001 277606 278997 283331 287104 287953 289137 291591
    292559 292946 295180 295566 297529 300400 304707 306931
    310638 313595 318449 319021 322082 323796 326266 326880
    327249 329914 334392 334575 336723 336734 338808 343269
    346040 350113 353631 357154 361633 361891 364889 365746
    365749 366426 369156 369444 369689 372896 374983 375223
    379163 380712 385640 386777 388599 389450 390178 392943
    394742 395921 398644 398832 401149 402219 405364 408088
    412901 417683 422267 424767 426613 430474 433910 435054
    440052 444630 447852 449116 453865 457631 461750 462985
    463328 466458 469601 473108 476773 477956 481991 482422
    486195 488359 489209 489388 491928 496569 496964 497901
    500877 502386 502715 507617 512526 512827 513796 518232
    521455 524277 528496 529345 531231 531766 535067 535183
    536593 537360 539055 540582 543708 547492 550779 551595
    556493 558807 559102 562050 564962 569677 570945 575447
    579937 580112 580680 582458 583012 585395 586244 587393
    590483 593112 593894 594293 597525 598184 600455 600953
    601523 605761 608618 609198 610141 610536 612636 615233
    618314 622752 626345 626632 628889 629457 629643 633673
    637656 641136 644176 644973 647617 652218 657143 659902
    662224 666265 668010 672480 672695 676868 677125 678315"""

def count_valid():
    count = 0
    """
        Write your code here, using the created variable
    """
    return count
====
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5_3" >

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="12" id="q5_12_en"  style="visibility: hidden;">
            <p>Q-12: Use the function created in the previous exercise to answer how many numbers from the list are valid according to the rules of <em>Ponteironuloville</em>?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [39, 39], "feedback": "<p>Your answer is correct, good job!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Your answer is incorrect, try again.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 4' data-toggle="tooltip">
  <a href="Quiz4_en.html" aria-label="Previous - Quiz - 4">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 6' data-toggle="tooltip" >
    <a href="Quiz6_en.html" aria-label="Next - Quiz - 6">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>