<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 2 &#8212; PyZombis aayush/issue#296@6f6cec7</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/runtime.bundle.js?v=e476ce06ce356667798d"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 3" href="Quiz3_en.html" />
    <link rel="prev" title="Quiz - 1" href="Quiz1_en.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis aayush/issue#296@6f6cec7
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz2_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 1" data-toggle="tooltip">
    <a href="Quiz1_en.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 3' data-toggle="tooltip" >              
    <a href="Quiz3_en.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 1' data-toggle="tooltip">
  <a href="Quiz1_en.html" aria-label="Previous - Quiz - 1">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 3' data-toggle="tooltip" >
    <a href="Quiz3_en.html" aria-label="Next - Quiz - 3">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-2">
<h1>Quiz - 2<a class="headerlink" href="#quiz-2" title="Permalink to this headline">¶</a></h1>
<div id='quiz2' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_1_en data-question_label="1">
<div id=q2_1_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">es_triangulo</span></code> that receives three positive integers <code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code>, and <code class="docutils literal notranslate"><span class="pre">c</span></code>. They represent the sides of a triangle. The function should verify that a triangle is formed with the given parameters. If the given parameters form a triangle, the function should return a string indicating its type, i.e., <code class="docutils literal notranslate"><span class="pre">&quot;Equilátero&quot;</span></code>, <code class="docutils literal notranslate"><span class="pre">&quot;Isósceles&quot;</span></code>, or <code class="docutils literal notranslate"><span class="pre">&quot;Escaleno&quot;</span></code>, otherwise, the function should return the string, <code class="docutils literal notranslate"><span class="pre">&quot;No</span> <span class="pre">es</span> <span class="pre">triángulo&quot;</span></code>.|br|
<strong>Note</strong>: remember that it is not a triangle when the longest side is greater than or equal to the sum of the other two. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">es_triangulo(2,</span> <span class="pre">2,</span> <span class="pre">2)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Equilátero&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_triangulo(3,</span> <span class="pre">2,</span> <span class="pre">2)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Isósceles&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_triangulo(4,</span> <span class="pre">2,</span> <span class="pre">6)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;No</span> <span class="pre">es</span> <span class="pre">triángulo&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_triangulo(2,</span> <span class="pre">1,</span> <span class="pre">8)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;No</span> <span class="pre">es</span> <span class="pre">triángulo&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q2_1_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def es_triangulo(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(es_triangulo(2, 2, 2), "Equilátero", "Expected: Equilátero")
        self.assertEqual(es_triangulo(2, 1, 2), "Isósceles", "Expected: Isósceles")
        self.assertEqual(es_triangulo(2, 1, 3), "No es triángulo", "Expected: No es triángulo")
        self.assertEqual(es_triangulo(2, 1, 8), "No es triángulo", "Expected: No es triángulo")
        self.assertEqual(es_triangulo(4, 2, 1), "No es triángulo", "Expected: No es triángulo")
        self.assertEqual(es_triangulo(4, 1000, 1000), "Isósceles", "Expected: Isósceles")
        self.assertEqual(es_triangulo(10000, 10000, 10000), "Equilátero", "Expected: Equilátero")
        self.assertEqual(es_triangulo(3, 2, 2), "Isósceles", "Expected: Isósceles")
        self.assertEqual(es_triangulo(10000, 1, 9999), "No es triángulo", "Expected: No es triángulo")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_2_en data-question_label="2">
<div id=q2_2_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">es_bisiesto</span></code> that receives the parameter <code class="docutils literal notranslate"><span class="pre">anio</span></code> which is a positive integer greater than zero and represents a year. The function should verify if the given parameter is a leap year, therefore, it should return <code class="docutils literal notranslate"><span class="pre">True</span></code> if it is, or <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. A year is a leap year if it is divisible by 400, or also if it is divisible by 4 but not divisible by 100. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">es_bisiesto(2014)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_bisiesto(2016)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_bisiesto(1900)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_bisiesto(2000)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q2_2_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def es_bisiesto(anio):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(es_bisiesto(2000), True, "Expected: True")
        self.assertEqual(es_bisiesto(2001), False, "Expected: False")
        self.assertEqual(es_bisiesto(2020), True, "Expected: True")
        self.assertEqual(es_bisiesto(2016), True, "Expected: True")
        self.assertEqual(es_bisiesto(2400), True, "Expected: True")
        self.assertEqual(es_bisiesto(1952), True, "Expected: True")
        self.assertEqual(es_bisiesto(1900), False, "Expected: False")
        self.assertEqual(es_bisiesto(2200), False, "Expected: False")
        self.assertEqual(es_bisiesto(2100), False, "Expected: False")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_3_en data-question_label="3">
<div id=q2_3_en_question class="ac_question col-md-12">
<p>Juan Pablo the fisherman, is a good man who works every day of the week. At the end of each work day, Juan must report to the state. Every time he brings a weight of fish greater than established by fishing regulations (50 kilograms) he must pay a fine of 4.00 units per additional kilogram. To keep track, he bought a computer to monitor his work income and asked you to make the function called <cite>generate_report</cite> that receives a list of positive floats <cite>weights</cite>, which represent the amount that Juan fished each day of the week. The function should return an array of tuples, where each tuple should be a pair of the form <cite>(additional, fine)</cite>, which represent the amount of additional kilograms that Juan fished during a day and the fine he had to pay respectively. If he did not pay a fine, the result will be <cite>0.0</cite> in both cases. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">generate_report([25.5,</span> <span class="pre">50.5,</span> <span class="pre">60.25,</span> <span class="pre">15,</span> <span class="pre">100,</span> <span class="pre">50,</span> <span class="pre">30.50])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">[(0.0,</span> <span class="pre">0.0),</span> <span class="pre">(0.5,</span> <span class="pre">2.0),</span> <span class="pre">(10.25,</span> <span class="pre">41.0),</span> <span class="pre">(0.0,</span> <span class="pre">0.0),</span> <span class="pre">(50.0,</span> <span class="pre">200.0),</span> <span class="pre">(0.0,</span> <span class="pre">0.0),</span> <span class="pre">(0.0,</span> <span class="pre">0.0)]</span></code></p>

</div>
<textarea data-lang="python" id="q2_3_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def generate_report(weights):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(
            generate_report([25.5, 50.5, 60.25, 15.0, 100.0, 50.0, 30.50]),
            [(0.0, 0.0), (0.5, 2.0), (10.25, 41.0), (0.0, 0.0), (50.0, 200.0), (0.0, 0.0), (0.0, 0.0)],
            "Expected: [(0.0, 0.0), (0.5, 2.0), (10.25, 41.0), (0.0, 0.0), (50.0, 200.0), (0.0, 0.0), (0.0, 0.0)]",
        )
        self.assertEqual(
            generate_report([5.55, 50.0, 10.25, 15.0, 0.0, 50.0, 51.0]),
            [(0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (1.0, 4.0)],
            "Expected: [(0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (0.0, 0.0), (1.0, 4.0)]",
        )


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_4_en data-question_label="4">
<div id=q2_4_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">mayor_tres</span></code> that receives three integers <code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code> and <code class="docutils literal notranslate"><span class="pre">c</span></code>. The function must return the highest of the three numbers, without using the functions <code class="docutils literal notranslate"><span class="pre">max</span></code> or <code class="docutils literal notranslate"><span class="pre">min</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">mayor_tres(5,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">5</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">mayor_tres(-5,</span> <span class="pre">0,</span> <span class="pre">-2)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q2_4_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def mayor_tres(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(mayor_tres(1, 2, 3), 3, "Expected: 3")
        self.assertEqual(mayor_tres(9, 2, 3), 9, "Expected: 9")
        self.assertEqual(mayor_tres(-9, 2, 1), 2, "Expected: 2")
        self.assertEqual(mayor_tres(1, 1, 1), 1, "Expected: 1")
        self.assertEqual(mayor_tres(5, 5, 4), 5, "Expected: 5")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_5_en data-question_label="5">
<div id=q2_5_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">mayor_menor_tres</span></code> that receives three integers <code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code> and <code class="docutils literal notranslate"><span class="pre">c</span></code>. The function must return a tuple of the form <code class="docutils literal notranslate"><span class="pre">(mayor,</span> <span class="pre">menor)</span></code> that represent the highest and the lowest of the three numbers, without using the functions <code class="docutils literal notranslate"><span class="pre">max</span></code> or <code class="docutils literal notranslate"><span class="pre">min</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">mayor_menor_tres(5,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(2,</span> <span class="pre">5)</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">mayor_menor_tres(-5,</span> <span class="pre">0,</span> <span class="pre">-2)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(-2,</span> <span class="pre">0)</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q2_5_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def mayor_menor_tres(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(mayor_menor_tres(1, 2, 3), (1, 3), "Expected: (1, 3)")
        self.assertEqual(mayor_menor_tres(9, 2, 3), (2, 9), "Expected: (2, 9)")
        self.assertEqual(mayor_menor_tres(-9, 2, 1), (-9, 2), "Expected: (-9, 2)")
        self.assertEqual(mayor_menor_tres(1, 1, 1), (1, 1), "se Esperado: (1, 1)")
        self.assertEqual(mayor_menor_tres(5, 5, 4), (4, 5), "Expected: (4, 5)")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_6_en data-question_label="6">
<div id=q2_6_en_question class="ac_question col-md-12">
<p>An employee of a company receives a monthly gross salary calculated by the amount of hours worked multiplied by its value. From this <code class="docutils literal notranslate"><span class="pre">salary</span></code>, the month’s <code class="docutils literal notranslate"><span class="pre">11%</span></code> for taxes, <code class="docutils literal notranslate"><span class="pre">8%</span></code> for health insurance, and <code class="docutils literal notranslate"><span class="pre">5%</span></code> for payment to the union are deducted. Develop the function <code class="docutils literal notranslate"><span class="pre">calcular_salario</span></code> that receives a float <code class="docutils literal notranslate"><span class="pre">valor_hora</span></code> and an integer <code class="docutils literal notranslate"><span class="pre">cantidad_horas</span></code> that represent how much he earns per hour and the amount of hours worked during the month. The function must return a dictionary of the form: <br /></p>
<p><code class="docutils literal notranslate"><span class="pre">{&quot;salario_bruto&quot;:</span> <span class="pre">A,</span> <span class="pre">&quot;impuestos&quot;:</span> <span class="pre">B,</span></code> <br /> <code class="docutils literal notranslate"><span class="pre">&quot;seguro_medico&quot;:</span> <span class="pre">C,</span> <span class="pre">&quot;sindicato&quot;:</span> <span class="pre">E,</span> <span class="pre">&quot;salario_neto&quot;:</span> <span class="pre">F}</span></code>. <br /></p>
<p>Where A, B, C, and D represent the amount of money corresponding to each item. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">calcular_salario(15.0,</span> <span class="pre">120)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">{&quot;salario_bruto&quot;:</span> <span class="pre">1800.00,</span> <span class="pre">&quot;impuestos&quot;:</span> <span class="pre">198.00,</span></code> <br /> <code class="docutils literal notranslate"><span class="pre">&quot;seguro_medico&quot;:</span> <span class="pre">144.0,</span> <span class="pre">&quot;sindicato&quot;:</span> <span class="pre">90.0,</span> <span class="pre">&quot;salario_neto&quot;:</span> <span class="pre">1368.00}</span></code>. <br /></p>

</div>
<textarea data-lang="python" id="q2_6_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def calcular_salario(valor_hora, cantidad_horas):


====
from unittest.gui import TestCaseGui
import random


class myTests(TestCaseGui):
    def get_expected_dictionary(self, salary):
        result = {
            "salario_bruto": salary,
            "impuestos": salary * 0.11,
            "seguro_medico": salary * 0.08,
            "sindicato": salary * 0.05,
        }
        result["salario_neto"] = (salary - result["impuestos"] -
            result["seguro_medico"] - result["sindicato"])
        return result

    def testOne(self):
        test_numbers = 10
        min_hour = 50
        max_hour = 200
        min_price_hour = 7
        max_price_hour = 30
        for test in range(test_numbers):
            hour = random.randint(min_hour, max_hour)
            price = round(random.uniform(min_price_hour, max_price_hour), 2)
            salary = price * hour
            expected = self.get_expected_dictionary(salary)
            result = calcular_salario(price, hour)
            current_test = test + 1

            self.assertEqual(
                round(result["salario_bruto"], 2),
                round(expected["salario_bruto"], 2),
                f"Test #{current_test} - Expected gross salary: {round(expected['salario_bruto'], 2)}",
            )
            self.assertEqual(
                round(result["impuestos"], 2),
                round(expected["impuestos"], 2),
                f"Test #{current_test} - Expected taxes: {round(expected['impuestos'], 2)}",
            )
            self.assertEqual(
                round(result["seguro_medico"], 2),
                round(expected["seguro_medico"], 2),
                f"Test #{current_test} - Expected health insurance: {round(expected['seguro_medico'], 2)}",
            )
            self.assertEqual(
                round(result["sindicato"], 2),
                round(expected["sindicato"], 2),
                f"Test #{current_test} - Expected union payment: {round(expected['sindicato'], 2)}",
            )
            self.assertEqual(
                round(result["salario_neto"], 2),
                round(expected["salario_neto"], 2),
                f"Test #{current_test} - Expected net salary: {round(expected['salario_neto'], 2)}",
            )


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 7" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q2_7_en data-question_label="7">
<div id=q2_7_en_question class="ac_question col-md-12">
<p>The paint sold at your trusted hardware store has a coverage of 1 liter per every 3 square meters and the paint is sold only in cans of 18 liters that cost each one <code class="docutils literal notranslate"><span class="pre">80.00</span></code> units. Develop the function <code class="docutils literal notranslate"><span class="pre">puedo_pintar</span></code> that receives an amount in square meters of an area to be painted as a positive integer <code class="docutils literal notranslate"><span class="pre">area</span></code>. The function should return a tuple with the amount of cans of paint that need to be bought to cover the entire area, as well as their total price, that is, using the form <code class="docutils literal notranslate"><span class="pre">(amount_cans,</span> <span class="pre">total_price)</span></code>. <br />
<strong>Note</strong>: only an integer number of cans of paint is sold. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">puedo_pintar(10)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(1,</span> <span class="pre">80.00)</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">puedo_pintar(100)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(2,</span> <span class="pre">160.00)</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">puedo_pintar(54)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(1,</span> <span class="pre">80.00)</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">puedo_pintar(55)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">(2,</span> <span class="pre">160.00)</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q2_7_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def puedo_pintar(area):

====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(puedo_pintar(10), (1, 80.00), "Expected: (1, 80.00)")
        self.assertEqual(puedo_pintar(100), (2, 160.00), "Expected: (2, 160.00)")
        self.assertEqual(puedo_pintar(54), (1, 80.00), "Expected: (1, 80.00)")
        self.assertEqual(puedo_pintar(55), (2, 160.00), "Expected: (2, 160.00)")
        self.assertEqual(puedo_pintar(1000), (19, 1520.00), "Expected: (19, 1520.00)")
        self.assertEqual(puedo_pintar(500), (10, 800.00), "Expected: (10, 800.00)")
        self.assertEqual(puedo_pintar(250), (5, 400.00), "Expected: (5, 400.00)")
        self.assertEqual(puedo_pintar(125), (3, 240.00), "Expected: (3, 240.00)")


myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 1' data-toggle="tooltip">
  <a href="Quiz1_en.html" aria-label="Previous - Quiz - 1">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 3' data-toggle="tooltip" >
    <a href="Quiz3_en.html" aria-label="Next - Quiz - 3">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>