<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 10 &#8212; PyZombis aayush/issue#331@d5fc5be</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/runtime.bundle.js?v=e476ce06ce356667798d"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=db4b57b8df6b9b8f470a"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 11" href="Quiz11_en.html" />
    <link rel="prev" title="Quiz - 9" href="Quiz9_en.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis aayush/issue#331@d5fc5be
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz10_en'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 9" data-toggle="tooltip">
    <a href="Quiz9_en.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 11' data-toggle="tooltip" >              
    <a href="Quiz11_en.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 9' data-toggle="tooltip">
  <a href="Quiz9_en.html" aria-label="Previous - Quiz - 9">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 11' data-toggle="tooltip" >
    <a href="Quiz11_en.html" aria-label="Next - Quiz - 11">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-10">
<h1>Quiz - 10<a class="headerlink" href="#quiz-10" title="Permalink to this headline">¶</a></h1>
<div id='quiz10' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Exercise 1" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_1_en data-question_label="1">
<div id=q10_1_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">cuantas_donas</span></code> that takes <code class="docutils literal notranslate"><span class="pre">n</span></code>, a positive integer, as a parameter,
and returns a string in the form of <code class="docutils literal notranslate"><span class="pre">&quot;Number</span> <span class="pre">of</span> <span class="pre">donuts:</span> <span class="pre">n&quot;</span></code>, where <code class="docutils literal notranslate"><span class="pre">n</span></code> is the value
passed to the function as an argument. However, if <code class="docutils literal notranslate"><span class="pre">n</span></code> &gt;= 10, <code class="docutils literal notranslate"><span class="pre">cuantas_donas</span></code> will
return <code class="docutils literal notranslate"><span class="pre">&quot;many&quot;</span></code> instead of <code class="docutils literal notranslate"><span class="pre">n</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">cuantas_donas(5)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Number</span> <span class="pre">of</span> <span class="pre">donuts:</span> <span class="pre">5&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantas_donas(23)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Number</span> <span class="pre">of</span> <span class="pre">donuts:</span> <span class="pre">many&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_1_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def cuantas_donas(n):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(cuantas_donas(4), "Number of donuts: 4", "Expected output: Number of donuts: 4")
        self.assertEqual(cuantas_donas(9), "Number of donuts: 9", "Expected output: Number of donuts: 9")
        self.assertEqual(
            cuantas_donas(10),
            "Number of donuts: many",
            "Expected output: Number of donuts: many",
        )
        self.assertEqual(
            cuantas_donas(99),
            "Number of donuts: many",
            "Expected output: Number of donuts: many",
        )


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 2" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_2_en data-question_label="2">
<div id=q10_2_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">cadena_de_extremos</span></code> that, given a string <code class="docutils literal notranslate"><span class="pre">s</span></code>,
returns a string with the first two and last two letters of <code class="docutils literal notranslate"><span class="pre">s</span></code>.
However, if the string has less than 2 letters, it returns an empty string. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">cadena_de_extremos(&quot;palmeras&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;paas&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cadena_de_extremos(&quot;a&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_2_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def cadena_de_extremos(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(cadena_de_extremos("palmeras"), "paas", "Expected output: paas")
        self.assertEqual(cadena_de_extremos("algoritmos"), "alos", "Expected output: alos")
        self.assertEqual(cadena_de_extremos("co"), "coco", "Expected output: coco")
        self.assertEqual(cadena_de_extremos("a"), "", "Expected output: ''")
        self.assertEqual(cadena_de_extremos("xyz"), "xyyz", "Expected output: xyyz")
        self.assertEqual(cadena_de_extremos(""), "", "Expected output: ''")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 3" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_3_en data-question_label="3">
<div id=q10_3_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter</span></code> that, given a string <code class="docutils literal notranslate"><span class="pre">s</span></code>,
returns a string in which all occurrences of the first character in <code class="docutils literal notranslate"><span class="pre">s</span></code>
are replaced by “*”, except for the first one. <strong>Note:</strong>
use the method <code class="docutils literal notranslate"><span class="pre">.replace(value_to_replace,</span> <span class="pre">new_value)</span></code> to solve the
exercise. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter(&quot;google&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;goo*le&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter(&quot;dona&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;dona&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_3_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def remplazar_primer_caracter(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(remplazar_primer_caracter("babble"), "ba**le", "Expected output: ba**le")
        self.assertEqual(remplazar_primer_caracter("aardvark"), "a*rdv*rk", "Expected output: a*rdv*rk")
        self.assertEqual(remplazar_primer_caracter("google"), "goo*le", "Expected output: goo*le")
        self.assertEqual(remplazar_primer_caracter("dona"), "dona", "Expected output: dona")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 4" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_4_en data-question_label="4">
<div id=q10_4_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas</span></code> that takes two strings as
arguments, <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code>, and returns a new string in the following way:</p>
<ul class="simple">
<li><p>The new string has to be a combination of <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code>.</p></li>
<li><p>The new string will have the form <code class="docutils literal notranslate"><span class="pre">&quot;&lt;a&gt;</span> <span class="pre">&lt;b&gt;&quot;</span></code>, note the space between both.</p></li>
<li><p>The new string will interchange the first two letters of <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code>.</p></li>
</ul>
<p>Suppose that <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">b</span></code> have more than 2 characters.
For better clarification, see the following examples. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas(&quot;mix&quot;,</span> <span class="pre">&quot;pod&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;pox</span> <span class="pre">mid&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas(&quot;pezzy&quot;,</span> <span class="pre">&quot;firm&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;fizzy</span> <span class="pre">perm&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_4_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def combinar_dos_cadenas(a, b):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(combinar_dos_cadenas("mix", "pod"), "pox mid", "Expected output: pox mid")
        self.assertEqual(combinar_dos_cadenas("dog", "dinner"), "dig donner", "Expected output: dig donner")
        self.assertEqual(
            combinar_dos_cadenas("gnash", "sport"),
            "spash gnort",
            "Expected output: spash gnort",
        )
        self.assertEqual(combinar_dos_cadenas("pezzy", "firm"), "fizzy perm", "Expected output: fizzy perm")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 5" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_5_en data-question_label="5">
<div id=q10_5_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">es_palindromo</span></code> that takes a string <code class="docutils literal notranslate"><span class="pre">s</span></code> as
parameter and checks if <code class="docutils literal notranslate"><span class="pre">s</span></code> is a palindrome or not, returning <code class="docutils literal notranslate"><span class="pre">True</span></code> or
<code class="docutils literal notranslate"><span class="pre">False</span></code> accordingly. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">es_palindromo(&quot;asa&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_palindromo(&quot;casa&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code>  <br /></p>

</div>
<textarea data-lang="python" id="q10_5_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def es_palindromo(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(es_palindromo("asa"), True, "Expected output: True")
        self.assertEqual(es_palindromo("casa"), False, "Expected output: False")
        self.assertEqual(es_palindromo("reconocer"), True, "Expected output: True")
        self.assertEqual(es_palindromo("palabra"), False, "Expected output: False")
        self.assertEqual(es_palindromo("radar"), True, "Expected output: True")
        self.assertEqual(es_palindromo("seres"), True, "Expected output: True")


myTests().main()
</textarea>
</div>
</div>
</div><div data-component="tab" data-tabname="Exercise 6" >

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_6_en data-question_label="6">
<div id=q10_6_en_question class="ac_question col-md-12">
<p>Develop the function <code class="docutils literal notranslate"><span class="pre">contar_ocurrencias</span></code> that takes two parameters:
<code class="docutils literal notranslate"><span class="pre">frase</span></code> and <code class="docutils literal notranslate"><span class="pre">palabra</span></code>, both of type string. The function should return
the number of times that <code class="docutils literal notranslate"><span class="pre">palabra</span></code> occurs in <code class="docutils literal notranslate"><span class="pre">frase</span></code>. <br /> <br />
Examples: <br />
<code class="docutils literal notranslate"><span class="pre">contar_ocurrencias(&quot;a</span> <span class="pre">ana</span> <span class="pre">y</span> <span class="pre">a</span> <span class="pre">mariana</span> <span class="pre">les</span> <span class="pre">gustan</span> <span class="pre">las</span> <span class="pre">manzanas&quot;,</span> <span class="pre">&quot;ana&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_6_en_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
       style="visibility: hidden;">
def contar_ocurrencias(frase, palabra):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(
            contar_ocurrencias("a ana y a mariana les gustan las manzanas", "ana"),
            3,
            "Expected output: 3",
        )
        self.assertEqual(contar_ocurrencias("Cats, rats, bats, and hats.", "ats"), 4, "Expected output: 4")


myTests().main()
</textarea>
</div>
</div>
</div>
    </div>

</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    <ul role="navigation" class="nextprev-list nextprev-list" aria-label="Change page">
<li id="relations-prev" class="navLink" title='Previous Section - Quiz - 9' data-toggle="tooltip">
  <a href="Quiz9_en.html" aria-label="Previous - Quiz - 9">
    <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
  </a>
</li>

  <li id="relations-next" class="navLink" title='Next Section - Quiz - 11' data-toggle="tooltip" >
    <a href="Quiz11_en.html" aria-label="Next - Quiz - 11">
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
    </a>
  </li>
</ul>

<script>
window.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement': 'right', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement': 'left', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>