<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 12 &#8212; PyZombis master@7e97442</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />



    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/runtime.bundle.js?v=f50940efe5177b9c81e6"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=02c928219c61d5ca913a"></script>
    <script src="../_static/runestone.bundle.js?v=2a0b78efc53bfc8dda40"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 13" href="Quiz13.html" />
    <link rel="prev" title="Quiz - 11" href="Quiz11.html" />


<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />


<script>
  eBookConfig = {};
  
    eBookConfig.useRunestoneServices = true;
    eBookConfig.host = 'http://127.0.0.1:8000' || 'http://127.0.0.1:8000';
    eBookConfig.app = eBookConfig.host+'/runestone';
    eBookConfig.course = 'PyZombis';
    eBookConfig.basecourse = 'PyZombis';
    eBookConfig.isLoggedIn = false;
    eBookConfig.enableCompareMe = eBookConfig.useRunestoneServices;
  
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.python3 = true;
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.enable_chatcodes = false ? false : false;

</script>

<!-- Ad Serving Headers Only serve ads to Anonymous Users -->



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse" aria-label="navbar toggle">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        
        
          <a class="brand-logo" href="/runestone/default/user/login" aria-label="Login"><img src="../_static/img/logo_small.png" alt=""></a>
        
        <a class="navbar-brand" href="../index.html" aria-label="index-page">
          
          PyZombis master@7e97442
        </a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>


        <!-- social media dropdown -->
        
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span class="visuallyhidden" aria-label="Social">Social</span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone">Follow @iRunestone</a>
                <a href="https://www.facebook.com/RunestoneInteractive">Our Facebook Page</a>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li><!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span class="visuallyhidden" aria-label="Search">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='../genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span class="visuallyhidden" aria-label="User">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='/runestone/assignments/chooseAssignment.html'>Assignments</a></li>
            <li><a href='/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            
            <li><a href='/runestone/default/courses'>Change Course</a></li>
			      <li class="divider"></li>
            <li id="ip_dropdown_link"><a href='/runestone/admin/index'>Instructor's Page</a></li>
            
            <li><a href='/runestone/dashboard/studentreport'>Progress Page</a></li>
            <li class="divider"></li>
            
            <li><a href="/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="/runestone/default/user/register" id="registerlink">Register</a></li>
            
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
			<!-- dark mode toggle -->
             <li><label class="theme-switch" for="checkbox">
					  <input type="checkbox" id="checkbox" onchange="runestoneComponents.switchTheme()" />
							<div class="slider round"></div>
					</label>Dark Mode</li>
			 </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:runestoneComponents.popupScratchAC()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Scratch Activecode" >Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        
        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span class="visuallyhidden" aria-label="Help">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='http://runestoneinteractive.org/pages/faq.html'>FAQ</a></li>
            <li><a href='https://runestone.academy/runestone/static/instructorguide/index.html'>Instructors Guide</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombis&page=Quiz12'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>
      
      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 11" data-toggle="tooltip">
    <a href="Quiz11.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 13' data-toggle="tooltip" >              
    <a href="Quiz13.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script>
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  document.addEventListener('load', (event) => {
    $('#relations-prev').tooltip(opts);
    $('#relations-next').tooltip(opts);
  });
</script>
        -->
        
          <li></li>
        
      </ul>
      </div> <!-- navbar-collapse -->
    </div> <!-- navbar -->
  </div> <!-- container -->


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content" role="main">

  
    
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - Quiz - 11' data-toggle="tooltip">
    <a href="Quiz11.html" aria-label="Quiz - 11">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - Quiz - 13' data-toggle="tooltip" >
    <a href="Quiz13.html" aria-label="Quiz - 13">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script>
document.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  

<!-- Ad Serving for Runestone Campaign -->


  
  <div class="section" id="quiz-12">
<h1>Quiz - 12<a class="headerlink" href="#quiz-12" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-1">
<h2>Ejercicio 1<a class="headerlink" href="#ejercicio-1" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_1 data-question_label="1">
<div id=q12_1_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">verbo</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> como parámetro. Si la longitud de la cadena es al menos 3, debe devolver la cadena original concatenando <code class="docutils literal notranslate"><span class="pre">&quot;ing&quot;</span></code> al final. Si la cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> ya termina en <code class="docutils literal notranslate"><span class="pre">&quot;ing&quot;</span></code>, concatene la cadena <code class="docutils literal notranslate"><span class="pre">&quot;ly&quot;</span></code>. Si la longitud de la cadena es menor que 3, devuelve la cadena original. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">verbo(&quot;singing&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;singingly&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">verbo(&quot;travel&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;traveling&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">verbo(&quot;do&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;do&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q12_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def verbo(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(verbo("hail"), "hailing", "Esperado: hailing")
        self.assertEqual(verbo("swiming"), "swimingly", "Esperado: swimingly")
        self.assertEqual(verbo("do"), "do", "Esperado: do")
        self.assertEqual(verbo("singing"), "singingly", "Esperado: singingly")
        self.assertEqual(verbo("travel"), "traveling", "Esperado: traveling")
        self.assertEqual(verbo("lly"), "llying", "Esperado: llying")
        self.assertEqual(verbo("ing"), "ingly", "Esperado: ingly")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-2">
<h2>Ejercicio 2<a class="headerlink" href="#ejercicio-2" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_2 data-question_label="2">
<div id=q12_2_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">no_es_malo</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> como parámetro. La función debe buscar la primera aparición de la cadena <code class="docutils literal notranslate"><span class="pre">&quot;no</span> <span class="pre">es&quot;</span></code> y la última aparición de la cadena <code class="docutils literal notranslate"><span class="pre">&quot;malo&quot;</span></code> o la cadena <code class="docutils literal notranslate"><span class="pre">&quot;mala&quot;</span></code>, si alguna de ellas aparece después de la primera, reemplace <code class="docutils literal notranslate"><span class="pre">&quot;no</span> <span class="pre">es&quot;</span> <span class="pre">...</span> <span class="pre">&quot;malo&quot;</span></code> o <code class="docutils literal notranslate"><span class="pre">&quot;no</span> <span class="pre">es&quot;</span> <span class="pre">...</span> <span class="pre">&quot;mala&quot;</span></code> por las cadenas <code class="docutils literal notranslate"><span class="pre">&quot;es</span> <span class="pre">bueno&quot;</span></code> o <code class="docutils literal notranslate"><span class="pre">&quot;es</span> <span class="pre">buena&quot;</span></code> respectivamente, luego devuelva el resultado. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">no_es_malo(&quot;El</span> <span class="pre">té</span> <span class="pre">no</span> <span class="pre">es</span> <span class="pre">malo&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;El</span> <span class="pre">té</span> <span class="pre">es</span> <span class="pre">bueno&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">no_es_malo(&quot;La</span> <span class="pre">película</span> <span class="pre">no</span> <span class="pre">es</span> <span class="pre">mala&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;La</span> <span class="pre">película</span> <span class="pre">es</span> <span class="pre">buena&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">no_es_malo(&quot;El</span> <span class="pre">precio</span> <span class="pre">de</span> <span class="pre">esta</span> <span class="pre">casa</span> <span class="pre">no</span> <span class="pre">es</span> <span class="pre">para</span> <span class="pre">nada</span> <span class="pre">malo&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;El</span> <span class="pre">precio</span> <span class="pre">de</span> <span class="pre">esta</span> <span class="pre">casa</span> <span class="pre">es</span> <span class="pre">bueno&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">no_es_malo(&quot;El</span> <span class="pre">teléfono</span> <span class="pre">es</span> <span class="pre">malo&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;El</span> <span class="pre">teléfono</span> <span class="pre">es</span> <span class="pre">malo&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q12_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def no_es_malo(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(
            no_es_malo("El televisor no es malo"),
            "El televisor es bueno",
            "Esperado: El televisor es bueno"
        )
        self.assertEqual(
            no_es_malo("El asado de la cena no es malo!"),
            "El asado de la cena es bueno!",
            "Esperado: El asado de la cena es bueno!"
        )
        self.assertEqual(
            no_es_malo("El té no está caliente"),
            "El té no está caliente",
            "Esperado: El té no está caliente"
        )
        self.assertEqual(
            no_es_malo("La película no es mala"),
            "La película es buena",
            "Esperado: La película es buena"
        )
        self.assertEqual(no_es_malo("no es para nada malo"), "es bueno", "Esperado: es bueno")
        self.assertEqual(no_es_malo("no es malo"), "es bueno", "Esperado: es bueno")
        self.assertEqual(no_es_malo("malo"), "malo", "Esperado: malo")
        self.assertEqual(no_es_malo("no"), "no", "Esperado: no")
        self.assertEqual(no_es_malo("NO"), "NO", "Esperado: NO")
        self.assertEqual(no_es_malo("MALO"), "MALO", "Esperado: MALO")
        self.assertEqual(no_es_malo("NO es MALO"), "NO es MALO", "Esperado: NO es MALO")
        self.assertEqual(no_es_malo("no es MALO"), "no es MALO", "Esperado: no es MALO")
        self.assertEqual(no_es_malo("NO es malo"), "NO es malo", "Esperado: NO es malo")
        self.assertEqual(no_es_malo("no es malo ni mala"), "es buena", "Esperado: es buena")
        self.assertEqual(no_es_malo("no es ni mala ni malo"), "es bueno", "Esperado: es bueno")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-3">
<h2>Ejercicio 3<a class="headerlink" href="#ejercicio-3" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_3 data-question_label="3">
<div id=q12_3_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">inicio_final</span></code> que recibe dos cadenas <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code>. Las cadenas tienen que ser dividas en dos, si alguna de las cadenas tiene un número impar de caracteres, la primera mitad será la subcadena más larga (por ejemplo <code class="docutils literal notranslate"><span class="pre">perro</span></code> se dividirá entre: <code class="docutils literal notranslate"><span class="pre">per</span></code> y <code class="docutils literal notranslate"><span class="pre">ro</span></code>). Dada las dos cadenas, devuelva una nueva cadena formada de la siguiente manera <code class="docutils literal notranslate"><span class="pre">a_inicio</span> <span class="pre">+</span> <span class="pre">b_inicio</span> <span class="pre">+</span> <span class="pre">a_final</span> <span class="pre">+</span> <span class="pre">b_final</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">inicio_final(&quot;abcd&quot;,</span> <span class="pre">&quot;1234&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;ab12cd34&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">inicio_final(&quot;abc&quot;,</span> <span class="pre">&quot;1234&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;ab12c34&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">inicio_final(&quot;abc&quot;,</span> <span class="pre">&quot;123&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;ab12c3&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q12_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def inicio_final(a, b):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(inicio_final("abcd", "xy"), "abxcdy", "Esperado: abxcdy")
        self.assertEqual(inicio_final("abcde", "xyz"), "abcxydez", "Esperado: abcxydez")
        self.assertEqual(inicio_final("a", "b"), "ab", "Esperado: ab")
        self.assertEqual(inicio_final("ac", "b"), "abc", "Esperado: abc")
        self.assertEqual(inicio_final("a", "bc"), "abc", "Esperado: abc")
        self.assertEqual(inicio_final("", ""), "", "Esperado: ''")
        self.assertEqual(inicio_final("a", ""), "a", "Esperado: 'a'")
        self.assertEqual(inicio_final("", "b"), "b", "Esperado: 'b'")
        self.assertEqual(
            inicio_final("Kitten", "Donut"),
            "KitDontenut",
            "Esperado: KitDontenut"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-4">
<h2>Ejercicio 4<a class="headerlink" href="#ejercicio-4" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_4 data-question_label="4">
<div id=q12_4_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">cuantos_ceros</span></code> que dado un entero <code class="docutils literal notranslate"><span class="pre">n</span></code> positivo, devuelva la cantidad de ceros al final del entero <code class="docutils literal notranslate"><span class="pre">n</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_ceros(10010)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_ceros(908007000)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q12_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def cuantos_ceros(n):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(cuantos_ceros(10100100010000), 4, "Esperado: 4")
        self.assertEqual(cuantos_ceros(90000000000000000010), 1, "Esperado: 1")
        self.assertEqual(cuantos_ceros(10), 1, "Esperado: 1")
        self.assertEqual(cuantos_ceros(1050051222), 0, "Esperado: 0")
        self.assertEqual(cuantos_ceros(1010101010), 1, "Esperado: 1")
        self.assertEqual(cuantos_ceros(5000), 3, "Esperado: 3")
        self.assertEqual(cuantos_ceros(10000000000), 10, "Esperado: 10")
        self.assertEqual(cuantos_ceros(555), 0, "Esperado: 0")
        self.assertEqual(cuantos_ceros(1), 0, "Esperado: 0")
        self.assertEqual(cuantos_ceros(0), 0, "Esperado: 0")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-5">
<h2>Ejercicio 5<a class="headerlink" href="#ejercicio-5" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_5 data-question_label="5">
<div id=q12_5_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">contar_2</span></code> que recibe un entero <code class="docutils literal notranslate"><span class="pre">n</span></code> positivo mayor que 0. La función debe devolver la cantidad de veces que el dígito 2 aparece en el intervalo <code class="docutils literal notranslate"><span class="pre">[0,</span> <span class="pre">n-1]</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">contar_2(20)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_2(5)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_2(1)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q12_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_2(n):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(contar_2(20), 2, "Esperado: 2")
        self.assertEqual(contar_2(1), 0, "Esperado: 0")
        self.assertEqual(contar_2(5), 1, "Esperado: 1")
        self.assertEqual(contar_2(999), 300, "Esperado: 300")
        self.assertEqual(contar_2(555), 216, "Esperado: 216")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-6">
<h2>Ejercicio 6<a class="headerlink" href="#ejercicio-6" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q12_6 data-question_label="6">
<div id=q12_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">inicio_potencia</span></code> que recibe un entero <code class="docutils literal notranslate"><span class="pre">n</span></code> positivo mayor que 0. La función debe devolver la primera potencia de 2 que comienza con <code class="docutils literal notranslate"><span class="pre">n</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">inicio_potencia(65)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">16</span></code> <br />
<em>Explicación</em>: para <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">=</span> <span class="pre">65</span></code> la potencia <code class="docutils literal notranslate"><span class="pre">2^16</span></code> da como resultado <code class="docutils literal notranslate"><span class="pre">65536</span></code> que contiene a <code class="docutils literal notranslate"><span class="pre">n</span></code> al comienzo. <br /> <br />
<code class="docutils literal notranslate"><span class="pre">inicio_potencia(4)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<em>Explicación</em>: para <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">=</span> <span class="pre">4</span></code> la potencia <code class="docutils literal notranslate"><span class="pre">2^2</span></code> da como resultado <code class="docutils literal notranslate"><span class="pre">4</span></code> que contiene a <code class="docutils literal notranslate"><span class="pre">n</span></code> al comienzo. <br /> <br />
<code class="docutils literal notranslate"><span class="pre">inicio_potencia(3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">5</span></code> <br />
<em>Explicación</em>: para <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">=</span> <span class="pre">3</span></code> la potencia <code class="docutils literal notranslate"><span class="pre">2^5</span></code> da como resultado <code class="docutils literal notranslate"><span class="pre">32</span></code> que contiene a <code class="docutils literal notranslate"><span class="pre">n</span></code> al comienzo. <br /></p>

</div>
<textarea data-lang="python" id="q12_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def inicio_potencia(n):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(inicio_potencia(7), 46, "Esperado: 46")
        self.assertEqual(inicio_potencia(3), 5, "Esperado: 5")
        self.assertEqual(inicio_potencia(133), 316, "Esperado: 316")
        self.assertEqual(inicio_potencia(1024), 10, "Esperado: 10")
        self.assertEqual(inicio_potencia(123), 90, "Esperado: 90")
        self.assertEqual(inicio_potencia(1), 0, "Esperado: 0")
        self.assertEqual(inicio_potencia(10), 10, "Esperado: 10")
        self.assertEqual(inicio_potencia(50), 102, "Esperado: 102")


myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  <div id="scprogresscontainer">
    You have attempted <span id="scprogresstotal"></span> of <span id="scprogressposs"></span> activities on this page <div id="subchapterprogress" aria-label="Page progress"></div>
  </div>
  
    
<ul>
  <li id="relations-prev" style="opacity:0.4" onmouseover="mouseOverPrev()" onmouseout="mouseOutPrev()"class="navLink" title='Previous Section - Quiz - 11' data-toggle="tooltip">
    <a href="Quiz11.html" aria-label="Quiz - 11">
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
            <i class='prevNav glyphicon glyphicon-chevron-left'  style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black;"></i>
        </div>
    </a>
  </li>
</ul>

  
<ul>
  <li id="relations-next" style="opacity:0.4" onmouseover="mouseOverNext()" onmouseout="mouseOutNext()" class="navLink" title='Next Section - Quiz - 13' data-toggle="tooltip" >
    <a href="Quiz13.html" aria-label="Quiz - 13">
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%; color:black; "></i>
      </div>
    </a>
  </li>
</ul>

<script>
document.addEventListener('load', (event) => {
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
});
</script>

<script>
	function mouseOverPrev() {
		document.getElementById("relations-prev").style.opacity=1;
	}

	function mouseOutPrev() {
		document.getElementById("relations-prev").style.opacity=0.4;
	}
	function mouseOverNext() {
		document.getElementById("relations-next").style.opacity=1;
	}

	function mouseOutNext() {
		document.getElementById("relations-next").style.opacity=0.4;
	}
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      
      <span id='numuserspan'></span><span class='loggedinuser'></span>
      
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>




<script> 
  window.addEventListener('load', (event) => {
    runestoneComponents.getSwitch(); 
  });
</script>


  </body>
</html>