<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 6 &#8212; PyZombis master@5267196</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.6.2" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=4011C77C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 7" href="Quiz7.html" />
    <link rel="prev" title="Quiz - 5" href="Quiz5mcq.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.6.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@5267196</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz6'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 6</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 5" data-toggle="tooltip">
    <a href="Quiz5mcq.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 7' data-toggle="tooltip" >              
    <a href="Quiz7.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-6">
<h1>Quiz - 6<a class="headerlink" href="#quiz-6" title="Permalink to this headline">¶</a></h1>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 10)</p>
<p>Duplicate ID – see lectures/Quiz6, line 9</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">dormir</span>
   <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

   <span class="n">A</span><span class="o">.</span> <span class="n">Dormir</span> <span class="p">:</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
   <span class="n">dia_semana</span> <span class="n">y</span> <span class="kc">True</span> <span class="n">para</span> <span class="n">días</span> <span class="n">de</span> <span class="n">semana</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">vacaciones</span> <span class="n">y</span> <span class="kc">True</span> <span class="n">en</span> <span class="n">vacaciones</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">puedes</span> <span class="n">quedarte</span> <span class="n">dormido</span> <span class="n">cuando</span> <span class="n">sea</span> <span class="n">festivo</span> <span class="n">o</span> <span class="n">no</span> <span class="n">entre</span> <span class="n">semana</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">devuelva</span> <span class="n">Verdadero</span> <span class="n">o</span> <span class="n">Falso</span> <span class="n">cuando</span> <span class="n">se</span> <span class="n">vaya</span> <span class="n">a</span> <span class="n">dormir</span> <span class="n">o</span> <span class="n">no</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>  <span class="n">br</span><span class="o">|</span>

   <span class="o">~~~~</span>
   <span class="k">def</span> <span class="nf">dormir</span><span class="p">(</span><span class="n">dia_semana</span><span class="p">,</span> <span class="n">feriado</span><span class="p">):</span>


   <span class="o">====</span>
   <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>

   <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>

       <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">dormir</span><span class="p">(</span><span class="kc">False</span><span class="p">,</span> <span class="kc">False</span><span class="p">),</span> <span class="kc">True</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: True esperado: True&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">dormir</span><span class="p">(</span><span class="kc">True</span><span class="p">,</span> <span class="kc">False</span><span class="p">),</span> <span class="kc">False</span><span class="p">,</span> <span class="s2">&quot;Excelente! obtenido: False esperado: False&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">dormir</span><span class="p">(</span><span class="kc">False</span><span class="p">,</span> <span class="kc">True</span><span class="p">),</span> <span class="kc">True</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: True esperado: True&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">dormir</span><span class="p">(</span><span class="kc">True</span><span class="p">,</span> <span class="kc">True</span><span class="p">),</span> <span class="kc">True</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: True esperado: True&quot;</span><span class="p">)</span>

   <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>
</pre></div>
</div>
</div>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 36)</p>
<p>Duplicate ID – see lectures/Quiz6, line 36</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">alumnos_problema</span>
   <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

   <span class="n">B</span><span class="o">.</span> <span class="n">problema_alumnos</span><span class="p">:</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">tenemos</span> <span class="n">dos</span> <span class="n">estudiantes</span> <span class="n">ayb</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">a_sorri</span> <span class="n">y</span> <span class="n">b_sorri</span> <span class="n">indican</span> <span class="n">si</span> <span class="n">ayb</span> <span class="n">sonríen</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">tenemos</span> <span class="n">problemas</span> <span class="n">cuando</span> <span class="n">ambos</span> <span class="n">sonríen</span> <span class="n">o</span> <span class="n">ambos</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">no</span> <span class="n">están</span> <span class="n">sonriendo</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">devuelve</span> <span class="kc">True</span> <span class="n">cuando</span> <span class="n">hay</span> <span class="n">problemas</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
   <span class="o">~~~~</span>
   <span class="k">def</span> <span class="nf">alumnos_problema</span><span class="p">(</span><span class="n">a_sonrrie</span><span class="p">,</span> <span class="n">b_sonrrie</span><span class="p">):</span>


   <span class="o">====</span>
   <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>

   <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>

       <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">alumnos_problema</span><span class="p">(</span><span class="kc">True</span><span class="p">,</span> <span class="kc">True</span><span class="p">),</span> <span class="kc">True</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: True esperado: True&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">alumnos_problema</span><span class="p">(</span><span class="kc">False</span><span class="p">,</span> <span class="kc">False</span><span class="p">),</span> <span class="kc">True</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: True esperado: True&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">alumnos_problema</span><span class="p">(</span><span class="kc">True</span><span class="p">,</span> <span class="kc">False</span><span class="p">),</span> <span class="kc">False</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: False esperado: False&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">alumnos_problema</span><span class="p">(</span><span class="kc">False</span><span class="p">,</span> <span class="kc">True</span><span class="p">),</span> <span class="kc">False</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: False esperado: False&quot;</span><span class="p">)</span>

   <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>

</pre></div>
</div>
</div>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 63)</p>
<p>Duplicate ID – see lectures/Quiz6, line 63</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">activecode</span><span class="p">::</span> <span class="n">suma_doble</span>
   <span class="p">:</span><span class="n">nocodelens</span><span class="p">:</span>

   <span class="n">C</span><span class="o">.</span> <span class="n">suma_doble</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">dados</span> <span class="n">dos</span> <span class="n">enteros</span> <span class="n">devuelve</span> <span class="n">su</span> <span class="n">suma</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">sin</span> <span class="n">embargo</span><span class="p">,</span> <span class="n">si</span> <span class="n">los</span> <span class="n">números</span> <span class="n">son</span> <span class="n">iguales</span><span class="p">,</span> <span class="n">devuelve</span> <span class="n">el</span> <span class="n">doble</span> <span class="n">de</span> <span class="n">la</span> <span class="n">suma</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">doble_</span> <span class="n">suma</span> <span class="p">(</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">)</span> <span class="o">-&gt;</span> <span class="mi">3</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>
    <span class="n">doble_</span> <span class="n">suma</span> <span class="p">(</span><span class="mi">2</span><span class="p">,</span> <span class="mi">2</span><span class="p">)</span> <span class="o">-&gt;</span> <span class="mi">8</span> <span class="o">|</span><span class="n">br</span><span class="o">|</span>


   <span class="o">~~~~</span>
   <span class="k">def</span> <span class="nf">suma_doble</span><span class="p">(</span><span class="n">a</span><span class="p">,</span> <span class="n">b</span><span class="p">):</span>


   <span class="o">====</span>
   <span class="kn">from</span> <span class="nn">unittest.gui</span> <span class="kn">import</span> <span class="n">TestCaseGui</span>

   <span class="k">class</span> <span class="nc">myTests</span><span class="p">(</span><span class="n">TestCaseGui</span><span class="p">):</span>

       <span class="k">def</span> <span class="nf">testOne</span><span class="p">(</span><span class="bp">self</span><span class="p">):</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">),</span> <span class="mi">3</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: 3 esperado: 3&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="mi">3</span><span class="p">,</span> <span class="mi">2</span><span class="p">),</span> <span class="mi">5</span><span class="p">,</span> <span class="s2">&quot;Excelente! obtenido: 5 esperado: 5&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="mi">2</span><span class="p">,</span> <span class="mi">2</span><span class="p">),</span> <span class="mi">8</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: 8 esperado: 8&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="o">-</span><span class="mi">1</span><span class="p">,</span> <span class="mi">0</span><span class="p">),</span> <span class="o">-</span><span class="mi">1</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: -1 esperado: -1&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="mi">0</span><span class="p">),</span> <span class="mi">0</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: 0 esperado: 0&quot;</span><span class="p">)</span>
           <span class="bp">self</span><span class="o">.</span><span class="n">assertEqual</span><span class="p">(</span><span class="n">suma_doble</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="mi">1</span><span class="p">),</span> <span class="mi">1</span><span class="p">,</span><span class="s2">&quot;Excelente! obtenido: 1 esperado: 1&quot;</span><span class="p">)</span>

   <span class="n">myTests</span><span class="p">()</span><span class="o">.</span><span class="n">main</span><span class="p">()</span>

</pre></div>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=diff21 data-question_label="1">
<div id=diff21_question class="ac_question col-md-12">
<p>D. diff21 <br />
dado un número entero n devuelve la diferencia absoluta entre n y 21 <br /></p>
<div class="system-message">
<p class="system-message-title">System Message: ERROR/3 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 98)</p>
<p>Unexpected indentation.</p>
</div>
<blockquote>
<div><p>sin embargo, si el número es mayor que 21, devuelve el doble de ladiferencia absoluta <br /></p>
</div></blockquote>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 99)</p>
<p>Block quote ends without a blank line; unexpected unindent.</p>
</div>
<p>diff21(19) -&gt; 2 <br />
diff21(25) -&gt; 8 <br />
dica: abs(x) devuelve el valor absoluto de x <br /></p>

</div>
<textarea data-lang="python" id="diff21_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def diff21(n):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(diff21(19),2,"Esperado: 2")
        self.assertEqual(diff21(10),11,"Esperado: 11")
        self.assertEqual(diff21(21),0,"Esperado: 0")
        self.assertEqual(diff21(22),2,"Esperado: 2")
        self.assertEqual(diff21(25),8,"Esperado: 8")
        self.assertEqual(diff21(30),18,"Esperado: 18")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=papagaio data-question_label="2">
<div id=papagaio_question class="ac_question col-md-12">
<ol class="upperalpha simple" start="5">
<li><p>loro <br /></p></li>
</ol>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 127)</p>
<p>Enumerated list ends without a blank line; unexpected unindent.</p>
</div>
<blockquote>
<div><p>tenemos un loro que habla fuerte <br />
hora es un parámetro entre 0 y 23 <br />
tenemos problemas si el loro esta hablando <br />
antes de las 7 o después de las 20 <br /></p>
</div></blockquote>

</div>
<textarea data-lang="python" id="papagaio_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def papagaio(falando, hora):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(loro(True, 6), True,"Esperado: True")
        self.assertEqual(loro(True, 7), False,"Esperado: False")
        self.assertEqual(loro(False, 6), False,"Esperado: False")
        self.assertEqual(loro(True, 21), True,"Esperado: True")
        self.assertEqual(loro(False, 21), False,"Esperado: False")
        self.assertEqual(loro(True, 23), True,"Esperado: True")
        self.assertEqual(loro(True, 20), False,"Esperado: False")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dez data-question_label="3">
<div id=dez_question class="ac_question col-md-12">
<ol class="upperalpha simple" start="6">
<li><p>dez | br |</p></li>
</ol>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 158)</p>
<p>Enumerated list ends without a blank line; unexpected unindent.</p>
</div>
<blockquote>
<div><p>dos enteros ayb <br />
devuelve Verdadero si uno de los dos es 10 o la suma es 10 <br /></p>
</div></blockquote>

</div>
<textarea data-lang="python" id="dez_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def dez(a,b):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(diez(9,10), True,"Esperado: True")
        self.assertEqual(diez(9,9), False,"Esperado: False")
        self.assertEqual(diez(1,9), True,"Esperado: True")
        self.assertEqual(diez(10,1), True,"Esperado: True")
        self.assertEqual(diez(10,10), True,"Esperado: True")
        self.assertEqual(diez(8,2), True,"Esperado: True")
        self.assertEqual(diez(8,3), False,"Esperado: False")
        self.assertEqual(diez(10,42), True,"Esperado: True") |br|
        self.assertEqual(diez(12,-2), True,"Esperado: True")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=dista10 data-question_label="4">
<div id=dista10_question class="ac_question col-md-12">
<ol class="upperalpha simple" start="7">
<li><p>dista10 <br /></p></li>
</ol>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 189)</p>
<p>Enumerated list ends without a blank line; unexpected unindent.</p>
</div>
<blockquote>
<div><p>ser un número entero n <br />
devuelve Verdadero si la diferencia absoluta entre n y 100 o n y 200 <br />
es menor o igual que 10 <br />
dista10 (93) -&gt; True <a href="#id1"><span class="problematic" id="id2">|en|</span></a>
dista10 (90) -&gt; True <a href="#id3"><span class="problematic" id="id4">|en|</span></a>
dista10 (89) -&gt; Falso <a href="#id5"><span class="problematic" id="id6">|en|</span></a></p>
</div></blockquote>

</div>
<textarea data-lang="python" id="dista10_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def dista10(n):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(dista10(93), True,"Esperado: True")
        self.assertEqual(dista10(90), True,"Esperado: True")
        self.assertEqual(dista10(89), False,"Esperado: False")
        self.assertEqual(dista10(110), True,"Esperado: True")
        self.assertEqual(dista10(111), False,"Esperado: False")
        self.assertEqual(dista10(121), False,"Esperado: False")
        self.assertEqual(dista10(0), False,"Esperado: False")
        self.assertEqual(dista10(5), False,"Esperado: False")
        self.assertEqual(dista10(191), True,"Esperado: True")
        self.assertEqual(dista10(189), False,"Esperado: False")
        self.assertEqual(dista10(190), True,"Esperado: True")
        self.assertEqual(dista10(200), True,"Esperado: True")
        self.assertEqual(dista10(210), True,"Esperado: True")
        self.assertEqual(dista10(211), False,"Esperado: False")
        self.assertEqual(dista10(290), False,"Esperado: False")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=apaga data-question_label="5">
<div id=apaga_question class="ac_question col-md-12">
<ol class="upperalpha simple" start="8">
<li><p>borra | br |</p></li>
</ol>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 230)</p>
<p>Enumerated list ends without a blank line; unexpected unindent.</p>
</div>
<blockquote>
<div><p>ser una cadena sy un entero n <br />
devuelve una nueva cadena sin la posición n <br />
delete (‘gatito’, 1) -&gt; ‘ktten’ <a href="#id7"><span class="problematic" id="id8">|en|</span></a>
delete (‘gatito’, 4) -&gt; ‘kittn’ <a href="#id9"><span class="problematic" id="id10">|en|</span></a></p>
</div></blockquote>

</div>
<textarea data-lang="python" id="apaga_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def apaga(s,n):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(apagar('kitten',1), "ktten","Esperado: ktten")
        self.assertEqual(apagar('kitten',0), "itten","Esperado: itten")
        self.assertEqual(apagar('kitten',4), "kittn","Esperado: kittn")
        self.assertEqual(apagar('Hi',0), "i","Esperado: i")
        self.assertEqual(apagar('Hi',1), "H","Esperado: H")
        self.assertEqual(apagar('code',0), "ode","Esperado: ode")
        self.assertEqual(apagar('code',1), "cde","Esperado: cde")
        self.assertEqual(apagar('code',2), "coe","Esperado: coe")
        self.assertEqual(apagar('code',3), "cod","Esperado: cod")
        self.assertEqual(apagar('chocolate',8), "chocolat","Esperado: chocolat")



myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=troca data-question_label="6">
<div id=troca_question class="ac_question col-md-12">
<ol class="upperroman simple">
<li><p>intercambio <br /></p></li>
</ol>
<div class="system-message">
<p class="system-message-title">System Message: WARNING/2 (<span class="docutils literal">/home/runner/work/PyZombis/PyZombis/_sources/quiz/Quiz6.rst</span>, line 265)</p>
<p>Enumerated list ends without a blank line; unexpected unindent.</p>
</div>
<blockquote>
<div><p>ser una cadena s | br |
si s tiene un tamaño &lt;= 1, se devuelve a sí mismo <br />
de lo contrario, cambie la primera y la última letra <br />
exchange (‘código’) -&gt; ‘eodc’ <br />
intercambio (‘a’) -&gt; ‘a’ <br />
intercambio (‘ab’) -&gt; ‘ba’ <br /></p>
</div></blockquote>

</div>
<textarea data-lang="python" id="troca_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def troca(s):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(intercambiar('code'), 'eodc',"Esperado: eodc")
        self.assertEqual(intercambiar('a'), 'a',"Esperado: a")
        self.assertEqual(intercambiar('ab'), 'ba',"Esperado: ba")
        self.assertEqual(intercambiar('abc'), 'cba',"Esperado: cba")
        self.assertEqual(intercambiar(' '), ' ',"Esperado: ' '")
        self.assertEqual(intercambiar('Chocolate'), 'ehocolatC',"Esperado: ehocolatC")
        self.assertEqual(intercambiar('nythoP'), 'Python',"Esperado: Python")
        self.assertEqual(intercambiar('hello'), 'oellh',"Esperado: oellh")


myTests().main()
</textarea>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 5' data-toggle="tooltip">
    <a href="Quiz5mcq.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz5mcq.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 7' data-toggle="tooltip" >
    <a href="Quiz7.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz7.html" >Next Section - Quiz - 7</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.6.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>