<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 10 &#8212; PyZombis master@dfc341b</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone-custom-sphinx-bootstrap.css?v=5.6.2" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/Blob.js?v=AEC2447"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../_static/runestone.js?v=4011C77C"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../_static/theme.js?v=D899FB97"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 11" href="Quiz11.html" />
    <link rel="prev" title="Quiz - 9" href="Quiz9.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.6.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@dfc341b</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz10'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 10</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 9" data-toggle="tooltip">
    <a href="Quiz9.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 11' data-toggle="tooltip" >              
    <a href="Quiz11.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-10">
<h1>Quiz - 10<a class="headerlink" href="#quiz-10" title="Permalink to this headline">¶</a></h1>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=near_ten data-question_label="1">
<div id=near_ten_question class="ac_question col-md-12">
<p>A. near_ten <br />
Seja um n não negativo, retorna True se o número está a distância de <br />
pelo menos dois de um múltiplo de dez. Use a função resto da divisão. <br />
near_ten(12) -&gt; True <br />
near_ten(17) -&gt; False <br />
near_ten(19) -&gt; True <br /></p>

</div>
<textarea data-lang="python" id="near_ten_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def near_ten(n):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(near_ten(12), True,"Esperado: True")
        self.assertEqual(near_ten(17), False,"Esperado: False")
        self.assertEqual(near_ten(19), True,"Esperado: True")
        self.assertEqual(near_ten(31), True,"Esperado: True")
        self.assertEqual(near_ten(6), False,"Esperado: False")
        self.assertEqual(near_ten(10), True,"Esperado: True")
        self.assertEqual(near_ten(11), True,"Esperado: True")
        self.assertEqual(near_ten(21), True,"Esperado: True")
        self.assertEqual(near_ten(22), True,"Esperado: True")
        self.assertEqual(near_ten(23), False,"Esperado: False")
        self.assertEqual(near_ten(54), False,"Esperado: False")
        self.assertEqual(near_ten(155), False,"Esperado: False")
        self.assertEqual(near_ten(158), True,"Esperado: True")
        self.assertEqual(near_ten(3), False,"Esperado: False")
        self.assertEqual(near_ten(1), True,"Esperado: True")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=lone_sum data-question_label="2">
<div id=lone_sum_question class="ac_question col-md-12">
<p>B. lone_sum <br />
Soma maluca: some os números inteiros a, b, e c <br />
Se algum número aparecer repetido ele não conta na soma <br />
lone_sum(1, 2, 3) -&gt; 6 <br />
lone_sum(3, 2, 3) -&gt; 2 <br />
lone_sum(3, 3, 3) -&gt; 0 <br /></p>

</div>
<textarea data-lang="python" id="lone_sum_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def lone_sum(a, b, c):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(lone_sum(1, 2, 3), 6,"Esperado: 6")
        self.assertEqual(lone_sum(3, 2, 3), 2,"Esperado: 2")
        self.assertEqual(lone_sum(3, 3, 3), 0,"Esperado: 0")
        self.assertEqual(lone_sum(9, 2, 2), 9,"Esperado: 9")
        self.assertEqual(lone_sum(2, 2, 9), 9,"Esperado: 9")
        self.assertEqual(lone_sum(2, 9, 2), 9,"Esperado: 9")
        self.assertEqual(lone_sum(2, 9, 3), 14,"Esperado: 14")
        self.assertEqual(lone_sum(4, 2, 3), 9,"Esperado: 9")
        self.assertEqual(lone_sum(1, 3, 1), 3,"Esperado: 3")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=luck_sum data-question_label="3">
<div id=luck_sum_question class="ac_question col-md-12">
<p>C. luck_sum <br />
Soma três inteiros a, b, c <br />
Se aparecer um 13 ele não conta e todos os da sua direita também <br />
lucky_sum(1, 2, 3) -&gt; 6 <br />
lucky_sum(1, 2, 13) -&gt; 3 <br />
lucky_sum(1, 13, 3) -&gt; 1 <br /></p>

</div>
<textarea data-lang="python" id="luck_sum_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def lucky_sum(a, b, c):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(lucky_sum(1, 2, 3), 6,"Esperado: 6")
        self.assertEqual(lucky_sum(1, 2, 13), 3,"Esperado: 3")
        self.assertEqual(lucky_sum(1, 13, 3), 1,"Esperado: 1")
        self.assertEqual(lucky_sum(1, 13, 13), 1,"Esperado: 1")
        self.assertEqual(lucky_sum(6, 5, 2), 13,"Esperado: 13")
        self.assertEqual(lucky_sum(13, 2, 3), 0,"Esperado: 0")
        self.assertEqual(lucky_sum(13, 2, 13), 0,"Esperado: 0")
        self.assertEqual(lucky_sum(13, 13, 2), 0,"Esperado: 0")
        self.assertEqual(lucky_sum(9, 4, 13), 13,"Esperado: 13")
        self.assertEqual(lucky_sum(8, 13, 2), 8,"Esperado: 8")
        self.assertEqual(lucky_sum(7, 2, 1), 10,"Esperado: 10")
        self.assertEqual(lucky_sum(3, 3, 13), 6,"Esperado: 6")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=double_char data-question_label="4">
<div id=double_char_question class="ac_question col-md-12">
<p>D. double_char <br />
retorna os caracteres da string original duplicados <br />
double_char(‘The’) -&gt; ‘TThhee’ <br />
double_char(‘AAbb’) -&gt; ‘AAAAbbbb’ <br />
double_char(‘Hi-There’) -&gt; ‘HHii–TThheerree’ <br /></p>

</div>
<textarea data-lang="python" id="double_char_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def double_char(s):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(double_char('The'), 'TThhee',"Esperado: TThhee")
        self.assertEqual(double_char('AAbb'), 'AAAAbbbb',"Esperado: AAAAbbbb")
        self.assertEqual(double_char('Hi-There'), 'HHii--TThheerree',"Esperado: HHii--TThheerree")
        self.assertEqual(double_char('Word!'), 'WWoorrdd!!',"Esperado: WWoorrdd!!")
        self.assertEqual(double_char('!!'), '!!!!',"Esperado: !!!!")
        self.assertEqual(double_char(''), '',"Esperado: ")
        self.assertEqual(double_char('a'), 'aa',"Esperado: aa")
        self.assertEqual(double_char('.'), '..',"Esperado: ..")
        self.assertEqual(double_char('aa'), 'aaaa',"Esperado: aaaa")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=count_hi data-question_label="5">
<div id=count_hi_question class="ac_question col-md-12">
<p>E. count_hi <br />
conta o número de vezes que aparece a string ‘hi’ <br />
count_hi(‘abc hi ho’) -&gt; 1 <br />
count_hi(‘ABChi hi’) -&gt; 2 <br />
count_hi(‘hihi’) -&gt; 2 <br /></p>

</div>
<textarea data-lang="python" id="count_hi_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def count_hi(s):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count_hi('abc hi ho'), 1,"Esperado: abc hi ho")
        self.assertEqual(count_hi('ABChi hi'), 2,"Esperado: ABChi hi")
        self.assertEqual(count_hi('hihi'), 2,"Esperado: hihi")
        self.assertEqual(count_hi('hiHIhi'), 2,"Esperado: hiHIhi")
        self.assertEqual(count_hi(''), 0,"Esperado: ")
        self.assertEqual(count_hi('h'), 0,"Esperado: h")
        self.assertEqual(count_hi('hi'), 1,"Esperado: hi")
        self.assertEqual(count_hi('Hi is no HI on ahI'), 0,"Esperado: Hi is no HI on ahI")
        self.assertEqual(count_hi('hiho not HOHIhi'), 2,"Esperado: hiho not HOHIhi")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=cat_dog data-question_label="6">
<div id=cat_dog_question class="ac_question col-md-12">
<p>F. cat_dog <br />
verifica se o aparece o mesmo número de vezes ‘cat’ e ‘dog’ <br />
cat_dog(‘catdog’) -&gt; True <br />
cat_dog(‘catcat’) -&gt; False <br />
cat_dog(‘1cat1cadodog’) -&gt; True <br /></p>

</div>
<textarea data-lang="python" id="cat_dog_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def cat_dog(s):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        seld.assertEqual(cat_dog('catdog'), True,"Esperado: True")
        seld.assertEqual(cat_dog('catcat'), False,"Esperado: False")
        seld.assertEqual(cat_dog('1cat1cadodog'), True,"Esperado: True")
        seld.assertEqual(cat_dog('catxxdogxxxdog'), False,"Esperado: False")
        seld.assertEqual(cat_dog('catxdogxdogxcat'), True,"Esperado: True")
        seld.assertEqual(cat_dog('catxdogxdogxca'), False,"Esperado: False")
        seld.assertEqual(cat_dog('dogdogcat'), False,"Esperado: False")
        seld.assertEqual(cat_dog('dogogcat'), True,"Esperado: True")
        seld.assertEqual(cat_dog('dog'), False,"Esperado: False")
        seld.assertEqual(cat_dog('cat'), False,"Esperado: False")
        seld.assertEqual(cat_dog('ca'), True,"Esperado: True")
        seld.assertEqual(cat_dog('c'), True,"Esperado: True")
        seld.assertEqual(cat_dog(''), True,"Esperado: True")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=count_code data-question_label="7">
<div id=count_code_question class="ac_question col-md-12">
<p>G. count_code <br />
conta quantas vezes aparece ‘code’ <br />
a letra ‘d’ pode ser trocada por outra qualquer <br />
assim ‘coxe’ ou ‘coze’ também são contadas como ‘code’ <br />
count_code(‘aaacodebbb’) -&gt; 1 <br />
count_code(‘codexxcode’) -&gt; 2 <br />
count_code(‘cozexxcope’) -&gt; 2 <br /></p>

</div>
<textarea data-lang="python" id="count_code_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def count_code(s):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count_code('aaacodebbb'), 1,"Esperado: 1")
        self.assertEqual(count_code('codexxcode'), 2,"Esperado: 2")
        self.assertEqual(count_code('cozexxcope'), 2,"Esperado: 2")
        self.assertEqual(count_code('cozfxxcope'), 1,"Esperado: 1")
        self.assertEqual(count_code('xxcozeyycop'), 1,"Esperado: 1")
        self.assertEqual(count_code('cozcop'), 0,"Esperado: 0")
        self.assertEqual(count_code('abcxyz'), 0,"Esperado: 0")
        self.assertEqual(count_code('code'), 1,"Esperado: 1")
        self.assertEqual(count_code('ode'), 0,"Esperado: 0")
        self.assertEqual(count_code('c'), 0,"Esperado: 0")
        self.assertEqual(count_code(''), 0,"Esperado: 0")
        self.assertEqual(count_code('AAcodeBBcoleCCccoreDD'), 3,"Esperado: 3")
        self.assertEqual(count_code('AAcodeBBcoleCCccorfDD'), 2,"Esperado: 2")
        self.assertEqual(count_code('coAcodeBcoleccoreDD'), 3,"Esperado: 3")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=end_other data-question_label="8">
<div id=end_other_question class="ac_question col-md-12">
<p>H. end_other <br />
as duas strings devem ser convertidas para minúsculo via lower() <br />
depois disso verifique que no final da string b ocorre a string a <br />
ou se no final da string a ocorre a string b <br />
end_other(‘Hiabc’, ‘abc’) -&gt; True <br />
end_other(‘AbC’, ‘HiaBc’) -&gt; True <br />
end_other(‘abc’, ‘abXabc’) -&gt; True <br /></p>

</div>
<textarea data-lang="python" id="end_other_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def end_other(a, b):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(end_other('Hiabc', 'abc'), True,"Esperado: True")
        self.assertEqual(end_other('AbC', 'HiaBc'), True,"Esperado: True")
        self.assertEqual(end_other('abc', 'abXabc'), True,"Esperado: True")
        self.assertEqual(end_other('Hiabc', 'abcd'), False,"Esperado: False")
        self.assertEqual(end_other('Hiabc', 'bc'), True,"Esperado: True")
        self.assertEqual(end_other('Hiabcx', 'bc'), False,"Esperado: False")
        self.assertEqual(end_other('abc', 'abc'), True,"Esperado: True")
        self.assertEqual(end_other('xyz', '12xyz'), True,"Esperado: True")
        self.assertEqual(end_other('yz', '12xz'), False,"Esperado: False")
        self.assertEqual(end_other('Z', '12xz'), True,"Esperado: True")
        self.assertEqual(end_other('12', '12'), True,"Esperado: True")
        self.assertEqual(end_other('abcXYZ', 'abcDEF'), False,"Esperado: False")
        self.assertEqual(end_other('ab', 'ab12'), False,"Esperado: False")
        self.assertEqual(end_other('ab', '12ab'), True,"Esperado: True")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=count_evens data-question_label="9">
<div id=count_evens_question class="ac_question col-md-12">
<p>I. count_evens <br />
conta os números pares da lista <br />
count_evens([2, 1, 2, 3, 4]) -&gt; 3 <br />
count_evens([2, 2, 0]) -&gt; 3 <br />
count_evens([1, 3, 5]) -&gt; 0 <br /></p>

</div>
<textarea data-lang="python" id="count_evens_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def count_evens(nums):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count_evens([2, 1, 2, 3, 4]), 3, "Esperado: 3")
        self.assertEqual(count_evens([2, 2, 0]), 3, "Esperado: 3")
        self.assertEqual(count_evens([1, 3, 5]), 0, "Esperado: 0")
        self.assertEqual(count_evens([]), 0, "Esperado: 0")
        self.assertEqual(count_evens([11, 9, 0, 1]), 1, "Esperado: 1")
        self.assertEqual(count_evens([2, 11, 9, 0]), 2, "Esperado: 2")
        self.assertEqual(count_evens([2]), 1, "Esperado: 1")
        self.assertEqual(count_evens([2, 5, 12]), 2, "Esperado: 2")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=sum13 data-question_label="10">
<div id=sum13_question class="ac_question col-md-12">
<p>J. sum13 <br />
retorna a soma dos números de uma lista <br />
13 dá azar, você deverá ignorar o 13 e todos os números à sua direita <br />
sum13([1, 2, 2, 1]) -&gt; 6 <br />
sum13([1, 1]) -&gt; 2 <br />
sum13([1, 2, 2, 1, 13]) -&gt; 6 <br />
sum13([13, 1, 2, 3, 4]) -&gt; 0 <br /></p>

</div>
<textarea data-lang="python" id="sum13_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def sum13(nums):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum13([1, 2, 2, 1]), 6,"Esperado: 6")
        self.assertEqual(sum13([1, 1]), 2,"Esperado: 2")
        self.assertEqual(sum13([1, 2, 2, 1, 13]), 6,"Esperado: 6")
        self.assertEqual(sum13([1, 2, 13, 2, 1, 13]), 3,"Esperado: 3")
        self.assertEqual(sum13([13, 1, 2, 13, 2, 1, 13]), 0,"Esperado: 0")
        self.assertEqual(sum13([]), 0,"Esperado: 0")
        self.assertEqual(sum13([13]), 0,"Esperado: 0")
        self.assertEqual(sum13([13, 13]), 0,"Esperado: 0")
        self.assertEqual(sum13([13, 0, 13]), 0,"Esperado: 0")
        self.assertEqual(sum13([13, 1, 13]), 0,"Esperado: 0")
        self.assertEqual(sum13([5, 7, 2]), 14,"Esperado: 14")
        self.assertEqual(sum13([5, 13, 2]), 5,"Esperado: 5")
        self.assertEqual(sum13([0]), 0,"Esperado: 0")
        self.assertEqual(sum13([13, 0]), 0,"Esperado: 0")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=has22 data-question_label="11">
<div id=has22_question class="ac_question col-md-12">
<p>K. has22 <br />
Verifica se na lista de números inteiros aparecem dois 2 consecutivos <br />
has22([1, 2, 2]) -&gt; True <br />
has22([1, 2, 1, 2]) -&gt; False <br />
has22([2, 1, 2]) -&gt; False <br /></p>

</div>
<textarea data-lang="python" id="has22_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def has22(nums):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(has22([1, 2, 2]), True,"Esperado: True")
        self.assertEqual(has22([1, 2, 1, 2]), False,"Esperado: False")
        self.assertEqual(has22([2, 1, 2]), False,"Esperado: False")
        self.assertEqual(has22([2, 2, 1, 2]), True,"Esperado: True")
        self.assertEqual(has22([1, 3, 2]), False,"Esperado: False")
        self.assertEqual(has22([1, 3, 2, 2]), True,"Esperado: True")
        self.assertEqual(has22([2, 3, 2, 2]), True,"Esperado: True")
        self.assertEqual(has22([4, 2, 4, 2, 2, 5]), True,"Esperado: True")
        self.assertEqual(has22([1, 2]), False,"Esperado: False")
        self.assertEqual(has22([2, 2]), True,"Esperado: True")
        self.assertEqual(has22([2]), False,"Esperado: False")
        self.assertEqual(has22([]), False,"Esperado: False")
        self.assertEqual(has22([3, 3, 2, 2]), True,"Esperado: True")
        self.assertEqual(has22([5, 2, 5, 2]), False,"Esperado: False")


myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=soma_na_lista data-question_label="12">
<div id=soma_na_lista_question class="ac_question col-md-12">
<p>L. soma_na_lista <br />
Verifica se n é soma de dois números distintos da lista <br />
soma_na_lista(5, [1, 2, 3, 4]) -&gt; True <br />
soma_na_lista(9, [1, 2, 3, 4]) -&gt; False <br />
soma_na_lista(0, [1, 2, 3, 4]) -&gt; False <br />
soma_na_lista(8, [1, 2, 3, 4]) -&gt; False <br />
soma_na_lista(4, [2, 2, 2, 2]) -&gt; False <br />
soma_na_lista(4, [2, 2, 1, 3]) -&gt; True <br />
Hall of Fame: Raphael Castro 1a Turma Python para Zumbis <br /></p>

</div>
<textarea data-lang="python" id="soma_na_lista_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def soma_na_lista(n, lista):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(soma_na_lista(5, [1, 2, 3, 4]), True,"Esperado: True")
        self.assertEqual(soma_na_lista(9, [1, 2, 3, 4]), False,"Esperado: False")
        self.assertEqual(soma_na_lista(0, [1, 2, 3, 4]), False,"Esperado: False")
        self.assertEqual(soma_na_lista(8, [1, 2, 3, 4]), False,"Esperado: False")
        self.assertEqual(soma_na_lista(4, [2, 2, 2, 2]), False,"Esperado: False")
        self.assertEqual(soma_na_lista(4, [2, 2, 1, 3]), True,"Esperado: True")
        self.assertEqual(soma_na_lista(42, [40, 2, 3, 39]), True,"Esperado: True")

myTests().main()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=M._desafio! data-question_label="13">
<div id=M._desafio!_question class="ac_question col-md-12">
<p>M. desafio! faça somente se já tiver acabado o EP1 e todas as listas <br />
Fila de tijolos sem usar loops <br />
queremos montar uma fila de tijolos de um tamanho denominado meta <br />
temos tijolos pequenos (tamanho 1) e tijolos grandes (tamanho 5) <br />
retorna True se for possível montar a fila de tijolos <br />
é possível uma solução sem usar for ou while <br />
fila_tijolos(3, 1, 8) -&gt; True <br />
fila_tijolos(3, 1, 9) -&gt; False <br />
fila_tijolos(3, 2, 10) -&gt; True <br /></p>

</div>
<textarea data-lang="python" id="M._desafio!_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def fila_tijolos(n_peq, n_gra, meta):


====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(fila_tijolos(3, 1, 8), True,"Esperado: True")
        self.assertEqual(fila_tijolos(3, 1, 9), False,"Esperado: False")
        self.assertEqual(fila_tijolos(3, 2, 10), True,"Esperado: True")
        self.assertEqual(fila_tijolos(3, 2, 8), True,"Esperado: True")
        self.assertEqual(fila_tijolos(3, 2, 9), False,"Esperado: False")
        self.assertEqual(fila_tijolos(6, 1, 11), True,"Esperado: True")
        self.assertEqual(fila_tijolos(6, 0, 11), False,"Esperado: False")
        self.assertEqual(fila_tijolos(3, 1, 7), True,"Esperado: True")
        self.assertEqual(fila_tijolos(1, 1, 7), False,"Esperado: False")
        self.assertEqual(fila_tijolos(2, 1, 7), True,"Esperado: True")
        self.assertEqual(fila_tijolos(7, 1, 11), True,"Esperado: True")
        self.assertEqual(fila_tijolos(7, 1, 8), True,"Esperado: True")
        self.assertEqual(fila_tijolos(7, 1, 13), False,"Esperado: False")
        self.assertEqual(fila_tijolos(43, 1, 46), True,"Esperado: True")
        self.assertEqual(fila_tijolos(40, 1, 46), False,"Esperado: False")
        self.assertEqual(fila_tijolos(22, 2, 33), False,"Esperado: False")
        self.assertEqual(fila_tijolos(0, 2, 10), True,"Esperado: True")
        self.assertEqual(fila_tijolos(1000000, 1000, 1000100), True,"Esperado: True")
        self.assertEqual(fila_tijolos(2, 1000000, 100003), False,"Esperado: False")
        self.assertEqual(fila_tijolos(12, 2, 21), True,"Esperado: True")

myTests().main()
</textarea>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 9' data-toggle="tooltip">
    <a href="Quiz9.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz9.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 11' data-toggle="tooltip" >
    <a href="Quiz11.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz11.html" >Next Section - Quiz - 11</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.6.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>