<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 11 &#8212; PyZombis master@92ad1e7</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=2bcd2c19954c99ce4506"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 12" href="Quiz12.html" />
    <link rel="prev" title="Quiz - 10" href="Quiz10.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.7.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@92ad1e7</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz11'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 11</a><ul>
<li><a class="reference internal" href="#ejercicio-1">Ejercicio 1</a></li>
<li><a class="reference internal" href="#ejercicio-2">Ejercicio 2</a></li>
<li><a class="reference internal" href="#ejercicio-3">Ejercicio 3</a></li>
<li><a class="reference internal" href="#ejercicio-4">Ejercicio 4</a></li>
<li><a class="reference internal" href="#ejercicio-5">Ejercicio 5</a></li>
<li><a class="reference internal" href="#ejercicio-6">Ejercicio 6</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 10" data-toggle="tooltip">
    <a href="Quiz10.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 12' data-toggle="tooltip" >              
    <a href="Quiz12.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-11">
<h1>Quiz - 11<a class="headerlink" href="#quiz-11" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-1">
<h2>Ejercicio 1<a class="headerlink" href="#ejercicio-1" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_1 data-question_label="1">
<div id=q11_1_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">cuantas_donas</span></code> que toma a <code class="docutils literal notranslate"><span class="pre">n</span></code>, un número entero
positivo, como parámetro, y devuelve una cadena de forma <code class="docutils literal notranslate"><span class="pre">&quot;Número</span> <span class="pre">de</span> <span class="pre">donas:</span> <span class="pre">n&quot;</span></code>,
donde <code class="docutils literal notranslate"><span class="pre">n</span></code> es el valor que se le pasó a la función como argumento. No obstante,
si <code class="docutils literal notranslate"><span class="pre">n</span></code> &gt;= 10, <code class="docutils literal notranslate"><span class="pre">cuantas_donas</span></code> devolverá <code class="docutils literal notranslate"><span class="pre">&quot;muchas&quot;</span></code> en vez de <code class="docutils literal notranslate"><span class="pre">n</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">cuantas_donas(5)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Número</span> <span class="pre">de</span> <span class="pre">donas:</span> <span class="pre">5&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantas_donas(23)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Número</span> <span class="pre">de</span> <span class="pre">donas:</span> <span class="pre">muchas&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q11_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def cuantas_donas(n):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(cuantas_donas(4), "Número de donas: 4", "Esperado: Número de donas: 4")
        self.assertEqual(cuantas_donas(9), "Número de donas: 9", "Esperado: Número de donas: 9")
        self.assertEqual(
            cuantas_donas(10),
            "Número de donas: muchas",
            "Esperado: Número de donas: muchas",
        )
        self.assertEqual(
            cuantas_donas(99),
            "Número de donas: muchas",
            "Esperado: Número de donas: muchas",
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-2">
<h2>Ejercicio 2<a class="headerlink" href="#ejercicio-2" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_2 data-question_label="2">
<div id=q11_2_question class="ac_question col-md-12">
<p>Desarrolla la función <code class="docutils literal notranslate"><span class="pre">cadena_de_extremos</span></code> que, dada una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code>,
devuelva una cadena con las dos primeras y las dos últimas letras de <code class="docutils literal notranslate"><span class="pre">s</span></code>.
Sin embargo, si la cadena tiene menos de 2 letras, devuelve una cadena vacía. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">cadena_de_extremos(&quot;palmeras&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;paas&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cadena_de_extremos(&quot;a&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q11_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def cadena_de_extremos(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(cadena_de_extremos("palmeras"), "paas", "Esperado: paas")
        self.assertEqual(cadena_de_extremos("algoritmos"), "alos", "Esperado: alos")
        self.assertEqual(cadena_de_extremos("co"), "coco", "Esperado: coco")
        self.assertEqual(cadena_de_extremos("a"), "", "Esperado: ''")
        self.assertEqual(cadena_de_extremos("xyz"), "xyyz", "Esperado: xyyz")
        self.assertEqual(cadena_de_extremos(""), "", "Esperado: ''")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-3">
<h2>Ejercicio 3<a class="headerlink" href="#ejercicio-3" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_3 data-question_label="3">
<div id=q11_3_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter</span></code> que, dada una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code>,
devuelva una cadena en la que todas las apariciones del primer carácter en <code class="docutils literal notranslate"><span class="pre">s</span></code>
se reemplacen por “*”, a excepción del primero. <strong>Nota:</strong>
use el método <code class="docutils literal notranslate"><span class="pre">.replace(valor_a_replazar,</span> <span class="pre">nuevo_valor)</span></code> para resolver el
ejercicio. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter(&quot;google&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;goo*le&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">remplazar_primer_caracter(&quot;dona&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;dona&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q11_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def remplazar_primer_caracter(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(remplazar_primer_caracter("babble"), "ba**le", "Esperado: ba**le")
        self.assertEqual(remplazar_primer_caracter("aardvark"), "a*rdv*rk", "Esperado: a*rdv*rk")
        self.assertEqual(remplazar_primer_caracter("google"), "goo*le", "Esperado: goo*le")
        self.assertEqual(remplazar_primer_caracter("dona"), "dona", "Esperado: dona")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-4">
<h2>Ejercicio 4<a class="headerlink" href="#ejercicio-4" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_4 data-question_label="4">
<div id=q11_4_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas</span></code> que tome dos cadenas como
argumentos, <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code>, y devuelva una nueva cadena de la siguiente forma:</p>
<ul class="simple">
<li><p>La nueva cadena tiene que ser una combinación de <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code>.</p></li>
<li><p>La nueva cadena tendrá la forma <code class="docutils literal notranslate"><span class="pre">&quot;&lt;a&gt;</span> <span class="pre">&lt;b&gt;&quot;</span></code>, note el espacio entre ambas.</p></li>
<li><p>La nueva cadena intercambiará las primeras dos letras de <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code>.</p></li>
</ul>
<p>Suponga que <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code> tienen más de 2 caracteres.
Para mayor claridad, observe los siguientes ejemplos. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas(&quot;mix&quot;,</span> <span class="pre">&quot;pod&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;pox</span> <span class="pre">mid&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">combinar_dos_cadenas(&quot;pezzy&quot;,</span> <span class="pre">&quot;firm&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;fizzy</span> <span class="pre">perm&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q11_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def combinar_dos_cadenas(a, b):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(combinar_dos_cadenas("mix", "pod"), "pox mid", "Esperado: pox mid")
        self.assertEqual(combinar_dos_cadenas("dog", "dinner"), "dig donner", "Esperado: dig donner")
        self.assertEqual(
            combinar_dos_cadenas("gnash", "sport"),
            "spash gnort",
            "Esperado: spash gnort",
        )
        self.assertEqual(combinar_dos_cadenas("pezzy", "firm"), "fizzy perm", "Esperado: fizzy perm")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-5">
<h2>Ejercicio 5<a class="headerlink" href="#ejercicio-5" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_5 data-question_label="5">
<div id=q11_5_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">es_palindromo</span></code> que toma una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> como
parámetro y verifica si <code class="docutils literal notranslate"><span class="pre">s</span></code> es palíndromo o no, devolviendo <code class="docutils literal notranslate"><span class="pre">True</span></code> o
<code class="docutils literal notranslate"><span class="pre">False</span></code> respectivamente. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">es_palindromo(&quot;asa&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_palindromo(&quot;casa&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code>  <br /></p>

</div>
<textarea data-lang="python" id="q11_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def es_palindromo(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(es_palindromo("asa"), True, "Esperado: True")
        self.assertEqual(es_palindromo("casa"), False, "Esperado: False")
        self.assertEqual(es_palindromo("reconocer"), True, "Esperado: True")
        self.assertEqual(es_palindromo("palabra"), False, "Esperado: False")
        self.assertEqual(es_palindromo("radar"), True, "Esperado: True")
        self.assertEqual(es_palindromo("seres"), True, "Esperado: True")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-6">
<h2>Ejercicio 6<a class="headerlink" href="#ejercicio-6" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q11_6 data-question_label="6">
<div id=q11_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">contar_ocurrencias</span></code> que toma dos parámetros:
<code class="docutils literal notranslate"><span class="pre">frase</span></code> y <code class="docutils literal notranslate"><span class="pre">palabra</span></code>, ambos de tipo cadena. La función debe devolver
el número de veces que <code class="docutils literal notranslate"><span class="pre">palabra</span></code> se encuentra en <code class="docutils literal notranslate"><span class="pre">frase</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">contar_ocurrencias(&quot;a</span> <span class="pre">ana</span> <span class="pre">y</span> <span class="pre">a</span> <span class="pre">mariana</span> <span class="pre">les</span> <span class="pre">gustan</span> <span class="pre">las</span> <span class="pre">manzanas&quot;,</span> <span class="pre">&quot;ana&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q11_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_ocurrencias(frase, palabra):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):

        self.assertEqual(
            contar_ocurrencias("a ana y a mariana les gustan las manzanas", "ana"),
            3,
            "Esperado: 3",
        )
        self.assertEqual(contar_ocurrencias("Cats, rats, bats, and hats.", "ats"), 4, "Esperado: 4")


myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 10' data-toggle="tooltip">
    <a href="Quiz10.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz10.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 12' data-toggle="tooltip" >
    <a href="Quiz12.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz12.html" >Next Section - Quiz - 12</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.7.1.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>