<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 7 &#8212; PyZombis master@8d046a9</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/runtime.bundle.js?v=ef1af912ec3415fea854"></script>
    <script defer="" src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script defer="" src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 9" href="Quiz9.html" />
    <link rel="prev" title="Quiz - 6" href="Quiz6.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.8.1';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@8d046a9</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz7'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 7</a><ul>
<li><a class="reference internal" href="#ejercicio-1">Ejercicio 1</a></li>
<li><a class="reference internal" href="#ejercicio-2">Ejercicio 2</a></li>
<li><a class="reference internal" href="#ejercicio-3">Ejercicio 3</a></li>
<li><a class="reference internal" href="#ejercicio-4">Ejercicio 4</a></li>
<li><a class="reference internal" href="#ejercicio-5">Ejercicio 5</a></li>
<li><a class="reference internal" href="#ejercicio-6">Ejercicio 6</a></li>
<li><a class="reference internal" href="#ejercicio-8">Ejercicio 8</a></li>
<li><a class="reference internal" href="#ejercicio-9">Ejercicio 9</a></li>
<li><a class="reference internal" href="#ejercicio-10">Ejercicio 10</a></li>
<li><a class="reference internal" href="#ejercicio-11">Ejercicio 11</a></li>
<li><a class="reference internal" href="#ejercicio-12">Ejercicio 12</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 6" data-toggle="tooltip">
    <a href="Quiz6.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 9' data-toggle="tooltip" >              
    <a href="Quiz9.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-7">
<h1>Quiz - 7<a class="headerlink" href="#quiz-7" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-1">
<h2>Ejercicio 1<a class="headerlink" href="#ejercicio-1" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_1 data-question_label="1">
<div id=q7_1_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">multi_cadena</span></code> que recibe como parámetros una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y un entero positivo <code class="docutils literal notranslate"><span class="pre">n</span></code> y devuelve una nueva cadena que contiene <code class="docutils literal notranslate"><span class="pre">n</span></code> copias de la cadena original <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">multi_cadena(&quot;Hola&quot;,</span> <span class="pre">2)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HolaHola&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">multi_cadena(&quot;Hola&quot;,</span> <span class="pre">5)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HolaHolaHolaHolaHola&quot;</span></code></p>

</div>
<textarea data-lang="python" id="q7_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def multi_cadena(s, n):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(multi_cadena("Hola", 2), "HolaHola", "Esperado: HolaHola")
        self.assertEqual(
                multi_cadena("Hola", 3), "HolaHolaHola", "Esperado: HolaHolaHola"
        )
        self.assertEqual(multi_cadena("Hi", 1), "Hi", "Esperado: Hi")
        self.assertEqual(multi_cadena("Hi", 0), "", "Esperado: ''")
        self.assertEqual(multi_cadena("Hi", 5), "HiHiHiHiHi", "Esperado: HiHiHiHiHi")
        self.assertEqual(
                multi_cadena("Oh Boy!", 2), "Oh Boy!Oh Boy!", "Esperado: Oh Boy!Oh Boy!"
        )
        self.assertEqual(multi_cadena("x", 4), "xxxx", "Esperado: xxxx")
        self.assertEqual(multi_cadena("", 4), "", "Esperado: ''")
        self.assertEqual(multi_cadena("code", 2), "codecode", "Esperado: codecode")
        self.assertEqual(
                multi_cadena("code", 3), "codecodecode", "Esperado: codecodecode"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-2">
<h2>Ejercicio 2<a class="headerlink" href="#ejercicio-2" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_2 data-question_label="2">
<div id=q7_2_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">expandir_cadena</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelva una nueva cadena siguiendo el patrón descrito en los ejemplos <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">expandir_cadena(&quot;Code&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;CCoCodCode&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">expandir_cadena(&quot;abc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;aababc&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">expandir_cadena(&quot;ab&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;aab&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def expandir_cadena(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(expandir_cadena("Code"), "CCoCodCode", "Esperado: CCoCodCode")
        self.assertEqual(expandir_cadena("abc"), "aababc", "Esperado: aababc")
        self.assertEqual(expandir_cadena("ab"), "aab", "Esperado: aab")
        self.assertEqual(expandir_cadena("x"), "x", "Esperado: x")
        self.assertEqual(expandir_cadena("fade"), "ffafadfade", "Esperado: ffafadfade")
        self.assertEqual(
            expandir_cadena("There"), "TThTheTherThere", "Esperado: TThTheTherThere"
        )
        self.assertEqual(
            expandir_cadena("Kitten"),
            "KKiKitKittKitteKitten",
            "Esperado: KKiKitKittKitteKitten",
        )
        self.assertEqual(expandir_cadena("Bye"), "BByBye", "Esperado: BByBye")
        self.assertEqual(expandir_cadena("Good"), "GGoGooGood", "Esperado: GGoGooGood")
        self.assertEqual(expandir_cadena("Bad"), "BBaBad", "Esperado: BBaBad")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-3">
<h2>Ejercicio 3<a class="headerlink" href="#ejercicio-3" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_3 data-question_label="3">
<div id=q7_3_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">contar_apariciones_9</span></code> que recibe como parámetro una lista no vacía de números enteros <code class="docutils literal notranslate"><span class="pre">numeros</span></code> y devuelva cuántas veces aparece el número 9 en la lista <br /> <br />
Ejemplo: <code class="docutils literal notranslate"><span class="pre">contar_apariciones_9([1,</span> <span class="pre">99,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code></p>

</div>
<textarea data-lang="python" id="q7_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_apariciones_9(numeros):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(contar_apariciones_9([1, 99, 9]), 1, "Esperado: 1")
        self.assertEqual(contar_apariciones_9([1, 9, 9]), 2, "Esperado: 2")
        self.assertEqual(contar_apariciones_9([1, 9, 9, 3, 9]), 3, "Esperado: 3")
        self.assertEqual(contar_apariciones_9([1, 2, 3]), 0, "Esperado: 0")
        self.assertEqual(contar_apariciones_9([1]), 0, "Esperado: 1")
        self.assertEqual(contar_apariciones_9([4, 2, 4, 3, 1]), 0, "Esperado: 0")
        self.assertEqual(contar_apariciones_9([9, 2, 99, 3, 1]), 1, "Esperado: 1")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-4">
<h2>Ejercicio 4<a class="headerlink" href="#ejercicio-4" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_4 data-question_label="4">
<div id=q7_4_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">verificar_comienzo_9</span></code> que recibe una lista de números enteros <code class="docutils literal notranslate"><span class="pre">numeros</span></code> y verifica si al menos uno de los primeros cuatro números es un 9 <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">verificar_comienzo_9([1,</span> <span class="pre">2,</span> <span class="pre">9,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">verificar_comienzo_9([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">verificar_comienzo_9([1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4,</span> <span class="pre">5])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">verificar_comienzo_9([1,</span> <span class="pre">2,</span> <span class="pre">9])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def verificar_comienzo_9(numeros):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            verificar_comienzo_9([1, 2, 9, 3, 4]), True, "Esperado: True"
        )
        self.assertEqual(
            verificar_comienzo_9([1, 2, 3, 4, 9]), False, "Esperado: False"
        )
        self.assertEqual(
            verificar_comienzo_9([1, 2, 3, 4, 5]), False, "Esperado: False"
        )
        self.assertEqual(
            verificar_comienzo_9([9, 2, 3]), True, "Esperado: True"
        )
        self.assertEqual(
            verificar_comienzo_9([1, 9, 9]), True, "Esperado: True"
        )
        self.assertEqual(
            verificar_comienzo_9([1, 2, 3]), False, "Esperado: False"
        )
        self.assertEqual(
            verificar_comienzo_9([1, 9]), True, "Esperado: True"
        )
        self.assertEqual(
            verificar_comienzo_9([5, 5]), False, "Esperado: False"
        )
        self.assertEqual(
            verificar_comienzo_9([2]), False, "Esperado: False"
        )
        self.assertEqual(verificar_comienzo_9([9]), True, "Esperado: True")
        self.assertEqual(verificar_comienzo_9([]), False, "Esperado: False")
        self.assertEqual(verificar_comienzo_9(
            [3, 9, 2, 3, 3]), True, "Esperado: True"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-5">
<h2>Ejercicio 5<a class="headerlink" href="#ejercicio-5" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_5 data-question_label="5">
<div id=q7_5_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">hola_usuario</span></code> que recibe como parámetro una cadena <code class="docutils literal notranslate"><span class="pre">nombre</span></code> que representa el nombre de un usuario y devuelva un saludo con este nombre <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">hola_usuario(&quot;Bob&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;¡Hola</span> <span class="pre">Bob!&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">hola_usuario(&quot;Alice&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;¡Hola</span> <span class="pre">Alice!&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">hola_usuario(&quot;X&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;¡Hola</span> <span class="pre">X!&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def hola_usuario(nombre):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            hola_usuario("Bob"), "¡Hola Bob!", "Esperado: ¡Hola Bob!"
        )
        self.assertEqual(
            hola_usuario("Alice"), "¡Hola Alice!", "Esperado: ¡Hola Alice!"
        )
        self.assertEqual(
            hola_usuario("X"), "¡Hola X!", "Esperado: ¡Hola X!"
        )
        self.assertEqual(
            hola_usuario("Hola"), "¡Hola Hola!", "Esperado: ¡Hola Hola!"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-6">
<h2>Ejercicio 6<a class="headerlink" href="#ejercicio-6" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_6 data-question_label="6">
<div id=q7_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">crear_etiquetas</span></code> que recibe dos cadenas <code class="docutils literal notranslate"><span class="pre">etiqueta</span></code> y <code class="docutils literal notranslate"><span class="pre">palabra</span></code> y devuelva una nueva cadena formateada siguiendo el patrón de los ejemplos <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">crear_etiquetas(&quot;i&quot;,</span> <span class="pre">&quot;Yay&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;&lt;i&gt;Yay&lt;/i&gt;&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">crear_etiquetas(&quot;i&quot;,</span> <span class="pre">&quot;Hello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;&lt;i&gt;Hello&lt;/i&gt;&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">crear_etiquetas(&quot;cite&quot;,</span> <span class="pre">&quot;Yay&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;&lt;cite&gt;Yay&lt;/cite&gt;&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def crear_etiquetas(etiqueta, palabra):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            crear_etiquetas("i", "Yay"),
            "<i>Yay</i>",
            "Esperado: <i>Yay</i>"
        )
        self.assertEqual(
            crear_etiquetas("i", "Hello"),
            "<i>Hello</i>",
            "Esperado:<i>Hello</i>"
        )
        self.assertEqual(
            crear_etiquetas("cite", "Yay"),
            "<cite>Yay</cite>",
            "Esperado: <cite>Yay</cite>",
        )
        self.assertEqual(
            crear_etiquetas("address", "here"),
            "<address>here</address>",
            "Esperado: <address>here</address>",
        )
        self.assertEqual(
            crear_etiquetas("body", "Heart"),
            "<body>Heart</body>",
            "Esperado: <body>Heart</body>",
        )
        self.assertEqual(
            crear_etiquetas("i", "i"),
            "<i>i</i>",
            "Esperado: <i>i</i>"
        )
        self.assertEqual(
            crear_etiquetas("i", ""),
            "<i></i>",
            "Esperado: <i></i>"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-8">
<h2>Ejercicio 8<a class="headerlink" href="#ejercicio-8" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_8 data-question_label="7">
<div id=q7_8_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">repetir_letras</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelve una nueva cadena con las dos últimas letras repetidas tres veces <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">repetir_letras(&quot;Hello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;lololo&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">repetir_letras(&quot;abb&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;bbbbbb&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">repetir_letras(&quot;Hi&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HiHiHi&quot;</span></code>  <br /></p>

</div>
<textarea data-lang="python" id="q7_8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def repetir_letras(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(repetir_letras("Hello"), "lololo", "Esperado: lololo")
        self.assertEqual(repetir_letras("ab"), "ababab", "Esperado: ababab")
        self.assertEqual(repetir_letras("Hi"), "HiHiHi", "Esperado: HiHiHi")
        self.assertEqual(repetir_letras("Candy"), "dydydy", "Esperado: dydydy")
        self.assertEqual(repetir_letras("Code"), "dedede", "Esperado: dedede")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-9">
<h2>Ejercicio 9<a class="headerlink" href="#ejercicio-9" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_9 data-question_label="8">
<div id=q7_9_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">otra_repetir_letras</span></code> (variante de la función del Ejercicio 8) que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y un entero positivo <code class="docutils literal notranslate"><span class="pre">n</span></code> y devuelva una nueva cadena con las dos últimas letras repetidas <code class="docutils literal notranslate"><span class="pre">n</span></code> veces <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">otra_repetir_letras(&quot;Hello&quot;,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;lololo&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">otra_repetir_letras(&quot;abb&quot;,</span> <span class="pre">1)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;bb&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">otra_repetir_letras(&quot;Hi&quot;,</span> <span class="pre">5)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HiHiHiHiHi&quot;</span></code>  <br /></p>

</div>
<textarea data-lang="python" id="q7_9_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def otra_repetir_letras(s, n):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(
            otra_repetir_letras("Hello", 3), "lololo", "Esperado: lololo"
        )
        self.assertEqual(otra_repetir_letras("ab", 2), "abab", "Esperado: abab")
        self.assertEqual(
            otra_repetir_letras("Hi", 6), "HiHiHiHiHiHi", "Esperado: HiHiHi"
        )
        self.assertEqual(otra_repetir_letras("Candy", 1), "dy", "Esperado: dy")
        self.assertEqual(otra_repetir_letras("Code", 0), "", "Esperado: ''")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-10">
<h2>Ejercicio 10<a class="headerlink" href="#ejercicio-10" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_10 data-question_label="9">
<div id=q7_10_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">primera_mitad</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelve la primera mitad de la cadena original <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">primera_mitad(&quot;WooHoo&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Woo&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">primera_mitad(&quot;HelloThere&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Hello&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">primera_mitad(&quot;abcdef&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;abc&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_10_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def primera_mitad(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(primera_mitad("WooHoo"), "Woo", "Esperado: Woo")
        self.assertEqual(
            primera_mitad("HelloThere"), "Hello", "Esperado: Hello"
        )
        self.assertEqual(primera_mitad("abcdef"), "abc", "Esperado: abc")
        self.assertEqual(primera_mitad(""), "", "Esperado: ")
        self.assertEqual(primera_mitad("ab"), "a", "Esperado: a")
        self.assertEqual(
            primera_mitad("0123456789"), "01234", "Esperado: 01234"
        )
        self.assertEqual(primera_mitad("kitten"), "kit", "Esperado: kit")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-11">
<h2>Ejercicio 11<a class="headerlink" href="#ejercicio-11" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_11 data-question_label="10">
<div id=q7_11_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">remover_primer_ultimo</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelva una cadena sin el primer ni el último carácter <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">remover_primer_ultimo(&quot;Hello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;ell&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">remover_primer_ultimo(&quot;python&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;ytho&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">remover_primer_ultimo(&quot;coding&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;odin&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_11_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def remover_primer_ultimo(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(remover_primer_ultimo("Hello"), "ell", "Esperado: ell")
        self.assertEqual(remover_primer_ultimo("Python"), "ytho", "Esperado: ytho")
        self.assertEqual(remover_primer_ultimo("coding"), "odin", "Esperado: odin")
        self.assertEqual(remover_primer_ultimo("code"), "od", "Esperado: od")
        self.assertEqual(remover_primer_ultimo("ab"), "", "Esperado: ")
        self.assertEqual(
            remover_primer_ultimo(" PyZombies "), "PyZombies", "Esperado: PyZombies"
        )
        self.assertEqual(
            remover_primer_ultimo("Chocolate!"), "hocolate", "Esperado: hocolate"
        )
        self.assertEqual(remover_primer_ultimo("kitten"), "itte", "Esperado: itte")
        self.assertEqual(remover_primer_ultimo("woohoo"), "ooho", "Esperado: ooho")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-12">
<h2>Ejercicio 12<a class="headerlink" href="#ejercicio-12" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q7_12 data-question_label="11">
<div id=q7_12_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">rotar_izq_2</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> de al menos dos caracteres y devuelva la cadena original rotada 2 posiciones a la izquierda <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">rotar_izq_2(&quot;Hello&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;lloHe&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">rotar_izq_2(&quot;Hi&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;Hi&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q7_12_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def rotar_izq_2(s):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(rotar_izq_2("Hello"), "lloHe", "Esperado: lloHe")
        self.assertEqual(rotar_izq_2("python"), "thonpy", "Esperado: thonpy")
        self.assertEqual(rotar_izq_2("Hi"), "Hi", "Esperado: Hi")
        self.assertEqual(rotar_izq_2("code"), "deco", "Esperado: deco")
        self.assertEqual(rotar_izq_2("cat"), "tca", "Esperado: tca")
        self.assertEqual(rotar_izq_2("12345"), "34512", "Esperado: 34512")
        self.assertEqual(
            rotar_izq_2("Chocolate"), "ocolateCh", "Esperado: ocolateCh"
        )
        self.assertEqual(rotar_izq_2("bricks"), "icksbr", "Esperado: icksbr")
        self.assertEqual(
            rotar_izq_2("isPyzomb"), "PyZombies", "Esperado: PyZombies"
        )


myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 6' data-toggle="tooltip">
    <a href="Quiz6.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz6.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 9' data-toggle="tooltip" >
    <a href="Quiz9.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz9.html" >Next Section - Quiz - 9</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.8.1.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>