<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 5 &#8212; PyZombis master@6914c51</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=ef1af912ec3415fea854"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 6" href="Quiz6.html" />
    <link rel="prev" title="Quiz - 4" href="Quiz4.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.7.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@6914c51</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz5'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 5</a><ul>
<li><a class="reference internal" href="#ejercicio-1">Ejercicio 1</a></li>
<li><a class="reference internal" href="#ejercicio-2">Ejercicio 2</a></li>
<li><a class="reference internal" href="#ejercicio-3">Ejercicio 3</a></li>
<li><a class="reference internal" href="#ejercicio-4">Ejercicio 4</a></li>
<li><a class="reference internal" href="#ejercicio-5">Ejercicio 5</a></li>
<li><a class="reference internal" href="#ejercicio-6">Ejercicio 6</a></li>
<li><a class="reference internal" href="#ejercicio-7">Ejercicio 7</a></li>
<li><a class="reference internal" href="#ejercicio-8">Ejercicio 8</a></li>
<li><a class="reference internal" href="#ejercicio-9">Ejercicio 9</a></li>
<li><a class="reference internal" href="#ejercicio-10">Ejercicio 10</a></li>
<li><a class="reference internal" href="#ejercicio-11">Ejercicio 11</a></li>
<li><a class="reference internal" href="#ejercicio-12">Ejercicio 12</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 4" data-toggle="tooltip">
    <a href="Quiz4.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 6' data-toggle="tooltip" >              
    <a href="Quiz6.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-5">
<h1>Quiz - 5<a class="headerlink" href="#quiz-5" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-1">
<h2>Ejercicio 1<a class="headerlink" href="#ejercicio-1" title="Permalink to this headline">¶</a></h2>

<div class="runestone">
<div data-component="clickablearea" class="runestone alert alert-warning" id="q5_1" data-question_label="1"     style="visibility: hidden;">
<span data-question>Q-1: El siguiente programa declara una variable en el lugar equivocado del código. Su tarea es seleccionar la declaración incorrecta y seleccionar la línea correcta donde debería ir.</span><span data-feedback>Recuerda, que no puedes usar una variable sin haberla declarado anteriormente.</span><pre><span data-incorrect>x = 2</span>
<span data-correct>     </span>
<span data-incorrect>si y > 8 entonces:</span>
    <span data-correct>y = 5</span>
    <span data-incorrect>y = y * 2</span>

<span data-incorrect>en caso contrario:</span>
    <span data-incorrect>x = x * 2:</span>
    <span data-incorrect>imprime (x + y):</span></pre>

</div>
</div>
</div>
<div class="section" id="ejercicio-2">
<h2>Ejercicio 2<a class="headerlink" href="#ejercicio-2" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_2 data-question_label="2">
<div id=q5_2_question class="ac_question col-md-12">
<p>Traduzca el programa anterior (con la declaración en el lugar correcto) al lenguaje Python utilizando el siguiente bloque de código. <br /></p>

</div>
<textarea data-lang="python" id="q5_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">

====
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-3">
<h2>Ejercicio 3<a class="headerlink" href="#ejercicio-3" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="3" id="q5_3"  style="visibility: hidden;">
            <p>Q-3: ¿Qué imprime el programa anterior?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [9, 9], "feedback": "<p>Tu respuesta es correcta, \u00a1buen trabajo!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Tu respuesta es incorrecta, intenta nuevamente.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="ejercicio-4">
<h2>Ejercicio 4<a class="headerlink" href="#ejercicio-4" title="Permalink to this headline">¶</a></h2>
<div class="highlight-none notranslate"><div class="highlight"><pre><span></span>para i = 1 hasta 9:
    si i != 3 entonces:
        para j = 1 hasta 6:
            imprimir(&quot;Hola&quot;)
</pre></div>
</div>
<br/>
<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_4 data-question_label="4">
<div id=q5_4_question class="ac_question col-md-12">
<p>Traduzca el programa anterior al lenguaje Python utilizando el siguiente bloque de código. <br />
<strong>Nota</strong>: en nuestro pseudolenguaje, el ciclo incluye los extremos, es decir, 1 a 4 significa 1, 2, 3, 4. <br /></p>

</div>
<textarea data-lang="python" id="q5_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">

====
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-5">
<h2>Ejercicio 5<a class="headerlink" href="#ejercicio-5" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="5" id="q5_5"  style="visibility: hidden;">
            <p>Q-5: ¿Cuántas veces el programa anterior imprime <code class="docutils literal notranslate"><span class="pre">&quot;Hola&quot;</span></code>?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [48, 48], "feedback": "<p>Tu respuesta es correcta, \u00a1buen trabajo!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Tu respuesta es incorrecta, intenta nuevamente.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="ejercicio-6">
<h2>Ejercicio 6<a class="headerlink" href="#ejercicio-6" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_6 data-question_label="6">
<div id=q5_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">pares_divisibles_7</span></code> que recibe como parámetros dos enteros positivos, <code class="docutils literal notranslate"><span class="pre">inicio</span></code> y <code class="docutils literal notranslate"><span class="pre">fin</span></code> que representan el inicio y el fin de un intervalo de números, la función tiene que devolver la cantidad de números que son pares y también divisibles por 7. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(1,</span> <span class="pre">7)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(25,</span> <span class="pre">123)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">7</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">pares_divisibles_7(13,</span> <span class="pre">245)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">17</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q5_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def pares_divisibles_7(inicio, fin):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(pares_divisibles_7(1, 7), 0, "Esperado: 0")
        self.assertEqual(pares_divisibles_7(25, 123), 7, "Esperado: 7")
        self.assertEqual(pares_divisibles_7(13, 245), 17, "Esperado: 17")
        self.assertEqual(pares_divisibles_7(1300, 2460), 83, "Esperado: 83")
        self.assertEqual(pares_divisibles_7(1, 2460), 175, "Esperado: 175")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-7">
<h2>Ejercicio 7<a class="headerlink" href="#ejercicio-7" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="7" id="q5_7"  style="visibility: hidden;">
            <p>Q-7: Utilice la función creada en el ejercicio anterior para responder ¿Cuántos números son pares y también divisibles por 7 entre 1067 y 3627 (inclusive)?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [183, 183], "feedback": "<p>Tu respuesta es correcta, \u00a1buen trabajo!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Tu respuesta es incorrecta, intenta nuevamente.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="ejercicio-8">
<h2>Ejercicio 8<a class="headerlink" href="#ejercicio-8" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_8 data-question_label="8">
<div id=q5_8_question class="ac_question col-md-12">
<p>Daniela es una persona muy supersticiosa. Para ella, un número tiene <strong>suerte</strong> si contiene el dígito <code class="docutils literal notranslate"><span class="pre">2</span></code> pero no el <code class="docutils literal notranslate"><span class="pre">7</span></code>. Ella además es muy curiosa y está interesada en saber ¿cuántos números de la suerte hay en un intervalo de números (incluido los extremos)?. Para ello le ha solicitado a usted que haga la función llamada <code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte</span></code>,  que recibe como parámetros dos enteros positivos, <code class="docutils literal notranslate"><span class="pre">inicio</span></code> y <code class="docutils literal notranslate"><span class="pre">fin</span></code> que representan el inicio y el fin de un intervalo de números, la función tiene que devolver la cantidad de números que tienen <strong>suerte</strong>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(1,</span> <span class="pre">7)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(1,</span> <span class="pre">20)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(25,</span> <span class="pre">123)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">16</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cuantos_tienen_suerte(13,</span> <span class="pre">245)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">74</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q5_8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def cuantos_tienen_suerte(inicio, fin):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(cuantos_tienen_suerte(1, 1), 0, "Esperado: 0")
        self.assertEqual(cuantos_tienen_suerte(1, 7), 1, "Esperado: 1")
        self.assertEqual(cuantos_tienen_suerte(1, 20), 3, "Esperado: 3")
        self.assertEqual(cuantos_tienen_suerte(25, 123), 16, "Esperado: 16")
        self.assertEqual(cuantos_tienen_suerte(1300, 2460), 481, "Esperado: 481")
        self.assertEqual(cuantos_tienen_suerte(13, 245), 74, "Esperado: 74")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-9">
<h2>Ejercicio 9<a class="headerlink" href="#ejercicio-9" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="9" id="q5_9"  style="visibility: hidden;">
            <p>Q-9: Utilice la función creada en el ejercicio anterior para responder ¿Cuántos números tienen suerte entre 18644 y 33087 (inclusive)?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [7995, 7995], "feedback": "<p>Tu respuesta es correcta, \u00a1buen trabajo!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Tu respuesta es incorrecta, intenta nuevamente.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="ejercicio-10">
<h2>Ejercicio 10<a class="headerlink" href="#ejercicio-10" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_10 data-question_label="10">
<div id=q5_10_question class="ac_question col-md-12">
<p>En el tranquilo pueblo rural de <em>Ponteironuloville</em>, todos los teléfonos tienen 6 dígitos. La compañía telefónica establece las siguientes reglas sobre los números: <br /></p>
<ol class="arabic simple">
<li><p>No puede haber dos dígitos idénticos consecutivos, porque esto es aburrido</p></li>
<li><p>La suma de los dígitos debe ser par, porque esto es legal</p></li>
<li><p>El último dígito no puede ser el mismo que el primero, porque eso es mala suerte.</p></li>
</ol>
<p>Entonces, dadas estas reglas perfectamente razonables, bien diseñadas y maduras, desarrolle la función que se llama <code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville</span></code> que recibe una cadena de enteros positivos y devuelve <code class="docutils literal notranslate"><span class="pre">True</span></code> si el número es valido de acuerdo a las reglas de <em>Ponteironuloville</em> y <code class="docutils literal notranslate"><span class="pre">False</span></code> en caso contrario.
<br /> <br /></p>
<p>Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;123457&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;234562&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;222222&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;123456&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;312214&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">es_numero_ponteironuloville(&quot;312312&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q5_10_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def es_numero_ponteironuloville(numero):



====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(es_numero_ponteironuloville("123457"), True, "Esperado: True")
        self.assertEqual(es_numero_ponteironuloville("234562"), False, "Esperado: False")
        self.assertEqual(es_numero_ponteironuloville("304706"), True, "Esperado: True")
        self.assertEqual(es_numero_ponteironuloville("222222"), False, "Esperado: False")
        self.assertEqual(es_numero_ponteironuloville("123456"), False, "Esperado: False")
        self.assertEqual(es_numero_ponteironuloville("312214"), False, "Esperado: False")
        self.assertEqual(es_numero_ponteironuloville("312312"), True, "Esperado: True")
        self.assertEqual(es_numero_ponteironuloville("131313"), True, "Esperado: True")
        self.assertEqual(es_numero_ponteironuloville("249184"), True, "Esperado: True")
        self.assertEqual(es_numero_ponteironuloville("012445"), False, "Esperado: False")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-11">
<h2>Ejercicio 11<a class="headerlink" href="#ejercicio-11" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q5_11 data-question_label="11">
<div id=q5_11_question class="ac_question col-md-12">
<p>Sabiendo que la función <code class="docutils literal notranslate"><span class="pre">split()</span></code> divide una cadena en múltiples cadenas, utilice la función creada en el ejercicio anterior (<strong>OBLIGATORIO</strong> haberla terminado antes y haber pasado todas las pruebas) para terminar la función <code class="docutils literal notranslate"><span class="pre">contar_validos</span></code> que devuelve la cantidad de números válidos de la lista de números dada en el siguiente bloque de código. <br /> <br />
<strong>Nota</strong>: la lista de números es una cadena, utilice sabiamente la función <code class="docutils literal notranslate"><span class="pre">split()</span></code> para obtener un arreglo de cadenas.</p>

</div>
<textarea data-lang="python" id="q5_11_editor" 
     data-include="q5_10" data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
lista_numeros = """236043 237330 239636 240138 242123 246224 249183 252936
    254711 257200 257607 261424 263814 266794 268649 273050
    275001 277606 278997 283331 287104 287953 289137 291591
    292559 292946 295180 295566 297529 300400 304707 306931
    310638 313595 318449 319021 322082 323796 326266 326880
    327249 329914 334392 334575 336723 336734 338808 343269
    346040 350113 353631 357154 361633 361891 364889 365746
    365749 366426 369156 369444 369689 372896 374983 375223
    379163 380712 385640 386777 388599 389450 390178 392943
    394742 395921 398644 398832 401149 402219 405364 408088
    412901 417683 422267 424767 426613 430474 433910 435054
    440052 444630 447852 449116 453865 457631 461750 462985
    463328 466458 469601 473108 476773 477956 481991 482422
    486195 488359 489209 489388 491928 496569 496964 497901
    500877 502386 502715 507617 512526 512827 513796 518232
    521455 524277 528496 529345 531231 531766 535067 535183
    536593 537360 539055 540582 543708 547492 550779 551595
    556493 558807 559102 562050 564962 569677 570945 575447
    579937 580112 580680 582458 583012 585395 586244 587393
    590483 593112 593894 594293 597525 598184 600455 600953
    601523 605761 608618 609198 610141 610536 612636 615233
    618314 622752 626345 626632 628889 629457 629643 633673
    637656 641136 644176 644973 647617 652218 657143 659902
    662224 666265 668010 672480 672695 676868 677125 678315"""

def contar_validos():
    contador = 0
    """
        Escriba su codigo aqui, use la variable creada
    """
    return contador
====
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-12">
<h2>Ejercicio 12<a class="headerlink" href="#ejercicio-12" title="Permalink to this headline">¶</a></h2>

        <div class="runestone">
        <div data-component="fillintheblank" data-question_label="12" id="q5_12"  style="visibility: hidden;">
            <p>Q-12: Utilice la función creada en el ejercicio anterior para responder ¿Cuántos números de la lista son válidos de acuerdo a las reglas de <em>Ponteironuloville</em>?</p>
<input type="text">
        <script type="application/json">
            [[{"number": [39, 39], "feedback": "<p>Tu respuesta es correcta, \u00a1buen trabajo!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Tu respuesta es incorrecta, intenta nuevamente.</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 4' data-toggle="tooltip">
    <a href="Quiz4.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz4.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 6' data-toggle="tooltip" >
    <a href="Quiz6.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz6.html" >Next Section - Quiz - 6</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.7.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>