<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quiz - 10 &#8212; PyZombis master@15721a0</title>
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../_static/runtime.bundle.js?v=ef1af912ec3415fea854"></script>
    <script src="../_static/vendors-node_modules_bootstrap_dist_js_bootstrap_js-node_modules_jquery-ui_jquery-ui_js-node_-e60694.bundle.js?v=0e91cff9e660624062d3"></script>
    <script src="../_static/runestone.bundle.js?v=766cad0b55ecfe66f111"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Quiz - 11" href="Quiz11.html" />
    <link rel="prev" title="Quiz - 9" href="Quiz9.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'PyZombisDev';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'PyZombisDev';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '5.7.2';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
  eBookConfig.allow_pairs = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html" aria-label="index-page">PyZombis master@15721a0</a>
      </div>
    </div>



    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html' aria-label="index-page">Table of Contents</a></li>
            
            <li><a href='/runestone/static/PyZombisDev/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/PyZombisDev/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=PyZombisDev&page=Quiz10'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Quiz - 10</a><ul>
<li><a class="reference internal" href="#ejercicio-1">Ejercicio 1</a></li>
<li><a class="reference internal" href="#ejercicio-2">Ejercicio 2</a></li>
<li><a class="reference internal" href="#ejercicio-3">Ejercicio 3</a></li>
<li><a class="reference internal" href="#ejercicio-4">Ejercicio 4</a></li>
<li><a class="reference internal" href="#ejercicio-5">Ejercicio 5</a></li>
<li><a class="reference internal" href="#ejercicio-6">Ejercicio 6</a></li>
<li><a class="reference internal" href="#ejercicio-7">Ejercicio 7</a></li>
<li><a class="reference internal" href="#ejercicio-8">Ejercicio 8</a></li>
<li><a class="reference internal" href="#ejercicio-9">Ejercicio 9</a></li>
<li><a class="reference internal" href="#ejercicio-10">Ejercicio 10</a></li>
<li><a class="reference internal" href="#ejercicio-11">Ejercicio 11</a></li>
<li><a class="reference internal" href="#ejercicio-12">Ejercicio 12</a></li>
<li><a class="reference internal" href="#ejercicio-13">Ejercicio 13</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Quiz - 9" data-toggle="tooltip">
    <a href="Quiz9.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Quiz - 11' data-toggle="tooltip" >              
    <a href="Quiz11.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="quiz-10">
<h1>Quiz - 10<a class="headerlink" href="#quiz-10" title="Permalink to this headline">¶</a></h1>
<div class="section" id="ejercicio-1">
<h2>Ejercicio 1<a class="headerlink" href="#ejercicio-1" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_1 data-question_label="1">
<div id=q10_1_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">cerca_de_10</span></code> que recibe un entero <code class="docutils literal notranslate"><span class="pre">n</span></code> y devuelve <code class="docutils literal notranslate"><span class="pre">True</span></code> cuando <code class="docutils literal notranslate"><span class="pre">n</span></code> sea múltiplo de 10, o cuando <code class="docutils literal notranslate"><span class="pre">n</span></code> esté a una o dos unidades de serlo. En caso contrario devuelva <code class="docutils literal notranslate"><span class="pre">False</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">cerca_de_10(12)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cerca_de_10(17)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">cerca_de_10(19)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_1_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def cerca_de_10(n):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(cerca_de_10(12), True, "Esperado: True")
        self.assertEqual(cerca_de_10(17), False, "Esperado: False")
        self.assertEqual(cerca_de_10(19), True, "Esperado: True")
        self.assertEqual(cerca_de_10(31), True, "Esperado: True")
        self.assertEqual(cerca_de_10(6), False, "Esperado: False")
        self.assertEqual(cerca_de_10(10), True, "Esperado: True")
        self.assertEqual(cerca_de_10(11), True, "Esperado: True")
        self.assertEqual(cerca_de_10(21), True, "Esperado: True")
        self.assertEqual(cerca_de_10(22), True, "Esperado: True")
        self.assertEqual(cerca_de_10(23), False, "Esperado: False")
        self.assertEqual(cerca_de_10(54), False, "Esperado: False")
        self.assertEqual(cerca_de_10(155), False, "Esperado: False")
        self.assertEqual(cerca_de_10(158), True, "Esperado: True")
        self.assertEqual(cerca_de_10(3), False, "Esperado: False")
        self.assertEqual(cerca_de_10(1), True, "Esperado: True")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-2">
<h2>Ejercicio 2<a class="headerlink" href="#ejercicio-2" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_2 data-question_label="2">
<div id=q10_2_question class="ac_question col-md-12">
<p>Desarrolle una función llamada <code class="docutils literal notranslate"><span class="pre">suma_loca</span></code> que recibe tres números enteros <code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code> y <code class="docutils literal notranslate"><span class="pre">c</span></code>. La función debe devolver la suma de los tres números con la condición de que si algún número aparece repetido no será contado en la suma. <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(1,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">6</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(3,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_loca(3,</span> <span class="pre">3,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_2_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def suma_loca(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_loca(1, 2, 3), 6, "Esperado: 6")
        self.assertEqual(suma_loca(3, 2, 3), 2, "Esperado: 2")
        self.assertEqual(suma_loca(3, 3, 3), 0, "Esperado: 0")
        self.assertEqual(suma_loca(9, 2, 2), 9, "Esperado: 9")
        self.assertEqual(suma_loca(2, 2, 9), 9, "Esperado: 9")
        self.assertEqual(suma_loca(2, 9, 2), 9, "Esperado: 9")
        self.assertEqual(suma_loca(2, 9, 3), 14, "Esperado: 14")
        self.assertEqual(suma_loca(4, 2, 3), 9, "Esperado: 9")
        self.assertEqual(suma_loca(1, 3, 1), 3, "Esperado: 3")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-3">
<h2>Ejercicio 3<a class="headerlink" href="#ejercicio-3" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_3 data-question_label="3">
<div id=q10_3_question class="ac_question col-md-12">
<p>Desarrolle la función llamada <code class="docutils literal notranslate"><span class="pre">suma_con_suerte</span></code> que recibe tres números enteros <code class="docutils literal notranslate"><span class="pre">a</span></code>, <code class="docutils literal notranslate"><span class="pre">b</span></code> y <code class="docutils literal notranslate"><span class="pre">c</span></code>. La función debe devolver la suma de los tres números con la condición de que si uno de los números es el <strong>13</strong>, éste no cuenta en la suma así como todos los números que se encuentren a su derecha. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">2,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">6</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">2,</span> <span class="pre">13)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_suerte(1,</span> <span class="pre">13,</span> <span class="pre">3)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def suma_con_suerte(a, b, c):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_con_suerte(1, 2, 3), 6, "Esperado: 6")
        self.assertEqual(suma_con_suerte(1, 2, 13), 3, "Esperado: 3")
        self.assertEqual(suma_con_suerte(1, 13, 3), 1, "Esperado: 1")
        self.assertEqual(suma_con_suerte(1, 13, 13), 1, "Esperado: 1")
        self.assertEqual(suma_con_suerte(6, 5, 2), 13, "Esperado: 13")
        self.assertEqual(suma_con_suerte(13, 2, 3), 0, "Esperado: 0")
        self.assertEqual(suma_con_suerte(13, 2, 13), 0, "Esperado: 0")
        self.assertEqual(suma_con_suerte(13, 13, 2), 0, "Esperado: 0")
        self.assertEqual(suma_con_suerte(9, 4, 13), 13, "Esperado: 13")
        self.assertEqual(suma_con_suerte(8, 13, 2), 8, "Esperado: 8")
        self.assertEqual(suma_con_suerte(7, 2, 1), 10, "Esperado: 10")
        self.assertEqual(suma_con_suerte(3, 3, 13), 6, "Esperado: 6")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-4">
<h2>Ejercicio 4<a class="headerlink" href="#ejercicio-4" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_4 data-question_label="4">
<div id=q10_4_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">caracteres_dobles</span></code> que recibe una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y debe devolver los caracteres de la cadena original duplicados. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('The')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;TThhee&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('AAbb')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;AAAAbbbb&quot;</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">caracteres_dobles('Hi-There')</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">&quot;HHii--TThheerree&quot;</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_4_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def caracteres_dobles(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(caracteres_dobles("The"), "TThhee", "Esperado: TThhee")
        self.assertEqual(caracteres_dobles("AAbb"), "AAAAbbbb", "Esperado: AAAAbbbb")
        self.assertEqual(caracteres_dobles("Hi-There"), "HHii--TThheerree", "Esperado: HHii--TThheerree")
        self.assertEqual(caracteres_dobles("Word!"), "WWoorrdd!!", "Esperado: WWoorrdd!!")
        self.assertEqual(caracteres_dobles("!!"), "!!!!", "Esperado: !!!!")
        self.assertEqual(caracteres_dobles(""), "", "Esperado: ")
        self.assertEqual(caracteres_dobles("a"), "aa", "Esperado: aa")
        self.assertEqual(caracteres_dobles("."), "..", "Esperado: ..")
        self.assertEqual(caracteres_dobles("aa"), "aaaa", "Esperado: aaaa")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-5">
<h2>Ejercicio 5<a class="headerlink" href="#ejercicio-5" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_5 data-question_label="5">
<div id=q10_5_question class="ac_question col-md-12">
<p>Desarrolle la función, <code class="docutils literal notranslate"><span class="pre">contar_saludo</span></code> que recibe como parámetro una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelve el número de veces que aparece la cadena <code class="docutils literal notranslate"><span class="pre">&quot;hola&quot;</span></code> en <code class="docutils literal notranslate"><span class="pre">s</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">contar_saludo(&quot;abc</span> <span class="pre">hola</span> <span class="pre">ho&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_saludo(&quot;ABChola</span> <span class="pre">hola&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_saludo(&quot;holahola&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_5_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_saludo(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(contar_saludo("abc hola ho"), 1, "Esperado: abc hola ho")
        self.assertEqual(contar_saludo("ABChola hola"), 2, "Esperado: ABChi hola")
        self.assertEqual(contar_saludo("holahola"), 2, "Esperado: holahola")
        self.assertEqual(contar_saludo("holaHoLAhoLa"), 1, "Esperado: holaHOLAhoLa")
        self.assertEqual(contar_saludo(""), 0, "Esperado: ")
        self.assertEqual(contar_saludo("h"), 0, "Esperado: h")
        self.assertEqual(contar_saludo("hola"), 1, "Esperado: hola")
        self.assertEqual(contar_saludo("Hola is no HOLA on aHolA"), 0, "Esperado: Hola is no HOLA on aHolA")
        self.assertEqual(contar_saludo("holaho not HOholaholA"), 2, "Esperado: holaho not HOholaholA")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-6">
<h2>Ejercicio 6<a class="headerlink" href="#ejercicio-6" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_6 data-question_label="6">
<div id=q10_6_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">gato_perro</span></code> que recibe como parámetro una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelve <code class="docutils literal notranslate"><span class="pre">True</span></code> si en la cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> aparece el mismo número de veces la cadena <code class="docutils literal notranslate"><span class="pre">&quot;gato&quot;</span></code> y la cadena <code class="docutils literal notranslate"><span class="pre">&quot;perro&quot;</span></code> o devuelve <code class="docutils literal notranslate"><span class="pre">False</span></code> en caso contrario. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">gato_perro(&quot;gatoperro&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">gato_perro(&quot;gatogato&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">gato_perro(&quot;1gato1cadoperro&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_6_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def gato_perro(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(gato_perro("gatoperro"), True, "Esperado: True")
        self.assertEqual(gato_perro("gatogato"), False, "Esperado: False")
        self.assertEqual(gato_perro("1gato1cadoperro"), True, "Esperado: True")
        self.assertEqual(gato_perro("gatoxxperroxxxperro"), False, "Esperado: False")
        self.assertEqual(gato_perro("gatoxperroxperroxgato"), True, "Esperado: True")
        self.assertEqual(gato_perro("gatoxperroxperroxca"), False, "Esperado: False")
        self.assertEqual(gato_perro("perroperrogato"), False, "Esperado: False")
        self.assertEqual(gato_perro("perrooggato"), True, "Esperado: True")
        self.assertEqual(gato_perro("perro"), False, "Esperado: False")
        self.assertEqual(gato_perro("gato"), False, "Esperado: False")
        self.assertEqual(gato_perro("ca"), True, "Esperado: True")
        self.assertEqual(gato_perro("c"), True, "Esperado: True")
        self.assertEqual(gato_perro(""), True, "Esperado: True")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-7">
<h2>Ejercicio 7<a class="headerlink" href="#ejercicio-7" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_7 data-question_label="7">
<div id=q10_7_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">contar_code</span></code> que recibe como parámetro una cadena <code class="docutils literal notranslate"><span class="pre">s</span></code> y devuelve el número de veces que aparece la cadena <code class="docutils literal notranslate"><span class="pre">&quot;code&quot;</span></code> en <code class="docutils literal notranslate"><span class="pre">s</span></code> pero con la condición de que la letra <code class="docutils literal notranslate"><span class="pre">'d'</span></code> puede ser intercambiada por cualquier otra. Así, las cadenas <code class="docutils literal notranslate"><span class="pre">&quot;coze&quot;</span></code> y <code class="docutils literal notranslate"><span class="pre">&quot;coxe&quot;</span></code> deben ser contadas igualmente. <br /> <br />
Ejemplos:<br />
<code class="docutils literal notranslate"><span class="pre">contar_code(&quot;aaacodebbb&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">1</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_code(&quot;codexxcode&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_code(&quot;cozexxcope&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_7_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_code(s):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(contar_code("aaacodebbb"), 1, "Esperado: 1")
        self.assertEqual(contar_code("codexxcode"), 2, "Esperado: 2")
        self.assertEqual(contar_code("cozexxcope"), 2, "Esperado: 2")
        self.assertEqual(contar_code("cozfxxcope"), 1, "Esperado: 1")
        self.assertEqual(contar_code("xxcozeyycop"), 1, "Esperado: 1")
        self.assertEqual(contar_code("cozcop"), 0, "Esperado: 0")
        self.assertEqual(contar_code("abcxyz"), 0, "Esperado: 0")
        self.assertEqual(contar_code("code"), 1, "Esperado: 1")
        self.assertEqual(contar_code("ode"), 0, "Esperado: 0")
        self.assertEqual(contar_code("c"), 0, "Esperado: 0")
        self.assertEqual(contar_code(""), 0, "Esperado: 0")
        self.assertEqual(contar_code("AAcodeBBcoleCCccoreDD"), 3, "Esperado: 3")
        self.assertEqual(contar_code("AAcodeBBcoleCCccorfDD"), 2, "Esperado: 2")
        self.assertEqual(contar_code("coAcodeBcoleccoreDD"), 3, "Esperado: 3")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-8">
<h2>Ejercicio 8<a class="headerlink" href="#ejercicio-8" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_8 data-question_label="8">
<div id=q10_8_question class="ac_question col-md-12">
<p>Sabiendo que la función <code class="docutils literal notranslate"><span class="pre">lower()</span></code> convierte los carácteres de una cadena en minúsculas. Desarrolle una función llamada <code class="docutils literal notranslate"><span class="pre">finales_iguales</span></code> que recibe dos cadenas <code class="docutils literal notranslate"><span class="pre">a</span></code> y <code class="docutils literal notranslate"><span class="pre">b</span></code> como parámetros. La función devolverá <code class="docutils literal notranslate"><span class="pre">True</span></code> si al final de la cadena <code class="docutils literal notranslate"><span class="pre">a</span></code> se encuentra la cadena <code class="docutils literal notranslate"><span class="pre">b</span></code> o si al final de la cadena <code class="docutils literal notranslate"><span class="pre">b</span></code> se encuentra la cadena <code class="docutils literal notranslate"><span class="pre">a</span></code>. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">finales_iguales(&quot;Hiabc&quot;,</span> <span class="pre">&quot;abc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">finales_iguales(&quot;AbC&quot;,</span> <span class="pre">&quot;HiaBc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">finales_iguales(&quot;abc&quot;,</span> <span class="pre">&quot;abXabc&quot;)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_8_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def finales_iguales(a, b):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(finales_iguales("Hiabc", "abc"), True, "Esperado: True")
        self.assertEqual(finales_iguales("AbC", "HiaBc"), True, "Esperado: True")
        self.assertEqual(finales_iguales("abc", "abXabc"), True, "Esperado: True")
        self.assertEqual(finales_iguales("Hiabc", "abcd"), False, "Esperado: False")
        self.assertEqual(finales_iguales("Hiabc", "bc"), True, "Esperado: True")
        self.assertEqual(finales_iguales("Hiabcx", "bc"), False, "Esperado: False")
        self.assertEqual(finales_iguales("abc", "abc"), True, "Esperado: True")
        self.assertEqual(finales_iguales("xyz", "12xyz"), True, "Esperado: True")
        self.assertEqual(finales_iguales("yz", "12xz"), False, "Esperado: False")
        self.assertEqual(finales_iguales("Z", "12xz"), True, "Esperado: True")
        self.assertEqual(finales_iguales("12", "12"), True, "Esperado: True")
        self.assertEqual(finales_iguales("abcXYZ", "abcDEF"), False, "Esperado: False")
        self.assertEqual(finales_iguales("ab", "ab12"), False, "Esperado: False")
        self.assertEqual(finales_iguales("ab", "12ab"), True, "Esperado: True")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-9">
<h2>Ejercicio 9<a class="headerlink" href="#ejercicio-9" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_9 data-question_label="9">
<div id=q10_9_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">contar_pares</span></code> que recibe una lista <code class="docutils literal notranslate"><span class="pre">numeros</span></code> de números enteros y devuelve la cantidad números pares que se encuentran en la lista. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">contar_pares([2,</span> <span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_pares([2,</span> <span class="pre">2,</span> <span class="pre">0])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">contar_pares([1,</span> <span class="pre">3,</span> <span class="pre">5])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_9_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def contar_pares(numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(contar_pares([2, 1, 2, 3, 4]), 3, "Esperado: 3")
        self.assertEqual(contar_pares([2, 2, 0]), 3, "Esperado: 3")
        self.assertEqual(contar_pares([1, 3, 5]), 0, "Esperado: 0")
        self.assertEqual(contar_pares([]), 0, "Esperado: 0")
        self.assertEqual(contar_pares([11, 9, 0, 1]), 1, "Esperado: 1")
        self.assertEqual(contar_pares([2, 11, 9, 0]), 2, "Esperado: 2")
        self.assertEqual(contar_pares([2]), 1, "Esperado: 1")
        self.assertEqual(contar_pares([2, 5, 12]), 2, "Esperado: 2")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-10">
<h2>Ejercicio 10<a class="headerlink" href="#ejercicio-10" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_10 data-question_label="10">
<div id=q10_10_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">suma_con_mas_suerte</span></code> que recibe una lista <code class="docutils literal notranslate"><span class="pre">numeros</span></code> de números enteros y devuelve la suma de los números de la lista con la condición de que si uno de los números es el <strong>13</strong>, éste no cuenta en la suma así como todos los números que se encuentren a su derecha. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_mas_suerte([1,</span> <span class="pre">2,</span> <span class="pre">2,</span> <span class="pre">1])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">6</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_mas_suerte([1,</span> <span class="pre">1])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">2</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_mas_suerte([1,</span> <span class="pre">2,</span> <span class="pre">13,</span> <span class="pre">1,</span> <span class="pre">13])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">3</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_con_mas_suerte([13,</span> <span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">0</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_10_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def suma_con_mas_suerte(numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_con_mas_suerte([1, 2, 2, 1]), 6, "Esperado: 6")
        self.assertEqual(suma_con_mas_suerte([1, 1]), 2, "Esperado: 2")
        self.assertEqual(suma_con_mas_suerte([1, 2, 2, 1, 13]), 6, "Esperado: 6")
        self.assertEqual(suma_con_mas_suerte([1, 2, 13, 2, 1, 13]), 3, "Esperado: 3")
        self.assertEqual(suma_con_mas_suerte([13, 1, 2, 13, 2, 1, 13]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([0, 1, 2, 13, 2, 1, 13]), 3, "Esperado: 3")
        self.assertEqual(suma_con_mas_suerte([]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([13]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([0, 5, 1, 2, 13, 100]), 8, "Esperado: 8")
        self.assertEqual(suma_con_mas_suerte([13, 13]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([13, 0, 13]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([13, 1, 13]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([5, 7, 2]), 14, "Esperado: 14")
        self.assertEqual(suma_con_mas_suerte([5, 13, 2]), 5, "Esperado: 5")
        self.assertEqual(suma_con_mas_suerte([0]), 0, "Esperado: 0")
        self.assertEqual(suma_con_mas_suerte([13, 0]), 0, "Esperado: 0")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-11">
<h2>Ejercicio 11<a class="headerlink" href="#ejercicio-11" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_11 data-question_label="11">
<div id=q10_11_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">tiene_2</span></code> que recibe como parámetro una lista <code class="docutils literal notranslate"><span class="pre">numeros</span></code> de números enteros y devuelve <code class="docutils literal notranslate"><span class="pre">True</span></code> si en la lista aparecen dos números iguales consecutivos y <code class="docutils literal notranslate"><span class="pre">False</span></code> en caso contrario. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([1,</span> <span class="pre">2,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([1,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">tiene_2([2,</span> <span class="pre">1,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_11_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def tiene_2(numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(tiene_2([1, 2, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([1, 2, 1, 2]), False, "Esperado: False")
        self.assertEqual(tiene_2([2, 1, 2]), False, "Esperado: False")
        self.assertEqual(tiene_2([2, 2, 1, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([1, 3, 2]), False, "Esperado: False")
        self.assertEqual(tiene_2([1, 3, 2, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([2, 3, 2, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([4, 2, 4, 2, 2, 5]), True, "Esperado: True")
        self.assertEqual(tiene_2([1, 2]), False, "Esperado: False")
        self.assertEqual(tiene_2([2, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([2]), False, "Esperado: False")
        self.assertEqual(tiene_2([]), False, "Esperado: False")
        self.assertEqual(tiene_2([3, 3, 2, 2]), True, "Esperado: True")
        self.assertEqual(tiene_2([5, 2, 5, 2]), False, "Esperado: False")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-12">
<h2>Ejercicio 12<a class="headerlink" href="#ejercicio-12" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_12 data-question_label="12">
<div id=q10_12_question class="ac_question col-md-12">
<p>Desarrolle la función <code class="docutils literal notranslate"><span class="pre">suma_en_lista</span></code> que recibe dos parámetros, un entero <code class="docutils literal notranslate"><span class="pre">n</span></code> y un lista <code class="docutils literal notranslate"><span class="pre">numeros</span></code> de números. La función debe devolver <code class="docutils literal notranslate"><span class="pre">True</span></code> si existen dos elementos distintos de la lista que su suma da como resultado <code class="docutils literal notranslate"><span class="pre">n</span></code> y devolver <code class="docutils literal notranslate"><span class="pre">False</span></code> en caso contrario. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(5,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(9,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(0,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(8,</span> <span class="pre">[1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(4,</span> <span class="pre">[2,</span> <span class="pre">2,</span> <span class="pre">2,</span> <span class="pre">2])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<code class="docutils literal notranslate"><span class="pre">suma_en_lista(4,</span> <span class="pre">[2,</span> <span class="pre">2,</span> <span class="pre">1,</span> <span class="pre">3])</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br /></p>

</div>
<textarea data-lang="python" id="q10_12_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def suma_en_lista(n, numeros):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(suma_en_lista(5, [1, 2, 3, 4]), True, "Esperado: True")
        self.assertEqual(suma_en_lista(9, [1, 2, 3, 4]), False, "Esperado: False")
        self.assertEqual(suma_en_lista(0, [1, 2, 3, 4]), False, "Esperado: False")
        self.assertEqual(suma_en_lista(8, [1, 2, 3, 4]), False, "Esperado: False")
        self.assertEqual(suma_en_lista(4, [2, 2, 2, 2]), False, "Esperado: False")
        self.assertEqual(suma_en_lista(4, [2, 2, 1, 3]), True, "Esperado: True")
        self.assertEqual(suma_en_lista(42, [40, 2, 3, 39]), True, "Esperado: True")
        self.assertEqual(suma_en_lista(10, [5, 5, 4, 6]), True, "Esperado: True")
        self.assertEqual(suma_en_lista(8, [5, 5, 4, 4]), False, "Esperado: False")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="ejercicio-13">
<h2>Ejercicio 13<a class="headerlink" href="#ejercicio-13" title="Permalink to this headline">¶</a></h2>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=q10_13 data-question_label="13">
<div id=q10_13_question class="ac_question col-md-12">
<p>En una construcción se desea construir una fila de ladrillos. Se tienen ladrillos pequeños y ladrillos grandes. Los ladrillos pequeños son de 1 unidad de longitud mientras que los ladrillos grandes son de 5 unidades de longitud. Desarrolle la función <code class="docutils literal notranslate"><span class="pre">construir_fila</span></code> que recibe tres parámetros, <code class="docutils literal notranslate"><span class="pre">n_lad_peq</span></code>, <code class="docutils literal notranslate"><span class="pre">n_lad_grand</span></code> y <code class="docutils literal notranslate"><span class="pre">longitud_fila</span></code>. Estos parámetros son números enteros que representan la cantidad de ladrillos pequeños, cantidad de ladrillos grandes y la longitud de la fila que se quiere construir, respectivamente. La función debe devolver <code class="docutils literal notranslate"><span class="pre">True</span></code> si es posible armar la fila, o <code class="docutils literal notranslate"><span class="pre">False</span></code> en caso contrario. <br /> <br />
Ejemplos: <br />
<code class="docutils literal notranslate"><span class="pre">construir_fila(3,</span> <span class="pre">1,</span> <span class="pre">8)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<em>Explicación</em>: Para construir una longitud de 8 metros, se usarían 1 ladrillo grande y 3 ladrillos pequeños <br />
<em>Por lo tanto</em>: 1 + 1 + 1 + 5 = 8 <br /></p>
<p><code class="docutils literal notranslate"><span class="pre">construir_fila(3,</span> <span class="pre">1,</span> <span class="pre">9)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">False</span></code> <br />
<em>Explicación</em>: Para construir una longitud de 9 metros no se podría porque solamente se tienen 3 ladrillos pequeños (de 1 metro de longitud) y 1 ladrillo grande (de 5 metros de longitud). <br />
<em>Por lo tanto</em>: 1 + 1 + 1 + 5 &lt; 9 en este caso se necesitaría 1 ladrillo grande y 4 pequeños. <br /></p>
<p><code class="docutils literal notranslate"><span class="pre">construir_fila(3,</span> <span class="pre">2,</span> <span class="pre">10)</span></code> -&gt; <code class="docutils literal notranslate"><span class="pre">True</span></code> <br />
<em>Explicación</em>: Para llenar una longitud de 10 metros, se usarían solamente los 2 ladrillos grandes disponibles. <br />
<em>Por lo tanto</em>: 5 + 5 = 10 <br /></p>

</div>
<textarea data-lang="python" id="q10_13_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
      style="visibility: hidden;">
def construir_fila(n_lad_peq, n_lad_grand, longitud_fila):


====
from unittest.gui import TestCaseGui


class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(construir_fila(3, 1, 8), True, "Esperado: True")
        self.assertEqual(construir_fila(3, 1, 9), False, "Esperado: False")
        self.assertEqual(construir_fila(3, 2, 10), True, "Esperado: True")
        self.assertEqual(construir_fila(3, 2, 8), True, "Esperado: True")
        self.assertEqual(construir_fila(3, 2, 9), False, "Esperado: False")
        self.assertEqual(construir_fila(6, 1, 11), True, "Esperado: True")
        self.assertEqual(construir_fila(6, 0, 11), False, "Esperado: False")
        self.assertEqual(construir_fila(3, 1, 7), True, "Esperado: True")
        self.assertEqual(construir_fila(1, 1, 7), False, "Esperado: False")
        self.assertEqual(construir_fila(2, 1, 7), True, "Esperado: True")
        self.assertEqual(construir_fila(7, 1, 11), True, "Esperado: True")
        self.assertEqual(construir_fila(7, 1, 8), True, "Esperado: True")
        self.assertEqual(construir_fila(7, 1, 13), False, "Esperado: False")
        self.assertEqual(construir_fila(43, 1, 46), True, "Esperado: True")
        self.assertEqual(construir_fila(40, 1, 46), False, "Esperado: False")
        self.assertEqual(construir_fila(22, 2, 33), False, "Esperado: False")
        self.assertEqual(construir_fila(0, 2, 10), True, "Esperado: True")
        self.assertEqual(construir_fila(1000000, 1000, 1000100), True, "Esperado: True")
        self.assertEqual(construir_fila(2, 1000000, 100003), False, "Esperado: False")
        self.assertEqual(construir_fila(12, 2, 21), True, "Esperado: True")


myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Quiz - 9' data-toggle="tooltip">
    <a href="Quiz9.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Quiz9.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Quiz - 11' data-toggle="tooltip" >
    <a href="Quiz11.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="Quiz11.html" >Next Section - Quiz - 11</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2021 Python Argentina et al..
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 5.7.2.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>